<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>2. Python Review &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../_static/runestone.e4d5592da655219f.css?v=32648AE5" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/&#34;&#34;/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script src="../_static/runestonebase.js"></script>
    <script data-url_root="../" id="documentation_options" src="../_static/documentation_options.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script defer="" src="../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../_static/runtime.c894f7032a5bccea.bundle.js?v=D20A5287"></script>
    <script defer="" src="../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../_static/runestone.5b7c431cb6d3e0db.bundle.js?v=36CB01A8"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Underlying Technology of Computing Devices and the Internet" href="../ComputingAndSociety/UnderlyingTechnology.html" />
    <link rel="prev" title="1. Reeborg Review" href="Reeborg.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '6.2.2';
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">ActiveCode Work Area</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">2. Python Review</a><ul>
<li><a class="reference internal" href="#variables-and-data-types">2.1. Variables and Data Types</a></li>
<li><a class="reference internal" href="#converting-between-data-types">2.2. Converting Between Data Types</a></li>
<li><a class="reference internal" href="#math-operators">2.3. Math Operators</a></li>
<li><a class="reference internal" href="#if">2.4. <code class="docutils literal notranslate"><span class="pre">if</span></code></a></li>
<li><a class="reference internal" href="#if-elif-else">2.5. <code class="docutils literal notranslate"><span class="pre">if/elif/else</span></code></a></li>
<li><a class="reference internal" href="#while-loop">2.6. <code class="docutils literal notranslate"><span class="pre">while</span></code> loop</a></li>
<li><a class="reference internal" href="#for-loop">2.7. <code class="docutils literal notranslate"><span class="pre">for</span></code> loop</a></li>
<li><a class="reference internal" href="#and-or-not">2.8. <code class="docutils literal notranslate"><span class="pre">and</span></code>, <code class="docutils literal notranslate"><span class="pre">or</span></code>, <code class="docutils literal notranslate"><span class="pre">not</span></code></a></li>
<li><a class="reference internal" href="#functions">2.9. Functions</a></li>
<li><a class="reference internal" href="#drawing-images-practice">2.10. Drawing Images Practice</a></li>
<li><a class="reference internal" href="#image-manipulation-with-nested-loops">2.11. Image Manipulation with Nested Loops</a></li>
<li><a class="reference internal" href="#strings">2.12. Strings</a></li>
<li><a class="reference internal" href="#lists">2.13. Lists</a></li>
<li><a class="reference internal" href="#number-guessing-game">2.14. Number Guessing Game</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">1. </span>Reeborg Review" data-toggle="tooltip">
    <a href="Reeborg.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Underlying Technology of Computing Devices and the Internet' data-toggle="tooltip" >
    <a href="../ComputingAndSociety/UnderlyingTechnology.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <section id="python-review">
<h1><span class="section-number">2. </span>Python Review<a class="headerlink" href="#python-review" title="Permalink to this headline">¶</a></h1>
<section id="variables-and-data-types">
<h2><span class="section-number">2.1. </span>Variables and Data Types<a class="headerlink" href="#variables-and-data-types" title="Permalink to this headline">¶</a></h2>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.1.1" data-multipleanswers="false"  id="review_data_types_1_1"  style="visibility: hidden;">
    <p>python-review1: What is the data type of <code class="docutils literal notranslate"><span class="pre">'what</span> <span class="pre">is</span> <span class="pre">your</span> <span class="pre">name?'</span></code>?</p>

    <li data-component="answer"  id="review_data_types_1_1_opt_a">boolean</li><li data-component="feedback">It is not True or False.</li>
    
    <li data-component="answer"  id="review_data_types_1_1_opt_b">integer</li><li data-component="feedback">The data is not numeric.</li>
    
    <li data-component="answer"  id="review_data_types_1_1_opt_c">float</li><li data-component="feedback">The value is not numeric with a decimal point.</li>
    
    <li data-component="answer" data-correct='yes' id="review_data_types_1_1_opt_d">string</li><li data-component="feedback">Great! Strings are enclosed in quotes.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.1.2" data-multipleanswers="false"  id="review_data_types_1_2"  style="visibility: hidden;">
    <p>python-review2: What is the data type of <code class="docutils literal notranslate"><span class="pre">42</span></code>?</p>

    <li data-component="answer"  id="review_data_types_1_2_opt_a">boolean</li><li data-component="feedback">It is not True or False.</li>
    
    <li data-component="answer" data-correct='yes' id="review_data_types_1_2_opt_b">integer</li><li data-component="feedback">Great! The data is numeric, without a decimal point.</li>
    
    <li data-component="answer"  id="review_data_types_1_2_opt_c">float</li><li data-component="feedback">The value is not numeric with a decimal point.</li>
    
    <li data-component="answer"  id="review_data_types_1_2_opt_d">string</li><li data-component="feedback">Strings are enclosed in quotes.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.1.3" data-multipleanswers="false"  id="review_data_types_1_3"  style="visibility: hidden;">
    <p>python-review3: What is the data type of <code class="docutils literal notranslate"><span class="pre">False</span></code>?</p>

    <li data-component="answer" data-correct='yes' id="review_data_types_1_3_opt_a">boolean</li><li data-component="feedback">Great! Boolean is either True or False.</li>
    
    <li data-component="answer"  id="review_data_types_1_3_opt_b">integer</li><li data-component="feedback">The data is not numeric.</li>
    
    <li data-component="answer"  id="review_data_types_1_3_opt_c">float</li><li data-component="feedback">The value is not numeric with a decimal point.</li>
    
    <li data-component="answer"  id="review_data_types_1_3_opt_d">string</li><li data-component="feedback">Strings are enclosed in quotes.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.1.4" data-multipleanswers="false"  id="review_data_types_1_4"  style="visibility: hidden;">
    <p>python-review4: What is the data type of <code class="docutils literal notranslate"><span class="pre">4.5</span></code>?</p>

    <li data-component="answer"  id="review_data_types_1_4_opt_a">boolean</li><li data-component="feedback">It is not True or False.</li>
    
    <li data-component="answer"  id="review_data_types_1_4_opt_b">integer</li><li data-component="feedback">The data is not numeric.</li>
    
    <li data-component="answer" data-correct='yes' id="review_data_types_1_4_opt_c">float</li><li data-component="feedback">Great! The value is numeric with a decimal point.</li>
    
    <li data-component="answer"  id="review_data_types_1_4_opt_d">string</li><li data-component="feedback">Strings are enclosed in quotes.</li>
    

    </ul>
    </div>
    </section>
<section id="converting-between-data-types">
<h2><span class="section-number">2.2. </span>Converting Between Data Types<a class="headerlink" href="#converting-between-data-types" title="Permalink to this headline">¶</a></h2>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_type_casting_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a block of code that correctly takes in a number from a user, then prints out twice the value that was given.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.2.1"          style="visibility: hidden;">
        number =  input("Please enter a number:")
number = float(number)
doubled = number * 2
print(str(doubled) + " is twice the value of the number you entered.")
        </pre>
        </div>
        </div>
    <p>The following example is not going to work when you try to run it. Try entering in <code class="docutils literal notranslate"><span class="pre">16</span></code>, then <code class="docutils literal notranslate"><span class="pre">15</span></code>. Notice that nothing is printed either time, even though it looks like the conditional should be causing it to print. Can you figure out what is wrong and fix it? <em>Hint: think about data types!</em></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=review_input_intro_3 data-question_label="2.2.2">
<div id=review_input_intro_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="review_input_intro_3_editor" 
      data-timelimit=25000     
    data-audio=''      
         data-caption='Can you figure out what is wrong?'  data-wasm=/_static
     style="visibility: hidden;">
age = input("How old are you?")

if age == 16:
    print("You can get your driver's license!")
elif age == 15:
    print("You can get your learner driver's license.")
</textarea>
</div>
</div>
</section>
<section id="math-operators">
<h2><span class="section-number">2.3. </span>Math Operators<a class="headerlink" href="#math-operators" title="Permalink to this headline">¶</a></h2>
<p>If you have trouble with any of the following, you might want to look back at the <a class="reference internal" href="../Python/Overview/QuickOverview.html#math-operator-list"><span class="std std-ref">Math Operators</span></a> list.</p>

        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.1" id="review_python_math_operators_1"  style="visibility: hidden;">
            <p>python-review7: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">12</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span> <span class="o">/</span> <span class="mi">4</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [3, 3], "feedback": "<p>Great!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.2" id="review_python_math_operators_2"  style="visibility: hidden;">
            <p>python-review8: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">12</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span> <span class="o">%</span> <span class="mi">5</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [2, 2], "feedback": "<p>Great!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.3" id="review_python_math_operators_3"  style="visibility: hidden;">
            <p>python-review9: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">42</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span> <span class="o">//</span> <span class="mi">5</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [8, 8], "feedback": "<p>Great!</p>\n"}, {"number": [8.4, 8.4], "feedback": "<p>Remember that // returns only an integer!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.4" id="review_python_math_operators_4"  style="visibility: hidden;">
            <p>python-review10: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">2</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span> <span class="o">**</span> <span class="mi">4</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [16, 16], "feedback": "<p>Great!</p>\n"}, {"number": [8, 8], "feedback": "<p>Remember that ** means \u201cto the power of\u201d.</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again!</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.5" id="review_python_math_operators_5"  style="visibility: hidden;">
            <p>python-review11: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">2</span>
<span class="n">number</span> <span class="o">+=</span> <span class="mi">4</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [6, 6], "feedback": "<p>Great!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again! Remember that number += 4 is the same as number = number + 4</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="runestone ">
        <div data-component="fillintheblank" data-question_label="2.3.6" id="review_python_math_operators_6"  style="visibility: hidden;">
            <p>python-review12: What would the following code print?:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">9</span>
<span class="n">number</span> <span class="o">-=</span> <span class="mi">2</span>
<span class="nb">print</span><span class="p">(</span><span class="n">number</span><span class="p">)</span>
</pre></div>
</div>
<input type="text">
        <script type="application/json">
            [[{"number": [7, 7], "feedback": "<p>Great!</p>\n"}, {"regex": "^\\s*.*\\s*$", "regexFlags": "", "feedback": "<p>Try again! Remember that number -= 2 is the same as number = number - 2</p>\n"}]]
        </script>

        </div>
        </div>
            </section>
<section id="if">
<h2><span class="section-number">2.4. </span><code class="docutils literal notranslate"><span class="pre">if</span></code><a class="headerlink" href="#if" title="Permalink to this headline">¶</a></h2>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>A single equal sign <code class="docutils literal notranslate"><span class="pre">=</span></code> is used to <strong>assign</strong> a value. Two equal signs <code class="docutils literal notranslate"><span class="pre">==</span></code> are used when <strong>comparing</strong> a value.</p>
</div>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.4.1" data-multipleanswers="false"  id="review_conditionals_if_mc_1"  style="visibility: hidden;">
    <p>python-review13: What does the following code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">x</span> <span class="o">=</span> <span class="o">-</span><span class="mi">10</span>
<span class="k">if</span> <span class="n">x</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;The negative number &quot;</span><span class="p">,</span>  <span class="n">x</span><span class="p">,</span> <span class="s2">&quot; is not valid here.&quot;</span><span class="p">)</span>
<span class="nb">print</span><span class="p">(</span><span class="s2">&quot;This is always printed&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-html notranslate"><div class="highlight"><pre><span></span>a.
This is always printed

b.
The negative number -10 is not valid here
This is always printed

c.
The negative number -10 is not valid here
</pre></div>
</div>

    <li data-component="answer"  id="review_conditionals_if_mc_1_opt_a">Output a</li><li data-component="feedback">Because -10 is less than 0, Python will execute the body of the if-statement here.</li>
    
    <li data-component="answer" data-correct='yes' id="review_conditionals_if_mc_1_opt_b">Output b</li><li data-component="feedback">Python executes the body of the if-block as well as the statement that follows the if-block.</li>
    
    <li data-component="answer"  id="review_conditionals_if_mc_1_opt_c">Output c</li><li data-component="feedback">Python will also execute the statement that follows the if-block (because it is not enclosed in an else-block, but rather just a normal statement).</li>
    
    <li data-component="answer"  id="review_conditionals_if_mc_1_opt_d">It will cause an error because every if must have an else clause.</li><li data-component="feedback">It is valid to have an if-block without a corresponding else-block (though you cannot have an else-block without a corresponding if-block).</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.4.2" data-multipleanswers="false"  id="review_conditionals_if_mc_2"  style="visibility: hidden;">
    <p>python-review14: What does the following code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">64</span>

<span class="k">if</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s positive!&quot;</span><span class="p">)</span>

<span class="k">if</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">50</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s pretty big.&quot;</span><span class="p">)</span>

<span class="k">if</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">100</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s really big!&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-html notranslate"><div class="highlight"><pre><span></span>a.
It&#39;s positive!

b.
It&#39;s positive!
It&#39;s really big!

c.
It&#39;s pretty big.

d.
It&#39;s positive!
It&#39;s pretty big.
</pre></div>
</div>

    <li data-component="answer"  id="review_conditionals_if_mc_2_opt_a">Output a</li><li data-component="feedback">Try again. Remember that any number of consecutive if statements can evaluate to True.</li>
    
    <li data-component="answer"  id="review_conditionals_if_mc_2_opt_b">Output b</li><li data-component="feedback">Try again. Remember that any number of consecutive if statements can evaluate to True.</li>
    
    <li data-component="answer"  id="review_conditionals_if_mc_2_opt_c">Output c</li><li data-component="feedback">Try again. It's less than 100!</li>
    
    <li data-component="answer" data-correct='yes' id="review_conditionals_if_mc_2_opt_d">Output d</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_conditionals_if_mc_2_opt_e">It will cause an error because every if must have an else clause.</li><li data-component="feedback">Try again. An if statement does not require an else!</li>
    

    </ul>
    </div>
    </section>
<section id="if-elif-else">
<h2><span class="section-number">2.5. </span><code class="docutils literal notranslate"><span class="pre">if/elif/else</span></code><a class="headerlink" href="#if-elif-else" title="Permalink to this headline">¶</a></h2>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.5.1" data-multipleanswers="false"  id="review_conditionals_mc_1"  style="visibility: hidden;">
    <p>python-review15: What does the following code print (choose from output a, b, c or nothing)?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span> <span class="mi">4</span> <span class="o">+</span> <span class="mi">5</span> <span class="o">==</span> <span class="mi">10</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;TRUE&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;FALSE&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span><span class="o">.</span>  <span class="n">TRUE</span>

<span class="n">b</span><span class="o">.</span>  <span class="n">FALSE</span>

<span class="n">c</span><span class="o">.</span>  <span class="n">TRUE</span>
    <span class="n">FALSE</span>

<span class="n">d</span><span class="o">.</span>  <span class="n">Nothing</span> <span class="n">will</span> <span class="n">be</span> <span class="n">printed</span><span class="o">.</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_conditionals_mc_1_opt_a">Output 1</li><li data-component="feedback">TRUE is printed by the if-block, which only executes if the conditional (in this case, 4+5 == 10) is true.  In this case 5+4 is not equal to 10.</li>
    
    <li data-component="answer" data-correct='yes' id="review_conditionals_mc_1_opt_b">Output 2</li><li data-component="feedback">Since 4+5==10 evaluates to False, Python will skip over the if block and execute the statement in the else block.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_1_opt_c">Output 3</li><li data-component="feedback">Python would never print both TRUE and FALSE because it will only execute one of the if-block or the else-block, but not both.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_1_opt_d">Output 4</li><li data-component="feedback">Python will always execute either the if-block (if the condition is true) or the else-block (if the condition is false).  It would never skip over both blocks.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.5.2" data-multipleanswers="false"  id="review_conditionals_mc_2"  style="visibility: hidden;">
    <p>python-review16: What does the following code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span> <span class="mi">4</span> <span class="o">+</span> <span class="mi">5</span> <span class="o">==</span> <span class="mi">10</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;TRUE&quot;</span><span class="p">)</span>
<span class="k">else</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;FALSE&quot;</span><span class="p">)</span>
<span class="nb">print</span><span class="p">(</span><span class="s2">&quot;TRUE&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span><span class="o">.</span> <span class="n">TRUE</span>

<span class="n">b</span><span class="o">.</span>
   <span class="n">TRUE</span>
   <span class="n">FALSE</span>

<span class="n">c</span><span class="o">.</span>
   <span class="n">FALSE</span>
   <span class="n">TRUE</span>
<span class="n">d</span><span class="o">.</span>
   <span class="n">TRUE</span>
   <span class="n">FALSE</span>
   <span class="n">TRUE</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_conditionals_mc_2_opt_a">Output a</li><li data-component="feedback">Although TRUE is printed after the if-else statement completes, both blocks within the if-else statement print something too.  In this case, Python would have had to have skipped both blocks in the if-else statement, which it never would do.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_2_opt_b">Output b</li><li data-component="feedback">Because there is a TRUE printed after the if-else statement ends, Python will always print TRUE as the last statement.</li>
    
    <li data-component="answer" data-correct='yes' id="review_conditionals_mc_2_opt_c">Output c</li><li data-component="feedback">Python will print FALSE from within the else-block (because 5+4 does not equal 10), and then print TRUE after the if-else statement completes.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_2_opt_d">Output d</li><li data-component="feedback">To print these three lines, Python would have to execute both blocks in the if-else statement, which it can never do.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.5.3" data-multipleanswers="false"  id="review_conditionals_mc_3"  style="visibility: hidden;">
    <p>python-review17: What does the following code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">number</span> <span class="o">=</span> <span class="mi">64</span>

<span class="k">if</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s positive!&quot;</span><span class="p">)</span>

<span class="k">elif</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">50</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s pretty big.&quot;</span><span class="p">)</span>

<span class="k">elif</span> <span class="n">number</span> <span class="o">&gt;</span> <span class="mi">100</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;It&#39;s really big!&quot;</span><span class="p">)</span>
</pre></div>
</div>
<div class="highlight-html notranslate"><div class="highlight"><pre><span></span>a.
It&#39;s positive!

b.
It&#39;s positive!
It&#39;s really big!

c.
It&#39;s pretty big.

d.
It&#39;s positive!
It&#39;s pretty big.
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_conditionals_mc_3_opt_a">Output a</li><li data-component="feedback">Great! Remember that only one branch of an if/elif/else block can execute.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_3_opt_b">Output b</li><li data-component="feedback">Try again. Remember that only one branch of an if/elif/else block can execute.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_3_opt_c">Output c</li><li data-component="feedback">Try again. Remember that only one branch of an if/elif/else block can execute.</li>
    
    <li data-component="answer"  id="review_conditionals_mc_3_opt_d">Output d</li><li data-component="feedback">Try again. Remember that only one branch of an if/elif/else block can execute.</li>
    

    </ul>
    </div>
    </section>
<section id="while-loop">
<h2><span class="section-number">2.6. </span><code class="docutils literal notranslate"><span class="pre">while</span></code> loop<a class="headerlink" href="#while-loop" title="Permalink to this headline">¶</a></h2>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_while_loop_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a block of code that correctly counts from 10 down to 1, then prints Blastoff!</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.6.1"          style="visibility: hidden;">
        counter = 10
while counter > 0:
    print(counter)
    counter = counter - 1
print("Blastoff!")
        </pre>
        </div>
        </div>
    <p>Write a program that asks the user to enter a password. Keep asking for the password until they enter “sask”.  Once they have successfully typed in “sask”, print out <em>What a great place!</em>.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=review_while_intro_6 data-question_label="2.6.2">
<div id=review_while_intro_6_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="review_while_intro_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
#put your code here!
</textarea>
</div>
</div>
</section>
<section id="for-loop">
<h2><span class="section-number">2.7. </span><code class="docutils literal notranslate"><span class="pre">for</span></code> loop<a class="headerlink" href="#for-loop" title="Permalink to this headline">¶</a></h2>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_for_loop_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a block of code that correctly counts from 1 up to 10, then prints <code class="docutils literal notranslate"><span class="pre">Made</span> <span class="pre">it!</span></code></p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.7.1"          style="visibility: hidden;">
        for counter in range(1, 11):
    print(counter)
print("Made it!")
        </pre>
        </div>
        </div>
    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_for_loop_2" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a block of code that correctly prints out a greeting for each person, using a for loop.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.7.2"          style="visibility: hidden;">
        people = ["Alice", "Bob", "Eve"]
for name in people:
    print("Hello there, " + name)
        </pre>
        </div>
        </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.7.3" data-multipleanswers="false"  id="review_test_question3_4_1"  style="visibility: hidden;">
    <p>python-review22: In the following code, how many lines does this code print?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">for</span> <span class="n">number</span> <span class="ow">in</span> <span class="p">[</span><span class="mi">5</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">0</span><span class="p">]:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;I have&quot;</span><span class="p">,</span> <span class="n">number</span><span class="p">,</span> <span class="s2">&quot;cookies.  I&#39;m going to eat one.&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_test_question3_4_1_opt_a">1</li><li data-component="feedback">The loop body prints one line, but the body will execute exactly one time for each element in the list [5, 4, 3, 2, 1, 0].</li>
    
    <li data-component="answer"  id="review_test_question3_4_1_opt_b">5</li><li data-component="feedback">Although the biggest number in the list is 5, there are actually 6 elements in the list.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question3_4_1_opt_c">6</li><li data-component="feedback">The loop body will execute (and print one line) for each of the 6 elements in the list [5, 4, 3, 2, 1, 0].</li>
    
    <li data-component="answer"  id="review_test_question3_4_1_opt_d">10</li><li data-component="feedback">The loop body will not execute more times than the number of elements in the list.</li>
    

    </ul>
    </div>
    </section>
<section id="and-or-not">
<h2><span class="section-number">2.8. </span><code class="docutils literal notranslate"><span class="pre">and</span></code>, <code class="docutils literal notranslate"><span class="pre">or</span></code>, <code class="docutils literal notranslate"><span class="pre">not</span></code><a class="headerlink" href="#and-or-not" title="Permalink to this headline">¶</a></h2>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.1" data-multipleanswers="false"  id="review_booleans_practice_quiz_1"  style="visibility: hidden;">
    <p>python-review23: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a</span> <span class="o">==</span> <span class="mi">6</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_1_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_1_opt_b">False</li><li data-component="feedback">Try again!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.2" data-multipleanswers="false"  id="review_booleans_practice_quiz_2"  style="visibility: hidden;">
    <p>python-review24: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="ow">not</span> <span class="p">(</span><span class="n">b</span> <span class="o">==</span> <span class="mi">6</span><span class="p">)</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_2_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_2_opt_b">False</li><li data-component="feedback">Try again!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.3" data-multipleanswers="false"  id="review_booleans_practice_quiz_3"  style="visibility: hidden;">
    <p>python-review25: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">10</span> <span class="ow">or</span> <span class="n">b</span> <span class="o">==</span> <span class="mi">10</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_3_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_3_opt_b">False</li><li data-component="feedback">Try again!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.4" data-multipleanswers="false"  id="review_booleans_practice_quiz_4"  style="visibility: hidden;">
    <p>python-review26: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">6</span> <span class="ow">and</span> <span class="mi">10</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_4_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_4_opt_b">False</li><li data-component="feedback">Try again! Notice that we did not ask a full question on the right side of the AND.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.5" data-multipleanswers="false"  id="review_booleans_practice_quiz_5"  style="visibility: hidden;">
    <p>python-review27: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="ow">not</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">10</span> <span class="ow">and</span> <span class="n">b</span> <span class="o">==</span> <span class="mi">10</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_5_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_5_opt_b">False</li><li data-component="feedback">Try again!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.6" data-multipleanswers="false"  id="review_booleans_practice_quiz_6"  style="visibility: hidden;">
    <p>python-review28: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">10</span> <span class="ow">or</span> <span class="ow">not</span> <span class="n">b</span> <span class="o">==</span> <span class="mi">10</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_booleans_practice_quiz_6_opt_a">True</li><li data-component="feedback">Try again!</li>
    
    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_6_opt_b">False</li><li data-component="feedback">Great!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.7" data-multipleanswers="false"  id="review_booleans_practice_quiz_7"  style="visibility: hidden;">
    <p>python-review29: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">6</span> <span class="ow">and</span> <span class="p">(</span><span class="ow">not</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">10</span><span class="p">)</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_7_opt_a">True</li><li data-component="feedback">Great!</li>
    
    <li data-component="answer"  id="review_booleans_practice_quiz_7_opt_b">False</li><li data-component="feedback">Try again!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.8.8" data-multipleanswers="false"  id="review_booleans_practice_quiz_8"  style="visibility: hidden;">
    <p>python-review30: What would the following print?:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="mi">6</span>
<span class="n">b</span> <span class="o">=</span> <span class="mi">10</span>
<span class="nb">print</span><span class="p">(</span> <span class="ow">not</span> <span class="p">(</span> <span class="ow">not</span> <span class="n">a</span> <span class="o">==</span> <span class="mi">10</span> <span class="ow">or</span> <span class="ow">not</span> <span class="n">b</span> <span class="o">==</span> <span class="mi">10</span><span class="p">)</span> <span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_booleans_practice_quiz_8_opt_a">True</li><li data-component="feedback">Try again!</li>
    
    <li data-component="answer" data-correct='yes' id="review_booleans_practice_quiz_8_opt_b">False</li><li data-component="feedback">Great!</li>
    

    </ul>
    </div>
    </section>
<section id="functions">
<h2><span class="section-number">2.9. </span>Functions<a class="headerlink" href="#functions" title="Permalink to this headline">¶</a></h2>

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_functions_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a function which returns the type of clothing you should wear, based on the parameter temperature. If the temperature is below -10, you will wear a parka and toque. If the temperature is between -10 and 0, wear a toque. If the temperature is greater than 0 but less than 10, wear a sweater. If the temperature is between 10 and 20, wear a t-shirt. If the temperature is greater than 20, wear shorts.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.9.1"          style="visibility: hidden;">
        def wear_the_right_thing(temperature):
    """Return a suggestion of what to wear, given a temperature."""
    if temperature < -10:
        return "parka and toque"
    elif temperature <= 0:
        return "toque"
    elif temperature < 10:
        return "sweater"
    elif temperature <= 20:
        return "t-shirt"
    else:
        return "shorts"
        </pre>
        </div>
        </div>
    <div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=review-double-it-positive data-question_label="2.9.2">
<div id=review-double-it-positive_question class="ac_question col-md-12">
<p>The parameter <code class="docutils literal notranslate"><span class="pre">the_number</span></code> needs to be doubled, but only if <code class="docutils literal notranslate"><span class="pre">the_number</span></code> is positive. Return the doubled value of the number that is passed in if <code class="docutils literal notranslate"><span class="pre">the_number</span></code> is positive. If <code class="docutils literal notranslate"><span class="pre">the_number</span></code> is negative, return -1. If <code class="docutils literal notranslate"><span class="pre">the_number</span></code> is 0, return 0.</p>
<p><code class="docutils literal notranslate"><span class="pre">double_it_positive(5)</span> <span class="pre">→</span> <span class="pre">10</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">double_it_positive(0)</span> <span class="pre">→</span> <span class="pre">0</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">double_it_positive(-4)</span> <span class="pre">→</span> <span class="pre">-1</span></code></p>

</div>
<textarea data-lang="python" id="review-double-it-positive_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def double_it_positive(the_number):
    return 0

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(double_it_positive(5),10,"double_it_positive(5)")
        self.assertEqual(double_it_positive(11),22,"double_it_positive(11)")
        self.assertEqual(double_it_positive(0),0,"double_it_positive(0)")
        self.assertEqual(double_it_positive(-12),-1,"double_it_positive(-12)")
        self.assertEqual(double_it_positive(-1),-1,"double_it_positive(-1)")

myTests().main()
</textarea>
</div>
</div>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.9.3" data-multipleanswers="false"  id="review_test_question5_3_1"  style="visibility: hidden;">
    <p>python-review33: What is a variable’s scope?</p>

    <li data-component="answer"  id="review_test_question5_3_1_opt_a">Its value</li><li data-component="feedback">Value is the contents of the variable.  Scope concerns where the variable is &quot;known&quot;.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question5_3_1_opt_b">The area in the code where a variable can be accessed.</li><li data-component="feedback">Great! Remember that a variable defined inside a function "disappears" after the function has finished executing.</li>
    
    <li data-component="answer"  id="review_test_question5_3_1_opt_c">Its name</li><li data-component="feedback">The name of a variable is just an identifier or alias.  Scope concerns where the variable is &quot;known&quot;.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.9.4" data-multipleanswers="false"  id="review_test_question5_3_2"  style="visibility: hidden;">
    <p>python-review34: What is a local variable?</p>

    <li data-component="answer" data-correct='yes' id="review_test_question5_3_2_opt_a">A temporary variable that is only used inside a function</li><li data-component="feedback">Yes, a local variable is a temporary variable that is only known (only exists) in the function it is defined in.</li>
    
    <li data-component="answer"  id="review_test_question5_3_2_opt_b">The same as a parameter</li><li data-component="feedback">While parameters may be considered local variables, functions may also define and use additional local variables.</li>
    
    <li data-component="answer"  id="review_test_question5_3_2_opt_c">Another name for any variable</li><li data-component="feedback">Variables that are used outside a function are not local, but rather global variables.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.9.5" data-multipleanswers="false"  id="review_test_question5_3_3"  style="visibility: hidden;">
    <p>python-review35: Can you use the same name for a local variable as a global variable?</p>

    <li data-component="answer"  id="review_test_question5_3_3_opt_a">Yes, and there is no reason not to.</li><li data-component="feedback">While there is no problem as far as Python is concerned, it is generally considered bad style because of the potential for the programmer to get confused.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question5_3_3_opt_b">Yes, but it is considered bad form.</li><li data-component="feedback">It is generally considered bad style because of the potential for the programmer to get confused.  If you must use global variables (also generally bad form) make sure they have unique names.</li>
    
    <li data-component="answer"  id="review_test_question5_3_3_opt_c">No, it will cause an error.</li><li data-component="feedback">Python manages global and local scope separately and has clear rules for how to handle variables with the same name in different scopes, so this will not cause a Python error.</li>
    

    </ul>
    </div>
    </section>
<section id="drawing-images-practice">
<h2><span class="section-number">2.10. </span>Drawing Images Practice<a class="headerlink" href="#drawing-images-practice" title="Permalink to this headline">¶</a></h2>
<img alt="../_images/confusion-52.png" src="../_images/confusion-52.png" />

        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_turtle_drawings_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a function that draws the image shown above:</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.10.1"          style="visibility: hidden;">
        import turtle
---
def draw_cross(some_turtle, side_length):
    """Draws a cross shape with the given some_turtle, with each side being of side_length."""
---
    for tower in range(4):
---
        for side in range(3):
---
            some_turtle.forward(side_length)
            some_turtle.left(90)
---
        some_turtle.left(180)
---
canvas = turtle.Screen()
ramisa = turtle.Turtle()
---
draw_cross(ramisa, 50)
        </pre>
        </div>
        </div>
    <p>Draw the following image. Be sure to define and use at least one function as part of your solution.</p>
<img alt="../_images/confusion-191.png" src="../_images/confusion-191.png" />

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=drawing_images_review_1 data-question_label="2.10.2">
<div id=drawing_images_review_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="drawing_images_review_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# first review shape
</textarea>
</div>
</div>
<p>Once you have completed the shape above, try the following. <em>Use the code you made above as a starting point.</em></p>
<img alt="../_images/confusion-271.png" src="../_images/confusion-271.png" />

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=drawing_images_review_2 data-question_label="2.10.3">
<div id=drawing_images_review_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="drawing_images_review_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# second review shape
</textarea>
</div>
</div>
</section>
<section id="image-manipulation-with-nested-loops">
<h2><span class="section-number">2.11. </span>Image Manipulation with Nested Loops<a class="headerlink" href="#image-manipulation-with-nested-loops" title="Permalink to this headline">¶</a></h2>
<img src="../../_static/skflag.png" id="skflag.png">
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_image_manipulation_1" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a function that draws the <em>negative</em> of the image shown above. <em>Note</em>: due to technical limitations with this question, you need to use the <em>x</em> coordinate as your outer loop.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.11.1"          style="visibility: hidden;">
        import image
---
img = image.Image("skflag.png")
---
width = img.get_width()
height = img.get_height()
---
canvas = image.ImageWin(width, height)
img.draw(canvas)
---
for x in range(width):
---
    for y in range(height):
---
        p = img.get_pixel(x, y)
---
        r = 255 - p.get_red()
        g = 255 - p.get_green()
        b = 255 - p.get_blue()
---
        new_pixel = image.Pixel(r, g, b)
---
        img.set_pixel(x, y, new_pixel)
---
img.draw(canvas)
        </pre>
        </div>
        </div>
    </section>
<section id="strings">
<h2><span class="section-number">2.12. </span>Strings<a class="headerlink" href="#strings" title="Permalink to this headline">¶</a></h2>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.1" data-multipleanswers="false"  id="review_string_index_test_question_1"  style="visibility: hidden;">
    <p>python-review40: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_index_test_question_1_opt_a">t</li><li data-component="feedback">Index locations do not start with 1, they start with 0.</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_index_test_question_1_opt_b">h</li><li data-component="feedback">Yes, index locations start with 0.</li>
    
    <li data-component="answer"  id="review_string_index_test_question_1_opt_c">c</li><li data-component="feedback">sentence[-3] would return c, counting from right to left.</li>
    
    <li data-component="answer"  id="review_string_index_test_question_1_opt_d">Error, you cannot use the [ ] operator with a string.</li><li data-component="feedback">[ ] is the index operator</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.2" data-multipleanswers="false"  id="review_string_index_test_question_2"  style="visibility: hidden;">
    <p>python-review41: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">+</span> <span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_string_index_test_question_2_opt_a">tr</li><li data-component="feedback">Yes, indexing operator has precedence over concatenation.</li>
    
    <li data-component="answer"  id="review_string_index_test_question_2_opt_b">ps</li><li data-component="feedback">p is at location 0, not 2.</li>
    
    <li data-component="answer"  id="review_string_index_test_question_2_opt_c">nn</li><li data-component="feedback">n is at location 5, not 2.</li>
    
    <li data-component="answer"  id="review_string_index_test_question_2_opt_d">Error, you cannot use the [ ] operator with the + operator.</li><li data-component="feedback">[ ] operator returns a string that can be concatenated with another string.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.3" data-multipleanswers="false"  id="review_string_length_test_question_1"  style="visibility: hidden;">
    <p>python-review42: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_length_test_question_1_opt_a">11</li><li data-component="feedback">The blank counts as a character.</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_length_test_question_1_opt_b">12</li><li data-component="feedback">Yes, there are 12 characters in the string.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.4" data-multipleanswers="false"  id="review_string_length_test_question_2"  style="visibility: hidden;">
    <p>python-review43: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">)</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_length_test_question_2_opt_a">o</li><li data-component="feedback">Take a look at the index calculation again, len(sentence)-5.</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_length_test_question_2_opt_b">r</li><li data-component="feedback">Yes, len(sentence) is 12 and 12-5 is 7.  Use 7 as index and remember to start counting with 0.</li>
    
    <li data-component="answer"  id="review_string_length_test_question_2_opt_c">s</li><li data-component="feedback">sentence is at index 11</li>
    
    <li data-component="answer"  id="review_string_length_test_question_2_opt_d">Error, len(sentence) is 12 and there is no index 12.</li><li data-component="feedback">You subtract 5 before using the index operator so it will work.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.5" data-multipleanswers="false"  id="review_string_length_test_question_3"  style="visibility: hidden;">
    <p>python-review44: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_string_length_test_question_3_opt_a">c</li><li data-component="feedback">Yes, 3 characters from the end.</li>
    
    <li data-component="answer"  id="review_string_length_test_question_3_opt_b">k</li><li data-component="feedback">Count backward 3 characters.</li>
    
    <li data-component="answer"  id="review_string_length_test_question_3_opt_c">s</li><li data-component="feedback">When expressed with a negative index the last character s is at index -1.</li>
    
    <li data-component="answer"  id="review_string_length_test_question_3_opt_d">Error, negative indices are illegal.</li><li data-component="feedback">Python does use negative indices to count backward from the end.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.6" data-multipleanswers="false"  id="review_string_length_test_question_4"  style="visibility: hidden;">
    <p>python-review45: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_length_test_question_4_opt_a">python</li><li data-component="feedback">That would be sentence[0:6].</li>
    
    <li data-component="answer"  id="review_string_length_test_question_4_opt_b">rocks</li><li data-component="feedback">That would be sentence[7:].</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_length_test_question_4_opt_c">hon r</li><li data-component="feedback">Yes, start with the character at index 3 and go up to but not include the character at index 8.</li>
    
    <li data-component="answer"  id="review_string_length_test_question_4_opt_d">Error, you cannot have two numbers inside the [ ].</li><li data-component="feedback">This is called slicing, not indexing.  It requires a start and an end.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.7" data-multipleanswers="false"  id="review_string_for_loop_by_item_test_question_1"  style="visibility: hidden;">
    <p>python-review46: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_1_opt_a">10</li><li data-component="feedback">Iteration by item will process once for each item in the sequence.</li>
    
    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_1_opt_b">11</li><li data-component="feedback">The blank is part of the sequence.</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_for_loop_by_item_test_question_1_opt_c">12</li><li data-component="feedback">Yes, there are 12 characters, including the blank.</li>
    
    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_1_opt_d">Error, the for statement needs to use the range function.</li><li data-component="feedback">The for statement can iterate over a sequence item by item.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.8" data-multipleanswers="false"  id="review_string_for_loop_by_item_test_question_2"  style="visibility: hidden;">
    <p>python-review47: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">]:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_2_opt_a">4</li><li data-component="feedback">Slice returns a sequence that can be iterated over.</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_for_loop_by_item_test_question_2_opt_b">5</li><li data-component="feedback">Yes, The blank is part of the sequence returned by slice</li>
    
    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_2_opt_c">6</li><li data-component="feedback">Check the result of s[3:8].  It does not include the item at index 8.</li>
    
    <li data-component="answer"  id="review_string_for_loop_by_item_test_question_2_opt_d">Error, the for statement cannot use slice.</li><li data-component="feedback">Slice returns a sequence.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.9" data-multipleanswers="false"  id="review_string_accumulator_pattern_test_question_1"  style="visibility: hidden;">
    <p>python-review48: What is printed by the following statements:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">some_string</span> <span class="o">=</span> <span class="s2">&quot;ball&quot;</span>
<span class="n">another_string</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
<span class="k">for</span> <span class="n">item</span> <span class="ow">in</span> <span class="n">some_string</span><span class="p">:</span>
    <span class="n">another_string</span> <span class="o">=</span> <span class="n">another_string</span> <span class="o">+</span> <span class="n">item</span>
<span class="nb">print</span><span class="p">(</span><span class="n">another_string</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_string_accumulator_pattern_test_question_1_opt_a">ball</li><li data-component="feedback">Yes, the repeated concatenation will cause another_string to become the same as some_string.</li>
    
    <li data-component="answer"  id="review_string_accumulator_pattern_test_question_1_opt_b">llab</li><li data-component="feedback">Look again at the *order* of the concatenation!</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.10" data-multipleanswers="false"  id="review_string_accumulator_pattern_test_question_2"  style="visibility: hidden;">
    <p>python-review49: What is printed by the following statements:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">some_string</span> <span class="o">=</span> <span class="s2">&quot;ball&quot;</span>
<span class="n">another_string</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
<span class="k">for</span> <span class="n">item</span> <span class="ow">in</span> <span class="n">some_string</span><span class="p">:</span>
    <span class="n">another_string</span> <span class="o">=</span> <span class="n">item</span> <span class="o">+</span> <span class="n">another_string</span>
<span class="nb">print</span><span class="p">(</span><span class="n">another_string</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_accumulator_pattern_test_question_2_opt_a">ball</li><li data-component="feedback">Look again at the *order* of the concatenation!</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_accumulator_pattern_test_question_2_opt_b">llab</li><li data-component="feedback">Yes, the order is reversed due to the order of the concatenation.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.12.11" data-multipleanswers="false"  id="review_string_methods_check_understanding1"  style="visibility: hidden;">
    <p>python-review50: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">s</span><span class="o">.</span><span class="n">count</span><span class="p">(</span><span class="s2">&quot;o&quot;</span><span class="p">)</span> <span class="o">+</span> <span class="n">s</span><span class="o">.</span><span class="n">count</span><span class="p">(</span><span class="s2">&quot;p&quot;</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_string_methods_check_understanding1_opt_a">0</li><li data-component="feedback">There are definitely o and p characters.</li>
    
    <li data-component="answer"  id="review_string_methods_check_understanding1_opt_b">2</li><li data-component="feedback">There are 2 o characters but what about p?</li>
    
    <li data-component="answer" data-correct='yes' id="review_string_methods_check_understanding1_opt_c">3</li><li data-component="feedback">Yes, add the number of o characters and the number of p characters.</li>
    

    </ul>
    </div>
    
        <div class="runestone parsons-container ">
        <div data-component="parsons" id="review_even_letters" class="parsons" >
        <div class="parsons_question parsons-text" >
    <p>Construct a block of code that correctly creates a function with a single parameter word that returns the even letters of the word, then calls the function and prints the result.</p>

        </div>
        <pre  class="parsonsblocks" data-question_label="2.12.12"          style="visibility: hidden;">
        def even_letters(word):
    new_word = ''
    counter = 0
    for letter in word:
        if counter % 2 == 0:
            new_word = new_word + letter
        counter = counter + 1
    return new_word
changed_word = even_letters("Saskatchewan")
print(changed_word)
        </pre>
        </div>
        </div>
    </section>
<section id="lists">
<h2><span class="section-number">2.13. </span>Lists<a class="headerlink" href="#lists" title="Permalink to this headline">¶</a></h2>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.1" data-multipleanswers="false"  id="review_test_question9_2_1"  style="visibility: hidden;">
    <p>python-review52: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_test_question9_2_1_opt_a">4</li><li data-component="feedback">len returns the actual number of items in the list, not the maximum index value.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question9_2_1_opt_b">5</li><li data-component="feedback">Yes, there are 5 items in this list.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.2" data-multipleanswers="false"  id="review_test_question9_2_2"  style="visibility: hidden;">
    <p>python-review53: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_test_question9_2_2_opt_a">7</li><li data-component="feedback">Yes, there are 7 items in this list even though two of them happen to also be lists.</li>
    
    <li data-component="answer"  id="review_test_question9_2_2_opt_b">8</li><li data-component="feedback">len returns the number of top level items in the list.  It does not count items in sublists.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.3" data-multipleanswers="false"  id="review_test_question9_3_1"  style="visibility: hidden;">
    <p>python-review54: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_test_question9_3_1_opt_a">[ ]</li><li data-component="feedback">The empty list is at index 4.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question9_3_1_opt_b">3.14</li><li data-component="feedback">Yes, 3.14 is at index 5 since we start counting at 0 and sublists count as one item.</li>
    
    <li data-component="answer"  id="review_test_question9_3_1_opt_c">False</li><li data-component="feedback">False is at index 6.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.4" data-multipleanswers="false"  id="review_test_question9_3_3"  style="visibility: hidden;">
    <p>python-review55: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">2</span><span class="p">][</span><span class="mi">0</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_test_question9_3_3_opt_a">56</li><li data-component="feedback">Indexes start with 0, not 1.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question9_3_3_opt_b">c</li><li data-component="feedback">Yes, the first character of the string at index 2 is c</li>
    
    <li data-component="answer"  id="review_test_question9_3_3_opt_c">cat</li><li data-component="feedback">cat is the item at index 2 but then we index into it further.</li>
    
    <li data-component="answer"  id="review_test_question9_3_3_opt_d">Error, you cannot have two index values unless you are using slicing.</li><li data-component="feedback">Using more than one index is fine.  You read it from left to right.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.5" data-multipleanswers="false"  id="review_test_question9_4_1"  style="visibility: hidden;">
    <p>python-review56: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="mf">3.14</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_test_question9_4_1_opt_a">True</li><li data-component="feedback">Yes, 3.14 is an item in the list a_list.</li>
    
    <li data-component="answer"  id="review_test_question9_4_1_opt_b">False</li><li data-component="feedback">There are 7 items in the list, 3.14 is one of them.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.6" data-multipleanswers="false"  id="review_test_question9_4_2"  style="visibility: hidden;">
    <p>python-review57: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="mi">57</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_test_question9_4_2_opt_a">True</li><li data-component="feedback">in returns True for top level items only.  57 is in a sublist.</li>
    
    <li data-component="answer" data-correct='yes' id="review_test_question9_4_2_opt_b">False</li><li data-component="feedback">Yes, 57 is not a top level item in a_list.  It is in a sublist.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.7" data-multipleanswers="false"  id="review_test_question9_6_1"  style="visibility: hidden;">
    <p>python-review58: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">4</span><span class="p">:])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="review_test_question9_6_1_opt_a">[ [ ], 3.14, False]</li><li data-component="feedback">Yes, the slice starts at index 4 and goes up to and including the last item.</li>
    
    <li data-component="answer"  id="review_test_question9_6_1_opt_b">[ [ ], 3.14]</li><li data-component="feedback">By leaving out the upper bound on the slice, we go up to and including the last item.</li>
    
    <li data-component="answer"  id="review_test_question9_6_1_opt_c">[ [56, 57, "dog"], [ ], 3.14, False]</li><li data-component="feedback">Index values start at 0.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.8" data-multipleanswers="false"  id="review_list_methods_check_1"  style="visibility: hidden;">
    <p>python-review59: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">4</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="n">a_list</span><span class="o">.</span><span class="n">append</span><span class="p">(</span><span class="kc">True</span><span class="p">)</span>
<span class="n">a_list</span><span class="o">.</span><span class="n">append</span><span class="p">(</span><span class="kc">False</span><span class="p">)</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_list_methods_check_1_opt_a">[4, 2, 8, 6, 5, False, True]</li><li data-component="feedback">True was added first, then False was added last.</li>
    
    <li data-component="answer" data-correct='yes' id="review_list_methods_check_1_opt_b">[4, 2, 8, 6, 5, True, False]</li><li data-component="feedback">Yes, each item is added to the end of the list.</li>
    
    <li data-component="answer"  id="review_list_methods_check_1_opt_c">[True, False, 4, 2, 8, 6, 5]</li><li data-component="feedback">append adds at the end, not the beginning.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.9" data-multipleanswers="false"  id="review_list_methods_check_2"  style="visibility: hidden;">
    <p>python-review60: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">4</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="n">temp</span> <span class="o">=</span> <span class="n">a_list</span><span class="o">.</span><span class="n">pop</span><span class="p">(</span><span class="mi">2</span><span class="p">)</span>
<span class="n">temp</span> <span class="o">=</span> <span class="n">a_list</span><span class="o">.</span><span class="n">pop</span><span class="p">()</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_list_methods_check_2_opt_a">[4, 8, 6]</li><li data-component="feedback">pop(2) removes the item at index 2, not the 2 itself.</li>
    
    <li data-component="answer"  id="review_list_methods_check_2_opt_b">[2, 6, 5]</li><li data-component="feedback">pop() removes the last item, not the first.</li>
    
    <li data-component="answer" data-correct='yes' id="review_list_methods_check_2_opt_c">[4, 2, 6]</li><li data-component="feedback">Yes, first the 8 was removed, then the last item, which was 5.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="2.13.10" data-multipleanswers="false"  id="review_list_methods_check_3"  style="visibility: hidden;">
    <p>python-review61: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">4</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="n">a_list</span> <span class="o">=</span> <span class="n">a_list</span><span class="o">.</span><span class="n">pop</span><span class="p">(</span><span class="mi">0</span><span class="p">)</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="review_list_methods_check_3_opt_a">[2, 8, 6, 5]</li><li data-component="feedback">a_list is now the value that was returned from pop(0).</li>
    
    <li data-component="answer"  id="review_list_methods_check_3_opt_b">[4, 2, 8, 6, 5]</li><li data-component="feedback">pop(0) changes the list by removing the first item.</li>
    
    <li data-component="answer" data-correct='yes' id="review_list_methods_check_3_opt_c">4</li><li data-component="feedback">Yes, first the 4 was removed from the list, then returned and assigned to a_list.  The list is lost.</li>
    
    <li data-component="answer"  id="review_list_methods_check_3_opt_d">None</li><li data-component="feedback">pop(0) returns the first item in the list so a_list has now been changed.</li>
    

    </ul>
    </div>
    <div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=review_list_overview_practice_problem_3 data-question_label="2.13.11">
<div id=review_list_overview_practice_problem_3_question class="ac_question col-md-12">
<p>Write a function to count how many odd numbers are in a list.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([1,3,5,7,9])</span> <span class="pre">→</span> <span class="pre">5</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([1,2,3,4,5])</span> <span class="pre">→</span> <span class="pre">3</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([2,4,6,8,10])</span> <span class="pre">→</span> <span class="pre">0</span></code></p>

</div>
<textarea data-lang="python" id="review_list_overview_practice_problem_3_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def count_odds(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(count_odds([1,3,5,7,9]),5,"count_odds([1,3,5,7,9])")
      self.assertEqual(count_odds([1,2,3,4,5]),3,"count_odds([-1,-2,-3,-4,-5])")
      self.assertEqual(count_odds([2,4,6,8,10]),0,"count_odds([2,4,6,8,10])")
      self.assertEqual(count_odds([0,-1,12,-33]),2,"count_odds([0,-1,12,-33])")

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="number-guessing-game">
<h2><span class="section-number">2.14. </span>Number Guessing Game<a class="headerlink" href="#number-guessing-game" title="Permalink to this headline">¶</a></h2>
<p>Implement a number guessing game in Python that does the following:</p>
<ul>
<li><p>generate a random number from 1 to 100 and stores it in a variable</p></li>
<li><p>repeat the following until the user guesses the number:</p>
<blockquote>
<div><ul class="simple">
<li><p>get the user to guess the number</p></li>
<li><p>tell the user if the number is too high or too low</p></li>
</ul>
</div></blockquote>
</li>
<li><p>congratulate the user when they guess the correct number with a message such as “Way to go! You guessed the right number in 9 tries!”</p></li>
</ul>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=review_number_guessing_game data-question_label="2.14.1">
<div id=review_number_guessing_game_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="review_number_guessing_game_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='Create a number guessing game here!'  data-wasm=/_static
     style="visibility: hidden;">
# your code here!
</textarea>
</div>
</div>
</section>
</section>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - <span class="section-number">1. </span>Reeborg Review' data-toggle="tooltip">
    <a href="Reeborg.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Reeborg.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Underlying Technology of Computing Devices and the Internet' data-toggle="tooltip" >
    <a href="../ComputingAndSociety/UnderlyingTechnology.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="../ComputingAndSociety/UnderlyingTechnology.html" >Next Section - Underlying Technology of Computing Devices and the Internet</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.
    </p>
  </div>
</footer>


  </body>
</html>