<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>1. Fruitful Functions and Variable Scope &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="2. Writing Fruitful Functions Practice" href="WritingFruitfulFunctionsPractice.html" />
    <link rel="prev" title="5. Loops Practice Quiz and Assignment Work Period" href="../MoreTurtles/PracticeQuizAndWorkPeriod.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">1. Fruitful Functions and Variable Scope</a><ul>
<li><a class="reference internal" href="#warmup-problem">1.1. Warmup Problem</a></li>
<li><a class="reference internal" href="#functions-that-return-values">1.2. Functions that Return Values</a><ul>
<li><a class="reference internal" href="#check-your-understanding">1.2.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#variables-and-parameters-are-local">1.3. Variables and Parameters are Local</a><ul>
<li><a class="reference internal" href="#id1">1.3.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#practice-problems">1.4. Practice Problems</a><ul>
<li><a class="reference internal" href="#double-it">1.4.1. Double It</a></li>
<li><a class="reference internal" href="#double-it-positive">1.4.2. Double It Positive</a></li>
<li><a class="reference internal" href="#wear-the-right-thing">1.4.3. Wear the Right Thing</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 5. Loops Practice Quiz and Assignment Work Period" data-toggle="tooltip">
    <a href="../MoreTurtles/PracticeQuizAndWorkPeriod.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 2. Writing Fruitful Functions Practice' data-toggle="tooltip" >
    <a href="WritingFruitfulFunctionsPractice.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="fruitful-functions-and-variable-scope">
<h1>1. Fruitful Functions and Variable Scope<a class="headerlink" href="#fruitful-functions-and-variable-scope" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>Practice creating a function that draws a shape with the turtle module. Introduce fruitful functions (functions that return values). Discuss local and global scope. Work on practice problems, in which students create functions that return values.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide">
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to encapsulate reusable pieces of code.</li>
</ul>

    </div>
    <div class="section" id="warmup-problem">
<h2>1.1. Warmup Problem<a class="headerlink" href="#warmup-problem" title="Permalink to this headline">¶</a></h2>
<img alt="../../_images/confusion-16.png" src="../../_images/confusion-16.png" />
<p>Draw the image above using the Python turtle module. <strong>You must define a function as part of your solution!</strong></p>

<div data-childcomponent="turtle_confusion_practice_16" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=turtle_confusion_practice_16 data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle
</textarea>
</div>
</div>
<div class="section" id="functions-that-return-values">
<h2>1.2. Functions that Return Values<a class="headerlink" href="#functions-that-return-values" title="Permalink to this headline">¶</a></h2>
<p>Most functions require arguments, values that control how the function does its job. For example, if you want to find the absolute value of a number (defined as the distance of a number from zero on a number line), you have to indicate what the number is. Python has a built-in function for computing the absolute value:</p>

<div data-childcomponent="more_function_details_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=more_function_details_1 data-lang="python"    data-timelimit=25000   data-audio=''     >
print(abs(5))

print(abs(-5))
</textarea>
</div>
<p>In this example, the arguments to the <code class="docutils literal"><span class="pre">abs</span></code> function are 5 and -5.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The variables inside the function are called <strong>parameters</strong>, and the values that are passed to them when the function is called are <strong>arguments</strong>. For example, when we defined a function like <code class="docutils literal"><span class="pre">def</span> <span class="pre">draw_square(some_turtle,</span> <span class="pre">side_length):</span></code>, <code class="docutils literal"><span class="pre">some_turtle</span></code> and <code class="docutils literal"><span class="pre">side_length</span></code> are parameters. When we call that function, like <code class="docutils literal"><span class="pre">draw_square(tess,</span> <span class="pre">100)</span></code>, <code class="docutils literal"><span class="pre">tess</span></code> and <code class="docutils literal"><span class="pre">100</span></code> are arguments.</p>
</div>
<p>Some functions take more than one argument. For example the math module contains a function called <code class="docutils literal"><span class="pre">pow</span></code> which takes two arguments, the base and the exponent.</p>

<div data-childcomponent="more_function_details_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=more_function_details_2 data-lang="python"    data-timelimit=25000   data-audio=''     >
import math
print(math.pow(2, 3))

print(math.pow(7, 2))
</textarea>
</div>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">Of course, we have already seen that raising a base to an exponent can be done with the ** operator.</p>
</div>
<p>Another built-in function that takes more than one argument is <code class="docutils literal"><span class="pre">max</span></code>.</p>

<div data-childcomponent="more_function_details_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=more_function_details_3 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
print(max(7, 11))
print(max(4, 1, 17, 2, 12))
print(max(11, 5 * 3, 25 - 9, 0))
</textarea>
</div>
<p><code class="docutils literal"><span class="pre">max</span></code> can be sent any number of arguments, separated by commas, and will
return the maximum value sent. The arguments can be either simple values or
expressions. In the last example, 16 is returned, since it is larger than 11,
15, and 0.  Note that <code class="docutils literal"><span class="pre">max</span></code> also works on lists of values.</p>
<p>Notice that each of the above functions return values that can be used to build more complex expressions.</p>
<p>So an important difference between these functions and ones like <code class="docutils literal"><span class="pre">draw_square</span></code> or <code class="docutils literal"><span class="pre">draw_rectangle</span></code> is that <code class="docutils literal"><span class="pre">draw_square</span></code> was not executed because we wanted it to compute a value — on the contrary, we wrote <code class="docutils literal"><span class="pre">draw_square</span></code> because we wanted it to execute a sequence of steps that caused the turtle to draw a specific shape.</p>
<p>Functions that return values are sometimes called <strong>fruitful functions</strong>.
A function that doesn’t return a value can be simply called a function, or if we want
to make it extremely clear, a <em>non-fruitful</em> function.</p>
<p>Fruitful functions still allow the user to provide information (arguments).  However there is now an additional piece of data that is returned from the function.</p>
<img alt="../../_images/blackboxfun.png" src="../../_images/blackboxfun.png" />
<p>How do we write our own fruitful function?  Let’s start by creating a very simple
mathematical function that we will call <code class="docutils literal"><span class="pre">square</span></code>.  The square function will take one number
as a parameter and return the result of squaring that number.  Here is the
black-box diagram with the Python code following.</p>
<img alt="../../_images/squarefun.png" src="../../_images/squarefun.png" />

<div data-childcomponent="more_function_details_4" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=more_function_details_4 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def square(original_number):
    squared_value = original_number * original_number
    return squared_value

to_square = 10
result = square(to_square)
print("The result of", to_square, "squared is", result)
</textarea>
</div>
<p>The <strong>return</strong> statement is followed by an expression which is evaluated.  Its
result is returned to the caller as the “fruit” of calling this function.
Because the return statement can contain any Python expression we could have
avoided creating the <strong>temporary variable</strong> <code class="docutils literal"><span class="pre">squared_value</span></code> and simply used
<code class="docutils literal"><span class="pre">return</span> <span class="pre">original_number</span> <span class="pre">*</span> <span class="pre">original_number</span></code>.
Try modifying the square function above to see that this works just the same.
On the other hand, using <strong>temporary variables</strong> like <code class="docutils literal"><span class="pre">squared_value</span></code> in the program above makes debugging easier.  These temporary variables are examples of <strong>local variables</strong>, which will be explained further in the next section.</p>
<p>Notice something important here. The name of the variable we pass as an
argument — <code class="docutils literal"><span class="pre">to_square</span></code> — has nothing to do with the name of the formal parameter
— <code class="docutils literal"><span class="pre">original_number</span></code>.  It is as if  <code class="docutils literal"><span class="pre">original_number</span> <span class="pre">=</span> <span class="pre">to_square</span></code> is executed when <code class="docutils literal"><span class="pre">square</span></code> is called. It doesn’t matter what the value was named in
the caller. In <code class="docutils literal"><span class="pre">square</span></code>, it’s name is <code class="docutils literal"><span class="pre">original_number</span></code>.  You can see this very clearly in
codelens, where the <strong>global variables</strong> (variables defined outside of any function) and the local variables for the square function are in separate boxes.</p>
<p>As you step through the example in codelens below, notice that the <strong>return</strong> statement not only causes the function to return a value, but it also returns the flow of control back to the place in the program where the function call was made.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The call to a function <em>terminates</em> after the execution of a return statement.
This is fairly obvious if the return statement is the last statement in the function, but
we will see later where it makes sense to have a return statement even when other statements
follow, and the further statements are <em>not</em> executed.</p>
</div>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="more_function_details_codelens_square"></div>
<p class="cl_caption"><span class="cl_caption_text"> (more_function_details_codelens_square)</span> </p>
</div>
<script type="text/javascript">
var more_function_details_codelens_square_trace = {"code": "def square(original_number):\n    squared_value = original_number * original_number\n    return squared_value\n\nto_square = 10\nresult = square(to_square)\nprint(\"The result of\", to_square, \"squared is\", result)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100}, "ordered_varnames": ["original_number", "squared_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100, "__return__": 100}, "ordered_varnames": ["original_number", "squared_value", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10, "result": 100}, "ordered_globals": ["square", "to_square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 7, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10, "result": 100}, "ordered_globals": ["square", "to_square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "The result of 10 squared is 100\n"}]};
var more_function_details_codelens_square_vis;

$(document).ready(function() {
    try {
        more_function_details_codelens_square_vis = new ExecutionVisualizer('more_function_details_codelens_square',more_function_details_codelens_square_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(more_function_details_codelens_square_vis,'more_function_details_codelens_square');
        styleButtons('more_function_details_codelens_square');
        allVisualizers.push(more_function_details_codelens_square_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component more_function_details_codelens_square_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#more_function_details_codelens_square_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'more_function_details_codelens_square'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (more_function_details_codelens_square_vis) {
        more_function_details_codelens_square_vis.redrawConnectors();
    }
});
</script>
</div>
<p>Another important thing to notice as you step through this codelens
demonstration is the movement of the red and green arrows.  Codelens uses these arrows to show you where it is currently executing.
Recall that the red arrow always points to the next line of code that will be executed.  The light green arrow points to the line
that was just executed in the last step.</p>
<p>When you first start running this codelens demonstration you will notice that there is only a red arrow and it points to
line 1.  This is because line 1 is the next line to be executed and since it is the first line, there is no previously executed line
of code.</p>
<p>When you click on the forward button, notice that the red arrow moves to line 5, skipping lines 2 and 3 of the function (and
the light green arrow has now appeared on line 1).  Why is this?
The answer is that function definition is not the same as function execution.  Lines 2
and 3 will not be executed until the function is called on line 6.  Line 1 defines the function and the name <code class="docutils literal"><span class="pre">square</span></code> is added to the
global variables, but that is all the <code class="docutils literal"><span class="pre">def</span></code> does at that point.  The body of the function will be executed later.  Continue to click
the forward button to see how the flow of control moves from the call, back up to the body of the function, and then finally back to line 7, after the function has returned its value and the value has been assigned to <code class="docutils literal"><span class="pre">result</span></code>.</p>
<p>Finally, there is one more aspect of function return values that should be noted.  All Python functions return the value <code class="docutils literal"><span class="pre">None</span></code> unless there is an explicit return statement with
a value other than <code class="docutils literal"><span class="pre">None.</span></code>
Consider the following common mistake made by beginning Python
programmers.  As you step through this example, pay very close attention to the return
value in the local variables listing.  Then look at what is printed when the
function returns.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="more_function_details_codelens_square_bad"></div>
<p class="cl_caption"><span class="cl_caption_text"> (more_function_details_codelens_square_bad)</span> </p>
</div>
<script type="text/javascript">
var more_function_details_codelens_square_bad_trace = {"code": "def square(original_number):\n    squared_value = original_number * original_number\n    print(squared_value)  # Bad! Should use return instead of print!\n\nto_square = 10\nresult = square(to_square)\nprint(\"The result of\", to_square, \"squared is\", result)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100}, "ordered_varnames": ["original_number", "squared_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "to_square": 10}, "ordered_globals": ["square", "to_square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100, "__return__": null}, "ordered_varnames": ["original_number", "squared_value", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "100\n"}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10, "result": null}, "ordered_globals": ["square", "to_square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "100\n"}, {"line": 7, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "to_square": 10, "result": null}, "ordered_globals": ["square", "to_square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "100\nThe result of 10 squared is None\n"}]};
var more_function_details_codelens_square_bad_vis;

$(document).ready(function() {
    try {
        more_function_details_codelens_square_bad_vis = new ExecutionVisualizer('more_function_details_codelens_square_bad',more_function_details_codelens_square_bad_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(more_function_details_codelens_square_bad_vis,'more_function_details_codelens_square_bad');
        styleButtons('more_function_details_codelens_square_bad');
        allVisualizers.push(more_function_details_codelens_square_bad_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component more_function_details_codelens_square_bad_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#more_function_details_codelens_square_bad_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'more_function_details_codelens_square_bad'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (more_function_details_codelens_square_bad_vis) {
        more_function_details_codelens_square_bad_vis.redrawConnectors();
    }
});
</script>
</div>
<p>The problem with this function is that even though it prints the value of the square, that value will not be returned to the place
where the call was made.  Since line 6 uses the return value as the right hand side of an assignment statement, the evaluation of the
function will be <code class="docutils literal"><span class="pre">None</span></code>.  In this case, <code class="docutils literal"><span class="pre">result</span></code> will refer to that value after the assignment statement and therefore the result printed in line 7 is incorrect.  Typically, functions will return values that can be printed or processed in some other way by the caller.</p>
<div class="section" id="check-your-understanding">
<h3>1.2.1. Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question5_2_1">
            <p>variable-scope1: What is wrong with the following function definition:</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">add_em</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="n">y</span><span class="p">,</span> <span class="n">z</span><span class="p">):</span>
    <span class="k">return</span> <span class="n">x</span> <span class="o">+</span> <span class="n">y</span> <span class="o">+</span> <span class="n">z</span>
    <span class="k">print</span><span class="p">(</span><span class="s1">&#39;the answer is&#39;</span><span class="p">,</span> <span class="n">x</span> <span class="o">+</span> <span class="n">y</span> <span class="o">+</span> <span class="n">z</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question5_2_1_opt_a">You should never use a print statement in a function definition.</li><li data-component="feedback" id="test_question5_2_1_opt_a">Although you should not mistake print for return, you may include print statements inside your functions.</li>
            
            <li data-component="answer" data-correct id="test_question5_2_1_opt_b">You should not have any statements in a function after the return statement.  Once the function gets to the return statement it will immediately stop executing the function.</li><li data-component="feedback" id="test_question5_2_1_opt_b">This is a very common mistake so be sure to watch out for it when you write your code!</li>
            
            <li data-component="answer"  id="test_question5_2_1_opt_c">You must calculate the value of x+y+z before you return it.</li><li data-component="feedback" id="test_question5_2_1_opt_c">Python will automatically calculate the value x+y+z and then return it in the statement as it is written</li>
            
            <li data-component="answer"  id="test_question5_2_1_opt_d">A function cannot return a number.</li><li data-component="feedback" id="test_question5_2_1_opt_d">Functions can return any legal data, including (but not limited to) numbers, strings, turtles, etc.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question5_2_2">
            <p>variable-scope2: What will the following function return?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">add_em</span><span class="p">(</span><span class="n">x</span><span class="p">,</span> <span class="n">y</span><span class="p">,</span> <span class="n">z</span><span class="p">):</span>
    <span class="k">print</span><span class="p">(</span><span class="n">x</span> <span class="o">+</span> <span class="n">y</span> <span class="o">+</span> <span class="n">z</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="test_question5_2_2_opt_a">None</li><li data-component="feedback" id="test_question5_2_2_opt_a">We have accidentally used print where we mean return.  Therefore, the function will return the value None by default.  This is a VERY COMMON mistake so watch out!  This mistake is also particularly difficult to find because when you run the function the output looks the same.  It is not until you try to assign its value to a variable that you can notice a difference.</li>
            
            <li data-component="answer"  id="test_question5_2_2_opt_b">The value of x + y + z</li><li data-component="feedback" id="test_question5_2_2_opt_b">Careful!  This is a very common mistake.  Here we have printed the value x+y+z but we have not returned it.  To return a value we MUST use the return keyword.</li>
            
            <li data-component="answer"  id="test_question5_2_2_opt_c">The string 'x + y + z'</li><li data-component="feedback" id="test_question5_2_2_opt_c">x+y+z calculates a number (assuming x+y+z are numbers) which represents the sum of the values x, y and z.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="variables-and-parameters-are-local">
<h2>1.3. Variables and Parameters are Local<a class="headerlink" href="#variables-and-parameters-are-local" title="Permalink to this headline">¶</a></h2>
<p>An assignment statement in a function creates a <strong>local variable</strong> for the
variable on the left hand side of the assignment operator. It is called local because this variable only
exists inside the function and you cannot use it outside. For example,
consider again the <code class="docutils literal"><span class="pre">square</span></code> function:</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="bad_local"></div>
<p class="cl_caption"><span class="cl_caption_text"> (bad_local)</span> </p>
</div>
<script type="text/javascript">
var bad_local_trace = {"code": "def square(original_number):\n    squared_value = original_number * original_number\n    return squared_value\n\nresult = square(10)\nprint(squared_value)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100}, "ordered_varnames": ["original_number", "squared_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 3, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 10, "squared_value": 100, "__return__": 100}, "ordered_varnames": ["original_number", "squared_value", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "result": 100}, "ordered_globals": ["square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 6, "event": "exception", "func_name": "<module>", "globals": {"square": ["REF", 1], "result": 100}, "ordered_globals": ["square", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "", "exception_msg": "NameError: name 'squared_value' is not defined"}]};
var bad_local_vis;

$(document).ready(function() {
    try {
        bad_local_vis = new ExecutionVisualizer('bad_local',bad_local_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(bad_local_vis,'bad_local');
        styleButtons('bad_local');
        allVisualizers.push(bad_local_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component bad_local_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#bad_local_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'bad_local'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (bad_local_vis) {
        bad_local_vis.redrawConnectors();
    }
});
</script>
</div>
<p>If you press the ‘last &gt;&gt;’ button you will see an error message.
When we try to use <code class="docutils literal"><span class="pre">squared_value</span></code> on line 6 (outside the function) Python looks for a global
variable named <code class="docutils literal"><span class="pre">squared_value</span></code> but does not find one.  This results in the
error: <code class="docutils literal"><span class="pre">Name</span> <span class="pre">Error:</span> <span class="pre">'squared_value'</span> <span class="pre">is</span> <span class="pre">not</span> <span class="pre">defined.</span></code></p>
<p>The variable <code class="docutils literal"><span class="pre">squared_value</span></code> only exists while the function is being executed —
we call this its <strong>lifetime</strong>.
When the execution of the function terminates (returns),
the local variables  are destroyed.  Codelens helps you  visualize this
because the local variables disappear after the function returns.  Go back and step through the
statements paying particular attention to the variables that are created when the function is called.
Note when they are subsequently destroyed as the function returns.</p>
<p>Formal parameters are also local and act like local variables.
For example, the lifetime of <code class="docutils literal"><span class="pre">original_number</span></code> begins when <code class="docutils literal"><span class="pre">square</span></code> is
called, and its lifetime ends when the function completes its execution.</p>
<p>So it is not possible for a function to set some local variable to a
value, complete its execution, and then when it is called again next
time, recover the local variable.  Each call of the function creates
new local variables, and their lifetimes expire when the function returns
to the caller.</p>
<p>On the other hand, it is legal for a function to access a global variable.  However, this is considered
<strong>bad form</strong> by nearly all programmers and should be avoided.  Look at the following,
nonsensical variation of the square function.</p>

<div data-childcomponent="badsquare_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=badsquare_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def bad_square(original_number):
    squared_value = original_number ** power
    return squared_value

power = 2
result = bad_square(10)
print(result)
</textarea>
</div>
<p>Although the <code class="docutils literal"><span class="pre">bad_square</span></code> function works, it is silly and poorly written.  We have done it here to illustrate
an important rule about how variables are looked up in Python.
First, Python looks at the variables that are defined as local variables in
the function.  We call this the <strong>local scope</strong>.  If the variable name is not
found in the local scope, then Python looks at the global variables,
or <strong>global scope</strong>.  This is exactly the case illustrated in the code above.
<code class="docutils literal"><span class="pre">power</span></code> is not found locally in <code class="docutils literal"><span class="pre">bad_square</span></code> but it does exist globally.
The appropriate way to write this function would be to pass power as a parameter.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">For practice, you should rewrite the bad_square example above to have a second parameter called <code class="docutils literal"><span class="pre">power</span></code>.</p>
</div>
<p>There is another variation on this theme of local versus global variables.  Assignment statements in the local function cannot
change variables defined outside the function, without further (discouraged) special syntax.  Consider the following
codelens example:</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="codelens_powerof_bad"></div>
<p class="cl_caption"><span class="cl_caption_text"> (codelens_powerof_bad)</span> </p>
</div>
<script type="text/javascript">
var codelens_powerof_bad_trace = {"code": "def power_of(original_number, p):\n    power = p   # Another dumb mistake\n    new_value = original_number ** power\n    return new_value\n\npower = 3\nresult = power_of(10, 2)\nprint(result)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "<module>", "globals": {"power_of": ["REF", 1]}, "ordered_globals": ["power_of"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "<module>", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "power_of", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [{"func_name": "power_of", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"p": 2, "original_number": 10}, "ordered_varnames": ["original_number", "p"], "is_zombie": false, "is_highlighted": true, "unique_hash": "power_of_f1"}], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "power_of", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [{"func_name": "power_of", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"p": 2, "original_number": 10}, "ordered_varnames": ["original_number", "p"], "is_zombie": false, "is_highlighted": true, "unique_hash": "power_of_f1"}], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "power_of", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [{"func_name": "power_of", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"p": 2, "original_number": 10, "power": 2}, "ordered_varnames": ["original_number", "p", "power"], "is_zombie": false, "is_highlighted": true, "unique_hash": "power_of_f1"}], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "power_of", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [{"func_name": "power_of", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"p": 2, "original_number": 10, "power": 2, "new_value": 100}, "ordered_varnames": ["original_number", "p", "power", "new_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "power_of_f1"}], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 4, "event": "return", "func_name": "power_of", "globals": {"power_of": ["REF", 1], "power": 3}, "ordered_globals": ["power_of", "power"], "stack_to_render": [{"func_name": "power_of", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"p": 2, "original_number": 10, "power": 2, "new_value": 100, "__return__": 100}, "ordered_varnames": ["original_number", "p", "power", "new_value", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "power_of_f1"}], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 8, "event": "step_line", "func_name": "<module>", "globals": {"power_of": ["REF", 1], "power": 3, "result": 100}, "ordered_globals": ["power_of", "power", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": ""}, {"line": 8, "event": "return", "func_name": "<module>", "globals": {"power_of": ["REF", 1], "power": 3, "result": 100}, "ordered_globals": ["power_of", "power", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "power_of(original_number, p)", null]}, "stdout": "100\n"}]};
var codelens_powerof_bad_vis;

$(document).ready(function() {
    try {
        codelens_powerof_bad_vis = new ExecutionVisualizer('codelens_powerof_bad',codelens_powerof_bad_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(codelens_powerof_bad_vis,'codelens_powerof_bad');
        styleButtons('codelens_powerof_bad');
        allVisualizers.push(codelens_powerof_bad_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component codelens_powerof_bad_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#codelens_powerof_bad_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'codelens_powerof_bad'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (codelens_powerof_bad_vis) {
        codelens_powerof_bad_vis.redrawConnectors();
    }
});
</script>
</div>
<p>Now step through the code.  What do you notice about the values of variable <code class="docutils literal"><span class="pre">power</span></code>
in the local scope compared to the variable <code class="docutils literal"><span class="pre">power</span></code> in the global scope?</p>
<p>The value of <code class="docutils literal"><span class="pre">power</span></code> in the local scope was different than the global scope.
That is because in this example <code class="docutils literal"><span class="pre">power</span></code> was used on the left hand side of the
assignment statement <code class="docutils literal"><span class="pre">power</span> <span class="pre">=</span> <span class="pre">p</span></code>.  When a variable name is used on the
left hand side of an assignment statement Python creates a local variable.
When a local variable has the same name as a global variable we say that the
local shadows the global.  A <strong>shadow</strong> means that the global variable cannot
be accessed by Python because the local variable will be found first. This is
another good reason not to use global variables. As you can see,
it makes your code confusing and difficult to
understand.</p>
<p>To cement all of these ideas even further lets look at one final example.
Inside the <code class="docutils literal"><span class="pre">square</span></code> function we are going to make an assignment to the
parameter <code class="docutils literal"><span class="pre">original_number</span></code>  <strong>There’s no good reason to do this</strong> other than to emphasize
the fact that the parameter <code class="docutils literal"><span class="pre">original_number</span></code> is a local variable.  If you step through
the example in codelens you will see that although <code class="docutils literal"><span class="pre">original_number</span></code> is 0 in the local
variables for <code class="docutils literal"><span class="pre">square</span></code>, the <code class="docutils literal"><span class="pre">original_number</span></code> in the global scope remains 2.  This is confusing
to many beginning programmers who think that an assignment to a
formal parameter will cause a change to the value of the variable that was
used as the actual parameter, especially when the two share the same name.
But this example demonstrates that that is clearly not how Python operates.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="codelens_change_parm"></div>
<p class="cl_caption"><span class="cl_caption_text"> (codelens_change_parm)</span> </p>
</div>
<script type="text/javascript">
var codelens_change_parm_trace = {"code": "def square(original_number):\n    squared_value = original_number * original_number\n\n    # assign a new value to the parameter original_number\n    original_number = 0\n\n    return squared_value\n\noriginal_number = 2\nresult = square(original_number)\nprint(result)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 9, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1]}, "ordered_globals": ["square"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 10, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "square", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 2}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 2}, "ordered_varnames": ["original_number"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 2, "squared_value": 4}, "ordered_varnames": ["original_number", "squared_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "square", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 0, "squared_value": 4}, "ordered_varnames": ["original_number", "squared_value"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 7, "event": "return", "func_name": "square", "globals": {"square": ["REF", 1], "original_number": 2}, "ordered_globals": ["square", "original_number"], "stack_to_render": [{"func_name": "square", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_number": 0, "squared_value": 4, "__return__": 4}, "ordered_varnames": ["original_number", "squared_value", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "square_f1"}], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 11, "event": "step_line", "func_name": "<module>", "globals": {"square": ["REF", 1], "original_number": 2, "result": 4}, "ordered_globals": ["square", "original_number", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": ""}, {"line": 11, "event": "return", "func_name": "<module>", "globals": {"square": ["REF", 1], "original_number": 2, "result": 4}, "ordered_globals": ["square", "original_number", "result"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "square(original_number)", null]}, "stdout": "4\n"}]};
var codelens_change_parm_vis;

$(document).ready(function() {
    try {
        codelens_change_parm_vis = new ExecutionVisualizer('codelens_change_parm',codelens_change_parm_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(codelens_change_parm_vis,'codelens_change_parm');
        styleButtons('codelens_change_parm');
        allVisualizers.push(codelens_change_parm_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component codelens_change_parm_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#codelens_change_parm_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'codelens_change_parm'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (codelens_change_parm_vis) {
        codelens_change_parm_vis.redrawConnectors();
    }
});
</script>
</div>
<div class="section" id="id1">
<h3>1.3.1. Check Your Understanding<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question5_3_1">
            variable-scope3: What is a variable’s scope?
            <li data-component="answer"  id="test_question5_3_1_opt_a">Its value</li><li data-component="feedback" id="test_question5_3_1_opt_a">Value is the contents of the variable.  Scope concerns where the variable is &quot;known&quot;.</li>
            
            <li data-component="answer" data-correct id="test_question5_3_1_opt_b">The area in the code where a variable can be accessed.</li><li data-component="feedback" id="test_question5_3_1_opt_b"></li>
            
            <li data-component="answer"  id="test_question5_3_1_opt_c">Its name</li><li data-component="feedback" id="test_question5_3_1_opt_c">The name of a variable is just an identifier or alias.  Scope concerns where the variable is &quot;known&quot;.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question5_3_2">
            variable-scope4: What is a local variable?
            <li data-component="answer" data-correct id="test_question5_3_2_opt_a">A temporary variable that is only used inside a function</li><li data-component="feedback" id="test_question5_3_2_opt_a">Yes, a local variable is a temporary variable that is only known (only exists) in the function it is defined in.</li>
            
            <li data-component="answer"  id="test_question5_3_2_opt_b">The same as a parameter</li><li data-component="feedback" id="test_question5_3_2_opt_b">While parameters may be considered local variables, functions may also define and use additional local variables.</li>
            
            <li data-component="answer"  id="test_question5_3_2_opt_c">Another name for any variable</li><li data-component="feedback" id="test_question5_3_2_opt_c">Variables that are used outside a function are not local, but rather global variables.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question5_3_3">
            variable-scope5: Can you use the same name for a local variable as a global variable?
            <li data-component="answer"  id="test_question5_3_3_opt_a">Yes, and there is no reason not to.</li><li data-component="feedback" id="test_question5_3_3_opt_a">While there is no problem as far as Python is concerned, it is generally considered bad style because of the potential for the programmer to get confused.</li>
            
            <li data-component="answer" data-correct id="test_question5_3_3_opt_b">Yes, but it is considered bad form.</li><li data-component="feedback" id="test_question5_3_3_opt_b">It is generally considered bad style because of the potential for the programmer to get confused.  If you must use global variables (also generally bad form) make sure they have unique names.</li>
            
            <li data-component="answer"  id="test_question5_3_3_opt_c">No, it will cause an error.</li><li data-component="feedback" id="test_question5_3_3_opt_c">Python manages global and local scope separately and has clear rules for how to handle variables with the same name in different scopes, so this will not cause a Python error.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="practice-problems">
<h2>1.4. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems to be sure you understand how to create fruitful functions. Your functions <strong>have to return</strong> the correct value – using <code class="docutils literal"><span class="pre">print()</span></code> will not work. When you run your code for these questions, your code will automatically be checked with a number of test cases to see if your function works in all situations. You will be able to see any situations in which your function does not provide the correct answer.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for the following is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>
<div class="section" id="double-it">
<h3>1.4.1. Double It<a class="headerlink" href="#double-it" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="double-it" class="runestone explainer ac_section alert alert-warning">
<p>The parameter <code class="docutils literal"><span class="pre">the_number</span></code> needs to be doubled. Return the doubled value of the number that is passed in.</p>
<p><code class="docutils literal"><span class="pre">double_it(5)</span> <span class="pre">→</span> <span class="pre">10</span></code></p>
<p><code class="docutils literal"><span class="pre">double_it(0)</span> <span class="pre">→</span> <span class="pre">0</span></code></p>
<p><code class="docutils literal"><span class="pre">double_it(-4)</span> <span class="pre">→</span> <span class="pre">-8</span></code></p>

<textarea data-component="activecode" id=double-it data-lang="python"    data-timelimit=25000   data-audio=''     >
def double_it(the_number):
    return 0

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(double_it(5),10,"double_it(5)")
        self.assertEqual(double_it(0),0,"double_it(0)")
        self.assertEqual(double_it(-12),-24,"double_it(-12)")
        self.assertEqual(double_it(-1),-2,"double_it(-1)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="double-it-positive">
<h3>1.4.2. Double It Positive<a class="headerlink" href="#double-it-positive" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="double-it-positive" class="runestone explainer ac_section alert alert-warning">
<p>The parameter <code class="docutils literal"><span class="pre">the_number</span></code> needs to be doubled, but only if <code class="docutils literal"><span class="pre">the_number</span></code> is positive. Return the doubled value of the number that is passed in if <code class="docutils literal"><span class="pre">the_number</span></code> is positive. If <code class="docutils literal"><span class="pre">the_number</span></code> is negative, return -1. If <code class="docutils literal"><span class="pre">the_number</span></code> is 0, return 0.</p>
<p><code class="docutils literal"><span class="pre">double_it_positive(5)</span> <span class="pre">→</span> <span class="pre">10</span></code></p>
<p><code class="docutils literal"><span class="pre">double_it_positive(0)</span> <span class="pre">→</span> <span class="pre">0</span></code></p>
<p><code class="docutils literal"><span class="pre">double_it_positive(-4)</span> <span class="pre">→</span> <span class="pre">-1</span></code></p>

<textarea data-component="activecode" id=double-it-positive data-lang="python"    data-timelimit=25000   data-audio=''     >
def double_it_positive(the_number):
    return 0

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(double_it_positive(5),10,"double_it_positive(5)")
        self.assertEqual(double_it_positive(11),22,"double_it_positive(11)")
        self.assertEqual(double_it_positive(0),0,"double_it_positive(0)")
        self.assertEqual(double_it_positive(-12),-1,"double_it_positive(-12)")
        self.assertEqual(double_it_positive(-1),-1,"double_it_positive(-1)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="wear-the-right-thing">
<h3>1.4.3. Wear the Right Thing<a class="headerlink" href="#wear-the-right-thing" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="wear-the-right-thing" class="runestone explainer ac_section alert alert-warning">
<p>Write a function which returns the type of clothing you should wear, based on the parameter <code class="docutils literal"><span class="pre">temperature</span></code>. If the <code class="docutils literal"><span class="pre">temperature</span></code> is below -10, you will wear a parka and toque (return “parka and toque”). If the <code class="docutils literal"><span class="pre">temperature</span></code> is between -10 and 0, wear a toque (return “toque”). If the <code class="docutils literal"><span class="pre">temperature</span></code> is greater than 0 but less than 10, wear a sweater (return “sweater”). If the <code class="docutils literal"><span class="pre">temperature</span></code> is between 10 and 20, wear a t-shirt (return “t-shirt”). If the <code class="docutils literal"><span class="pre">temperature</span></code> is greater than 20, wear shorts (return “shorts”).</p>
<p><code class="docutils literal"><span class="pre">wear_the_right_thing(25)</span> <span class="pre">→</span> <span class="pre">&quot;shorts&quot;</span></code></p>
<p><code class="docutils literal"><span class="pre">wear_the_right_thing(-25)</span> <span class="pre">→</span> <span class="pre">&quot;parka</span> <span class="pre">and</span> <span class="pre">toque&quot;</span></code></p>
<p><code class="docutils literal"><span class="pre">wear_the_right_thing(-5)</span> <span class="pre">→</span> <span class="pre">&quot;toque&quot;</span></code></p>

<textarea data-component="activecode" id=wear-the-right-thing data-lang="python"    data-timelimit=25000   data-audio=''     >
def wear_the_right_thing(temperature):
    return "some clothes"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(wear_the_right_thing(25),"shorts","wear_the_right_thing(25)")
        self.assertEqual(wear_the_right_thing(20),"t-shirt","wear_the_right_thing(20)")
        self.assertEqual(wear_the_right_thing(15),"t-shirt","wear_the_right_thing(15)")
        self.assertEqual(wear_the_right_thing(10),"t-shirt","wear_the_right_thing(10)")
        self.assertEqual(wear_the_right_thing(9),"sweater","wear_the_right_thing(9)")
        self.assertEqual(wear_the_right_thing(1),"sweater","wear_the_right_thing(1)")
        self.assertEqual(wear_the_right_thing(0),"toque","wear_the_right_thing(0)")
        self.assertEqual(wear_the_right_thing(-10),"toque","wear_the_right_thing(-10)")
        self.assertEqual(wear_the_right_thing(-11),"parka and toque","wear_the_right_thing(-11)")
        self.assertEqual(wear_the_right_thing(-30),"parka and toque","wear_the_right_thing(-30)")

myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 5. Loops Practice Quiz and Assignment Work Period' data-toggle="tooltip">
    <a href="../MoreTurtles/PracticeQuizAndWorkPeriod.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../MoreTurtles/PracticeQuizAndWorkPeriod.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 2. Writing Fruitful Functions Practice' data-toggle="tooltip" >
    <a href="WritingFruitfulFunctionsPractice.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="WritingFruitfulFunctionsPractice.html" >Next Section - 2. Writing Fruitful Functions Practice</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>