<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>2. Writing Fruitful Functions Practice &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.min.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-3.4.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/presenter_mode.css?v=9FF1320F" />
    <link rel="stylesheet" type="text/css" href="../../_static/jquery-ui-1.10.3.custom.min.css?v=97E0178A" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/runestone-custom-sphinx-bootstrap.css?v=5.6.1" />
    <link rel="stylesheet" type="text/css" href="../../_static/theme-overrides.css?v=2DAFE138" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/&#34;&#34;/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script id="documentation_options" data-url_root="../../" src="../../_static/documentation_options.js"></script>
    <script src="../../_static/runestonebase.js"></script>
    <script src="../../_static/jquery.js"></script>
    <script src="../../_static/underscore.js"></script>
    <script src="../../_static/doctools.js"></script>
    <script src="../../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script src="../../_static/sql-wasm.js?v=D71B5933"></script>
    <script src="https://cdn.jsdelivr.net/npm/handsontable@7.2.2/dist/handsontable.full.js"></script>
    <script src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script src="../../_static/html4css1.css?v=4533ACD3"></script>
    <script src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script src="../../_static/Blob.js?v=AEC2447"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.bidi.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.emitter.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.fallbacks.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.messagestore.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.parser.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jquery.i18n/1.0.5/jquery.i18n.language.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script src="../../_static/runestone.js?v=5A173BAB"></script>
    <script src="../../_static/jquery-ui-1.10.3.custom.min.js?v=2B07FE4F"></script>
    <script src="../../_static/bootstrap-3.4.1/js/bootstrap.min.js"></script>
    <script src="../../_static/jquery-fix.js?v=1226591B"></script>
    <script src="../../_static/bootstrap-sphinx.js"></script>
    <script src="../../_static/jquery.idle-timer.js?v=E942DAF1"></script>
    <script src="../../_static/presenter_mode.js?v=2DF16FB4"></script>
    <script src="../../_static/theme.js?v=D899FB97"></script>
    <script src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../../_static/jquery-fix.js"></script>
    <script src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../../_static/bootstrap-sphinx.js"></script>
    <script src="../../_static/waypoints.min.js"></script>
    <script src="../../_static/rangy-core.js"></script>
    <script src="../../_static/rangy-textrange.js"></script>
    <script src="../../_static/rangy-cssclassapplier.js"></script>
    <script src="../../_static/user-highlights.js"></script>
    <script src="../../_static/jquery.idle-timer.js"></script>
    <script src="../../_static/processing-1.4.1.min.js"></script>
    <script src="../../_static/jquery.hotkey.js"></script>
    <script src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="3. The Accumulator Pattern" href="TheAccumulatorPattern.html" />
    <link rel="prev" title="1. Fruitful Functions and Variable Scope" href="MoreFunctionDetails.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '5.6.1';
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">ActiveCode Work Area</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">2. Writing Fruitful Functions Practice</a><ul>
<li><a class="reference internal" href="#warmup-problem">2.1. Warmup Problem</a></li>
<li><a class="reference internal" href="#what-does-this-program-do">2.2. What Does This Program Do?</a></li>
<li><a class="reference internal" href="#practice-problems">2.3. Practice Problems</a><ul>
<li><a class="reference internal" href="#area-of-rectangle">2.3.1. Area of rectangle</a></li>
<li><a class="reference internal" href="#letter-grade">2.3.2. Letter Grade</a></li>
<li><a class="reference internal" href="#find-the-smallest">2.3.3. Find the Smallest</a></li>
<li><a class="reference internal" href="#is-even">2.3.4. Is Even</a></li>
<li><a class="reference internal" href="#leap-year">2.3.5. Leap Year</a></li>
<li><a class="reference internal" href="#using-your-is-even-function">2.3.6. Using Your Is Even Function</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">1. </span>Fruitful Functions and Variable Scope" data-toggle="tooltip">
    <a href="MoreFunctionDetails.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">3. </span>The Accumulator Pattern' data-toggle="tooltip" >
    <a href="TheAccumulatorPattern.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="writing-fruitful-functions-practice">
<h1><span class="section-number">2. </span>Writing Fruitful Functions Practice<a class="headerlink" href="#writing-fruitful-functions-practice" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title">Quick Overview of Day</p>
<p>Warmup drawing problem. WDTPD questions about functions. Students practice writing fruitful functions.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed_fruitful_functions"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide"  style="visibility: hidden;">
    <ul class="simple">
<li><p><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</p></li>
<li><p><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</p></li>
<li><p><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</p></li>
<li><p><strong>CS20-FP3</strong> Construct and utilize functions to create reusable pieces of code.</p></li>
</ul>

    </div>
    <div class="section" id="warmup-problem">
<h2><span class="section-number">2.1. </span>Warmup Problem<a class="headerlink" href="#warmup-problem" title="Permalink to this headline">¶</a></h2>
<img alt="../../_images/confusion-25.png" src="../../_images/confusion-25.png" />
<p>Draw the image above using the Python turtle module. <strong>You must define a function as part of your solution!</strong></p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=turtle_confusion_practice_25 data-question_label="2.1.1">
<div id=turtle_confusion_practice_25_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="turtle_confusion_practice_25_editor" 
      data-timelimit=25000     
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
import turtle
</textarea>
</div>
</div>
</div>
<div class="section" id="what-does-this-program-do">
<h2><span class="section-number">2.2. </span>What Does This Program Do?<a class="headerlink" href="#what-does-this-program-do" title="Permalink to this headline">¶</a></h2>
<p>Remember that variables created inside of a function have <strong>local scope</strong> (can only be used inside that function), whereas variables created outside of a function have <strong>global scope</strong> (can be accessed from anywhere).</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>Your teacher may choose to use the following examples as a class activity, by displaying the  examples, and having you take a guess as to what you think each will do before running the code.</p>
</div>
<p>What will the following programs output? Why?</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=wdtpd_functions_5 data-question_label="2.2.1">
<div id=wdtpd_functions_5_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="wdtpd_functions_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='What will this program print?'  data-wasm=/_static 
     style="visibility: hidden;">
def foo():
    print(a)

a = 10
foo()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=wdtpd_functions_6 data-question_label="2.2.2">
<div id=wdtpd_functions_6_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="wdtpd_functions_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='What will this program print?'  data-wasm=/_static 
     style="visibility: hidden;">
def foo():
    a = a + 1
    print(a)

a = 10
foo()
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=wdtpd_functions_7 data-question_label="2.2.3">
<div id=wdtpd_functions_7_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="wdtpd_functions_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='What will this program print?'  data-wasm=/_static 
     style="visibility: hidden;">
def doubler(some_number):
    new_value = some_number * 2
    return new_value

initial_value = 15
new_value = doubler(15)
print(new_value)
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=wdtpd_functions_8 data-question_label="2.2.4">
<div id=wdtpd_functions_8_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="wdtpd_functions_8_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='What will this program print?'  data-wasm=/_static 
     style="visibility: hidden;">
def greeting(first_name, last_name):
    message = "Hello there, " + first_name + " " + last_name + "!"
    return message

first = "Sherlock"
last = "Holmes"
print(greeting(first, last))
</textarea>
</div>
</div>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=wdtpd_functions_9 data-question_label="2.2.5">
<div id=wdtpd_functions_9_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="wdtpd_functions_9_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
         data-caption='What will this program print?'  data-wasm=/_static 
     style="visibility: hidden;">
def is_big(number):
    if number > 42:
        return True
    else:
        return False

my_number = 50
result = is_big(my_number)
print(result)
</textarea>
</div>
</div>
</div>
<div class="section" id="practice-problems">
<h2><span class="section-number">2.3. </span>Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems to be sure you understand how to create fruitful functions. Your functions <strong>have</strong> to return the correct value – using <code class="docutils literal notranslate"><span class="pre">print()</span></code> will not work. When you run your code for these questions, your code will automatically be checked with a number of test cases to see if your function works in all situations. You will be able to see any situations in which your function does not provide the correct answer.</p>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for the following is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>
<div class="section" id="area-of-rectangle">
<h3><span class="section-number">2.3.1. </span>Area of rectangle<a class="headerlink" href="#area-of-rectangle" title="Permalink to this headline">¶</a></h3>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=rectangle-area data-question_label="2.3.1.1">
<div id=rectangle-area_question class="ac_question col-md-12">
<p>The parameters <code class="docutils literal notranslate"><span class="pre">length</span></code> and <code class="docutils literal notranslate"><span class="pre">width</span></code> represent the lengths of the sides of a rectangle. Calculate the area of the rectangle with the given values, and return the result.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">rectangle_area(5,</span> <span class="pre">10)</span> <span class="pre">→</span> <span class="pre">50</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">rectangle_area(1,</span> <span class="pre">10)</span> <span class="pre">→</span> <span class="pre">10</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">rectangle_area(2,</span> <span class="pre">6)</span> <span class="pre">→</span> <span class="pre">12</span></code></p>

</div>
<textarea data-lang="python" id="rectangle-area_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def rectangle_area(length, width):
    return 0

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(rectangle_area(5, 10),50,"rectangle_area(5, 10)")
        self.assertEqual(rectangle_area(1, 10),10,"rectangle_area(1, 10)")
        self.assertEqual(rectangle_area(2, 6),12,"rectangle_area(2, 6)")
        self.assertEqual(rectangle_area(1, 1),1,"rectangle_area(1, 1)")
        self.assertEqual(rectangle_area(3, 5),15,"rectangle_area(3, 5)")
        self.assertEqual(rectangle_area(15, 4),60,"rectangle_area(15, 4)")

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="letter-grade">
<h3><span class="section-number">2.3.2. </span>Letter Grade<a class="headerlink" href="#letter-grade" title="Permalink to this headline">¶</a></h3>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=letter-grade data-question_label="2.3.2.1">
<div id=letter-grade_question class="ac_question col-md-12">
<p>Write a function that returns the letter grade, given an exam mark as the parameter. The grading scheme is:</p>
<table class="docutils align-default">
<colgroup>
<col style="width: 45%" />
<col style="width: 55%" />
</colgroup>
<thead>
<tr class="row-odd"><th class="head"><p>Letter Grade</p></th>
<th class="head"><p>Interval</p></th>
</tr>
</thead>
<tbody>
<tr class="row-even"><td><p>A</p></td>
<td><p>&gt;= 90</p></td>
</tr>
<tr class="row-odd"><td><p>B</p></td>
<td><p>[80, 90)</p></td>
</tr>
<tr class="row-even"><td><p>C</p></td>
<td><p>[70, 80)</p></td>
</tr>
<tr class="row-odd"><td><p>D</p></td>
<td><p>[60, 70)</p></td>
</tr>
<tr class="row-even"><td><p>F</p></td>
<td><p>&lt; 60</p></td>
</tr>
</tbody>
</table>
<p>The square and round brackets denote closed and open intervals. A closed interval includes the number, and open interval excludes it. So 79.99999 gets grade C, but 80 gets grade B.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_grade(83)</span> <span class="pre">→</span> <span class="pre">&quot;B&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_grade(73)</span> <span class="pre">→</span> <span class="pre">&quot;C&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_grade(80)</span> <span class="pre">→</span> <span class="pre">&quot;B&quot;</span></code></p>

</div>
<textarea data-lang="python" id="letter-grade_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def letter_grade(exam_grade):
    return "A"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(letter_grade(83),"B","letter_grade(83)")
        self.assertEqual(letter_grade(75),"C","letter_grade(75)")
        self.assertEqual(letter_grade(80),"B","letter_grade(80)")
        self.assertEqual(letter_grade(90),"A","letter_grade(90)")
        self.assertEqual(letter_grade(89.9),"B","letter_grade(89.9)")
        self.assertEqual(letter_grade(55),"F","letter_grade(55)")
        self.assertEqual(letter_grade(45),"F","letter_grade(45)")
        self.assertEqual(letter_grade(69.7),"D","letter_grade(69.7)")

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="find-the-smallest">
<h3><span class="section-number">2.3.3. </span>Find the Smallest<a class="headerlink" href="#find-the-smallest" title="Permalink to this headline">¶</a></h3>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=find-min data-question_label="2.3.3.1">
<div id=find-min_question class="ac_question col-md-12">
<p>The function find_min(a, b, c) will take three numbers as parameters and return the <em>smallest</em> value. If more than one number is tied for the smallest, still return that smallest number. Note that you cannot use the <code class="docutils literal notranslate"><span class="pre">min</span></code> function in this solution.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">find_min(4,</span> <span class="pre">7,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">4</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">find_min(4,</span> <span class="pre">5,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">4</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">find_min(4,</span> <span class="pre">-7,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">-7</span></code></p>

</div>
<textarea data-lang="python" id="find-min_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def find_min(a, b, c):
    return a

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(find_min(4, 7, 5),4,"find_min(4, 7, 5)")
        self.assertEqual(find_min(4, 5, 5),4,"find_min(4, 5, 5)")
        self.assertEqual(find_min(4, 4, 4),4,"find_min(4, 4, 4)")
        self.assertEqual(find_min(-2, -6, -100),-100,"find_min(-2, -6, -100)")
        self.assertEqual(find_min(4, -7, 5),-7,"find_min(4, -7, 5)")
        self.assertEqual(find_min(-4, 7, -5),-5,"find_min(-4, 7, -5)")

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="is-even">
<span id="is-even-fruitful-function"></span><h3><span class="section-number">2.3.4. </span>Is Even<a class="headerlink" href="#is-even" title="Permalink to this headline">¶</a></h3>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=is-even data-question_label="2.3.4.1">
<div id=is-even_question class="ac_question col-md-12">
<p>The function is_even(number) will return <code class="docutils literal notranslate"><span class="pre">True</span></code> if the number passed in is even, and <code class="docutils literal notranslate"><span class="pre">False</span></code> if it is odd. Hint: You might want to look back at the <a class="reference internal" href="../Overview/QuickOverview.html#math-operator-list"><span class="std std-ref">Math Operators</span></a> list.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">is_even(4)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">is_even(-4)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">is_even(5)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

</div>
<textarea data-lang="python" id="is-even_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def is_even(number):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_even(4),True,"is_even(4)")
        self.assertEqual(is_even(-4),True,"is_even(-4)")
        self.assertEqual(is_even(5),False,"is_even(5)")
        self.assertEqual(is_even(-5),False,"is_even(-5)")
        self.assertEqual(is_even(0),True,"is_even(0)")
        self.assertEqual(is_even(513),False,"is_even(513)")

myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="leap-year">
<h3><span class="section-number">2.3.5. </span>Leap Year<a class="headerlink" href="#leap-year" title="Permalink to this headline">¶</a></h3>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=leap-year data-question_label="2.3.5.1">
<div id=leap-year_question class="ac_question col-md-12">
<p>A year is a leap year if it is divisible by 4 unless it is a century that is not divisible by 400. Write a function that takes a year as a parameter and returns True if the year is a leap year, False otherwise. The following pseudocode determines whether a year is a leap year or a common year in the Gregorian calendar (from <a class="reference external" href="http://en.wikipedia.org/wiki/Leap_year#Algorithm">Wikipedia</a>):</p>
<blockquote>
<div><div class="line-block">
<div class="line">if (year is not divisible by 4) then (it is a common year)</div>
<div class="line">else if (year is not divisible by 100) then (it is a leap year)</div>
<div class="line">else if (year is not divisible by 400) then (it is a common year)</div>
<div class="line">else (it is a leap year)</div>
</div>
</div></blockquote>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">leap_year(2001)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">leap_year(2020)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">leap_year(1900)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

</div>
<textarea data-lang="python" id="leap-year_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
def leap_year(year):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(leap_year(2001),False,"leap_year(2001)")
        self.assertEqual(leap_year(2020),True,"leap_year(2020)")
        self.assertEqual(leap_year(1900),False,"leap_year(1900)")
        self.assertEqual(leap_year(2004),True,"leap_year(2004)")
        self.assertEqual(leap_year(2000),True,"leap_year(2000)")


myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="using-your-is-even-function">
<h3><span class="section-number">2.3.6. </span>Using Your Is Even Function<a class="headerlink" href="#using-your-is-even-function" title="Permalink to this headline">¶</a></h3>
<p>Write a program that continues to take in a number from the user until the number given is NOT even. For example, the user might enter 4, 10, 42, 5. The program would only stop when the non-even number 5 is entered. You need to use the <code class="docutils literal notranslate"><span class="pre">is_even</span></code> function you defined above.</p>

<div class="runestone explainer ac_section alert alert-warning">
<div data-component="activecode" id=using-is-even data-question_label="2.3.6.1">
<div id=using-is-even_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="using-is-even_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static 
     style="visibility: hidden;">
# your code here!
</textarea>
</div>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - <span class="section-number">1. </span>Fruitful Functions and Variable Scope' data-toggle="tooltip">
    <a href="MoreFunctionDetails.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="MoreFunctionDetails.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - <span class="section-number">3. </span>The Accumulator Pattern' data-toggle="tooltip" >
    <a href="TheAccumulatorPattern.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="TheAccumulatorPattern.html" >Next Section - <span class="section-number">3. </span>The Accumulator Pattern</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.
    </p>
  </div>
</footer>


  </body>
</html>