<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>18. Writing Fruitful Functions Practice &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="19. Reaction Game with Micro:bit" href="MicrobitReactionGame.html" />
    <link rel="prev" title="17. More Details About Functions" href="MoreFunctionDetails.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">18. Writing Fruitful Functions Practice</a><ul>
<li><a class="reference internal" href="#warmup-problem">18.1. Warmup Problem</a></li>
<li><a class="reference internal" href="#what-does-this-program-do">18.2. What Does This Program Do?</a></li>
<li><a class="reference internal" href="#practice-problems">18.3. Practice Problems</a><ul>
<li><a class="reference internal" href="#area-of-rectangle">18.3.1. Area of rectangle</a></li>
<li><a class="reference internal" href="#letter-grade">18.3.2. Letter Grade</a></li>
<li><a class="reference internal" href="#find-the-smallest">18.3.3. Find the Smallest</a></li>
<li><a class="reference internal" href="#is-even">18.3.4. Is Even</a></li>
<li><a class="reference internal" href="#leap-year">18.3.5. Leap Year</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 17. More Details About Functions" data-toggle="tooltip">
    <a href="MoreFunctionDetails.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 19. Reaction Game with Micro:bit' data-toggle="tooltip" >
    <a href="MicrobitReactionGame.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="writing-fruitful-functions-practice">
<h1>18. Writing Fruitful Functions Practice<a class="headerlink" href="#writing-fruitful-functions-practice" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>Warmup drawing problem. WDTPD questions about functions. Students practice writing fruitful functions.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed"   data-showtitle="Curriculum Objectives Addressed In This Section" data-hidetitle="Hide">
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to encapsulate reusable pieces of code.</li>
</ul>

    </div>
    <div class="section" id="warmup-problem">
<h2>18.1. Warmup Problem<a class="headerlink" href="#warmup-problem" title="Permalink to this headline">¶</a></h2>
<img alt="../../_images/confusion-25.png" src="../../_images/confusion-25.png" />
<p>Draw the image above using the Python turtle module. <strong>You must define a function as part of your solution!</strong></p>

<div data-childcomponent="turtle_confusion_practice_25" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=turtle_confusion_practice_25 data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle
</textarea>
</div>
</div>
<div class="section" id="what-does-this-program-do">
<h2>18.2. What Does This Program Do?<a class="headerlink" href="#what-does-this-program-do" title="Permalink to this headline">¶</a></h2>
<p>Remember that variables created inside of a function have <strong>local scope</strong> (can only be used inside that function), whereas variables created outside of a function have <strong>global scope</strong> (can be accessed from anywhere).</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">Your teacher may choose to use the following examples as a class activity, by displaying the  examples, and having you take a guess as to what you think each will do before running the code.</p>
</div>
<p>What will the following programs output? Why?</p>

<div data-childcomponent="wdtpd_functions_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def a():
    print("A")

def b():
    print("B")

def c():
    print("C")

a()
</textarea>
</div>

<div data-childcomponent="wdtpd_functions_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def a():
    b()
    print("A")

def b():
    c()
    print("B")

def c():
    print("C")

a()
</textarea>
</div>

<div data-childcomponent="wdtpd_functions_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_3 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def a():
    print("A")
    b()

def b():
    print("B")
    c()

def c():
    print("C")

a()
</textarea>
</div>

<div data-childcomponent="wdtpd_functions_4" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_4 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def a():
    print("A start")
    b()
    print("A end")

def b():
    print("B start")
    c()
    print("B end")

def c():
    print("C start and end")

a()
</textarea>
</div>

<div data-childcomponent="wdtpd_functions_5" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_5 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def foo():
    print(a)

a = 10
foo()
</textarea>
</div>

<div data-childcomponent="wdtpd_functions_6" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_functions_6 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     data-caption='What will this program print?'>
def foo():
    a = a + 1
    print(a)

a = 10
foo()
</textarea>
</div>
</div>
<div class="section" id="practice-problems">
<h2>18.3. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems to be sure you understand how to create fruitful functions. Your functions <strong>have</strong> to return the correct value – using <code class="docutils literal"><span class="pre">print()</span></code> will not work. When you run your code for these questions, your code will automatically be checked with a number of test cases to see if your function works in all situations. You will be able to see any situations in which your function does not provide the correct answer.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for the following is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>
<div class="section" id="area-of-rectangle">
<h3>18.3.1. Area of rectangle<a class="headerlink" href="#area-of-rectangle" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="rectangle-area" class="runestone explainer ac_section alert alert-warning">
<p>The parameters <code class="docutils literal"><span class="pre">length</span></code> and <code class="docutils literal"><span class="pre">width</span></code> represent the lengths of the sides of a rectangle. Calculate the area of the rectangle with the given values, and return the result.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">rectangle_area(5,</span> <span class="pre">10)</span> <span class="pre">→</span> <span class="pre">50</span></code></p>
<p><code class="docutils literal"><span class="pre">rectangle_area(1,</span> <span class="pre">10)</span> <span class="pre">→</span> <span class="pre">10</span></code></p>
<p><code class="docutils literal"><span class="pre">rectangle_area(2,</span> <span class="pre">6)</span> <span class="pre">→</span> <span class="pre">12</span></code></p>

<textarea data-component="activecode" id=rectangle-area data-lang="python"    data-timelimit=25000   data-audio=''     >
def rectangle_area(length, width):
    return 0

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(rectangle_area(5, 10),50,"rectangle_area(5, 10)")
        self.assertEqual(rectangle_area(1, 10),10,"rectangle_area(1, 10)")
        self.assertEqual(rectangle_area(2, 6),12,"rectangle_area(2, 6)")
        self.assertEqual(rectangle_area(1, 1),1,"rectangle_area(1, 1)")
        self.assertEqual(rectangle_area(3, 5),15,"rectangle_area(3, 5)")
        self.assertEqual(rectangle_area(15, 4),60,"rectangle_area(15, 4)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="letter-grade">
<h3>18.3.2. Letter Grade<a class="headerlink" href="#letter-grade" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="letter-grade" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that returns the letter grade, given an exam mark as the parameter. The grading scheme is:</p>
<table border="1" class="docutils">
<colgroup>
<col width="45%" />
<col width="55%" />
</colgroup>
<thead valign="bottom">
<tr class="row-odd"><th class="head">Letter Grade</th>
<th class="head">Interval</th>
</tr>
</thead>
<tbody valign="top">
<tr class="row-even"><td>A</td>
<td>&gt;= 90</td>
</tr>
<tr class="row-odd"><td>B</td>
<td>[80, 90)</td>
</tr>
<tr class="row-even"><td>C</td>
<td>[70, 80)</td>
</tr>
<tr class="row-odd"><td>D</td>
<td>[60, 70)</td>
</tr>
<tr class="row-even"><td>F</td>
<td>&lt; 60</td>
</tr>
</tbody>
</table>
<p>The square and round brackets denote closed and open intervals. A closed interval includes the number, and open interval excludes it. So 79.99999 gets grade C, but 80 gets grade B.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">letter_grade(83)</span> <span class="pre">→</span> <span class="pre">&quot;B&quot;</span></code></p>
<p><code class="docutils literal"><span class="pre">letter_grade(73)</span> <span class="pre">→</span> <span class="pre">&quot;C&quot;</span></code></p>
<p><code class="docutils literal"><span class="pre">letter_grade(80)</span> <span class="pre">→</span> <span class="pre">&quot;B&quot;</span></code></p>

<textarea data-component="activecode" id=letter-grade data-lang="python"    data-timelimit=25000   data-audio=''     >
def letter_grade(exam_grade):
    return "A"

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(letter_grade(83),"B","letter_grade(83)")
        self.assertEqual(letter_grade(75),"C","letter_grade(75)")
        self.assertEqual(letter_grade(80),"B","letter_grade(80)")
        self.assertEqual(letter_grade(90),"A","letter_grade(90)")
        self.assertEqual(letter_grade(89.9),"B","letter_grade(89.9)")
        self.assertEqual(letter_grade(55),"F","letter_grade(55)")
        self.assertEqual(letter_grade(45),"F","letter_grade(45)")
        self.assertEqual(letter_grade(69.7),"D","letter_grade(69.7)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="find-the-smallest">
<h3>18.3.3. Find the Smallest<a class="headerlink" href="#find-the-smallest" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="find-min" class="runestone explainer ac_section alert alert-warning">
<p>The function find_min(a, b, c) will take three numbers as parameters and return the <em>smallest</em> value. If more than one number is tied for the smallest, still return that smallest number. Note that you cannot use the <code class="docutils literal"><span class="pre">min</span></code> function in this solution.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">find_min(4,</span> <span class="pre">7,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">4</span></code></p>
<p><code class="docutils literal"><span class="pre">find_min(4,</span> <span class="pre">5,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">4</span></code></p>
<p><code class="docutils literal"><span class="pre">find_min(4,</span> <span class="pre">-7,</span> <span class="pre">5)</span> <span class="pre">→</span> <span class="pre">-7</span></code></p>

<textarea data-component="activecode" id=find-min data-lang="python"    data-timelimit=25000   data-audio=''     >
def find_min(a, b, c):
    return a

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(find_min(4, 7, 5),4,"find_min(4, 7, 5)")
        self.assertEqual(find_min(4, 5, 5),4,"find_min(4, 5, 5)")
        self.assertEqual(find_min(4, 4, 4),4,"find_min(4, 4, 4)")
        self.assertEqual(find_min(-2, -6, -100),-100,"find_min(-2, -6, -100)")
        self.assertEqual(find_min(4, -7, 5),-7,"find_min(4, -7, 5)")
        self.assertEqual(find_min(-4, 7, -5),-5,"find_min(-4, 7, -5)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="is-even">
<h3>18.3.4. Is Even<a class="headerlink" href="#is-even" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="is-even" class="runestone explainer ac_section alert alert-warning">
<p>The function is_even(number) will return <code class="docutils literal"><span class="pre">True</span></code> if the number passed in is even, and <code class="docutils literal"><span class="pre">False</span></code> if it is odd. Hint: You might want to look back at the <a class="reference internal" href="../Overview/QuickOverview.html#math-operator-list"><span class="std std-ref">Math Operators</span></a> list.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">is_even(4)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">is_even(-4)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">is_even(5)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

<textarea data-component="activecode" id=is-even data-lang="python"    data-timelimit=25000   data-audio=''     >
def is_even(number):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_even(4),True,"is_even(4)")
        self.assertEqual(is_even(-4),True,"is_even(-4)")
        self.assertEqual(is_even(5),False,"is_even(5)")
        self.assertEqual(is_even(-5),False,"is_even(-5)")
        self.assertEqual(is_even(0),True,"is_even(0)")
        self.assertEqual(is_even(513),False,"is_even(513)")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="leap-year">
<h3>18.3.5. Leap Year<a class="headerlink" href="#leap-year" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="leap-year" class="runestone explainer ac_section alert alert-warning">
<p>A year is a leap year if it is divisible by 4 unless it is a century that is not divisible by 400. Write a function that takes a year as a parameter and returns True if the year is a leap year, False otherwise. The following pseudocode determines whether a year is a leap year or a common year in the Gregorian calendar (from <a class="reference external" href="http://en.wikipedia.org/wiki/Leap_year#Algorithm">Wikipedia</a>):</p>
<blockquote>
<div><div class="line-block">
<div class="line">if (year is not divisible by 4) then (it is a common year)</div>
<div class="line">else if (year is not divisible by 100) then (it is a leap year)</div>
<div class="line">else if (year is not divisible by 400) then (it is a common year)</div>
<div class="line">else (it is a leap year)</div>
</div>
</div></blockquote>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">leap_year(2001)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal"><span class="pre">leap_year(2020)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">leap_year(1900)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

<textarea data-component="activecode" id=leap-year data-lang="python"    data-timelimit=25000   data-audio=''     >
def leap_year(year):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(leap_year(2001),False,"leap_year(2001)")
        self.assertEqual(leap_year(2020),True,"leap_year(2020)")
        self.assertEqual(leap_year(1900),False,"leap_year(1900)")
        self.assertEqual(leap_year(2004),True,"leap_year(2004)")
        self.assertEqual(leap_year(2000),True,"leap_year(2000)")


myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 17. More Details About Functions' data-toggle="tooltip">
    <a href="MoreFunctionDetails.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="MoreFunctionDetails.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 19. Reaction Game with Micro:bit' data-toggle="tooltip" >
    <a href="MicrobitReactionGame.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="MicrobitReactionGame.html" >Next Section - 19. Reaction Game with Micro:bit</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>