<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>7.30. Exercises &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">7.30. Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>7.30. Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" id="strings_ex_1" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='q1' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Draw a reference diagram for <code class="docutils literal"><span class="pre">a</span></code> and <code class="docutils literal"><span class="pre">b</span></code> before and after the third line of
the following python code is executed:</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a</span> <span class="o">=</span> <span class="p">[</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">]</span>
<span class="n">b</span> <span class="o">=</span> <span class="n">a</span><span class="p">[:]</span>
<span class="n">b</span><span class="p">[</span><span class="mi">0</span><span class="p">]</span> <span class="o">=</span> <span class="mi">5</span>
</pre></div>
</div>
</div><div data-component="tab" data-tabname="Answer" >
Your diagram should show two variables referring to two different lists.  <code class="docutils literal"><span class="pre">a</span></code> refers to the original list with 1,2, and 3.
<code class="docutils literal"><span class="pre">b</span></code> refers to a list with 5,2, and 3 since the zero-eth element was replaced with 5.</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_12314cf40dbe407cb145f029870c0347_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_12314cf40dbe407cb145f029870c0347_' + eBookConfig.course;
                    this.page.url = 'http://disqus_12314cf40dbe407cb145f029870c0347_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_12314cf40dbe407cb145f029870c0347_' + eBookConfig.course;
            disqus_url = 'http://disqus_12314cf40dbe407cb145f029870c0347_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_12314cf40dbe407cb145f029870c0347" onclick="disqus_12314cf40dbe407cb145f029870c0347_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_12314cf40dbe407cb145f029870c0347']").attr('data-disqus-identifier', 'disqus_12314cf40dbe407cb145f029870c0347_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_2" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_2" class="runestone explainer ac_section alert alert-warning">
Create a list called <code class="docutils literal"><span class="pre">myList</span></code> with the following six items: 76, 92.3, “hello”, True, 4, 76.  Do it with both append and with concatenation, one item at a time.
<textarea data-component="activecode" id=ex_9_2 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_3" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='q3' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_9_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ex_9_3 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
Starting with the list of the previous exercise, write Python statements to do the following:

a. Append "apple" and 76 to the list.
#. Insert the value "cat" at position 3.
#. Insert the value 99 at the start of the list.
#. Find the index of "hello".
#. Count the number of 76s in the list.
#. Remove the first occurrence of 76 from the list.
#. Remove True from the list using ``pop`` and ``index``.
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q3_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q3_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
myList = [76, 92.3, 'hello', True, 4, 76]

myList.append("apple")         # a
myList.append(76)              # a
myList.insert(3, "cat")        # b
myList.insert(0, 99)           # c

print(myList.index("hello"))   # d
print(myList.count(76))        # e
myList.remove(76)              # f
myList.pop(myList.index(True)) # g

print (myList)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_b9034b274ebe4c55a58c44315ee681a4_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_b9034b274ebe4c55a58c44315ee681a4_' + eBookConfig.course;
                    this.page.url = 'http://disqus_b9034b274ebe4c55a58c44315ee681a4_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_b9034b274ebe4c55a58c44315ee681a4_' + eBookConfig.course;
            disqus_url = 'http://disqus_b9034b274ebe4c55a58c44315ee681a4_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_b9034b274ebe4c55a58c44315ee681a4" onclick="disqus_b9034b274ebe4c55a58c44315ee681a4_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_b9034b274ebe4c55a58c44315ee681a4']").attr('data-disqus-identifier', 'disqus_b9034b274ebe4c55a58c44315ee681a4_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_4" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_4" class="runestone explainer ac_section alert alert-warning">
Create a list containing 100 random integers between 0 and 1000 (use iteration, append, and the random module).  Write a function called <code class="docutils literal"><span class="pre">average</span></code> that will take the list as a parameter and return the average.
<textarea data-component="activecode" id=ex_9_4 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_5" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='q5' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a Python function that will take a the list of 100 random integers between 0 and 1000 and return the maximum value.  (Note: there is a builtin function named <code class="docutils literal"><span class="pre">max</span></code> but pretend you cannot use it.)</p>

<div data-childcomponent="ex_9_5" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ex_9_5 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q5_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q5_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
import random

def max(lst):
    max = 0
    for e in lst:
        if e > max:
            max = e
    return max

lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(max(lst))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
                    this.page.url = 'http://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course;
            disqus_url = 'http://disqus_714fd5537ebf41189ce5fb6fb16d1d26_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_714fd5537ebf41189ce5fb6fb16d1d26" onclick="disqus_714fd5537ebf41189ce5fb6fb16d1d26_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_714fd5537ebf41189ce5fb6fb16d1d26']").attr('data-disqus-identifier', 'disqus_714fd5537ebf41189ce5fb6fb16d1d26_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_6" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_7_11" class="runestone explainer ac_section alert alert-warning">
Write a function <code class="docutils literal"><span class="pre">sum_of_squares(xs)</span></code> that computes the sum
of the squares of the numbers in the list <code class="docutils literal"><span class="pre">xs</span></code>.  For example,
<code class="docutils literal"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 4+9+16 which is 29:
<textarea data-component="activecode" id=ex_7_11 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def sum_of_squares(xs):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2,3,4]),29,"Tested sum_of_squares on input [2,3,4]")
        self.assertEqual(sum_of_squares([0,1,-1]),2,"Tested sum_of_squares on input [0,1,-1]")
        self.assertEqual(sum_of_squares([5,12,14]),365,"Tested sum_of_squares on input [5,12,14]")

myTests().main()
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_7" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='q7' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_9_6" class="runestone explainer ac_section alert alert-warning">
Write a function to count how many odd numbers are in a list.
<textarea data-component="activecode" id=ex_9_6 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def countOdd(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(countOdd([1,3,5,7,9]),5,"Tested countOdd on input [1,3,5,7,9]")
        self.assertEqual(countOdd([1,2,3,4,5]),3,"Tested countOdd on input [-1,-2,-3,-4,-5]")
        self.assertEqual(countOdd([2,4,6,8,10]),0,"Tested countOdd on input [2,4,6,8,10]")
        self.assertEqual(countOdd([0,-1,12,-33]),2,"Tested countOdd on input [0,-1,12,-33]")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q7_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q7_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
import random

def countOdd(lst):
    odd = 0
    for e in lst:
        if e % 2 != 0:
            odd = odd + 1
    return odd

# make a random list to test the function
lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(countOdd(lst))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_fdd366b1b4c8494082a385e1e1197844_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
                    this.page.url = 'http://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course;
            disqus_url = 'http://disqus_fdd366b1b4c8494082a385e1e1197844_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_fdd366b1b4c8494082a385e1e1197844" onclick="disqus_fdd366b1b4c8494082a385e1e1197844_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_fdd366b1b4c8494082a385e1e1197844']").attr('data-disqus-identifier', 'disqus_fdd366b1b4c8494082a385e1e1197844_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_8" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_7" class="runestone explainer ac_section alert alert-warning">
Sum up all the even numbers in a list.
<textarea data-component="activecode" id=ex_9_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def sumEven(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumEven([1,3,5,7,9]),0,"Tested sumEven on input [1,3,5,7,9]")
        self.assertEqual(sumEven([-1,-2,-3,-4,-5]),-6,"Tested sumEven on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumEven([2,4,6,7,9]),12,"Tested sumEven on input [2,4,6,7,9]")
        self.assertEqual(sumEven([0,1,12,33]),12,"Tested sumEven on input [0,1,12,33]")

myTests().main()
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_9" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='q9' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_9_8" class="runestone explainer ac_section alert alert-warning">
Sum up all the negative numbers in a list.
<textarea data-component="activecode" id=ex_9_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def sumNegatives(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumNegatives([-1,-2,-3,-4,-5]),-15,"Tested sumNegatives on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumNegatives([1,-3,5,-7,9]),-10,"Tested sumNegatives on input [1,-3,5,-7,9]")
        self.assertEqual(sumNegatives([-2,-4,6,-7,9]),-13,"Tested sumNegatives on input [-2,-4,6,-7,9]")
        self.assertEqual(sumNegatives([0,1,2,3,4]),0,"Tested sumNegatives on input [0,1,2,3,4]")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q9_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q9_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
import random

def sumNegative(lst):
    sum = 0
    for e in lst:
        if e < 0:
            sum = sum + e
    return sum

lst = []
for i in range(100):
    lst.append(random.randrange(-1000, 1000))

print(sumNegative(lst))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
                    this.page.url = 'http://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfe671ac1e0942f2be4de7179921f83f_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_bfe671ac1e0942f2be4de7179921f83f" onclick="disqus_bfe671ac1e0942f2be4de7179921f83f_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_bfe671ac1e0942f2be4de7179921f83f']").attr('data-disqus-identifier', 'disqus_bfe671ac1e0942f2be4de7179921f83f_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_10" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_9" class="runestone explainer ac_section alert alert-warning">
Count how many words in a list have length 5.
<textarea data-component="activecode" id=ex_9_9 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def countWords(lst):
    # your code here
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_11" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    <div id='q11' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Sum all the elements in a list up to but not including the first even number.</p>

<div data-childcomponent="ex_9_10" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ex_9_10 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def sumUntilEven(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumUntilEven([1,2,3,4,5]),1,"Tested sumUntilEven on input [1,2,3,4.5]")
        self.assertEqual(sumUntilEven([1,3,5,7,9]),25,"Tested sumUntilEven on input [1,3,5,7,9]")
        self.assertEqual(sumUntilEven([2,4,6,7,9]),0,"Tested sumUntilEven on input [2,4,6,7,9]")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q11_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q11_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
import random

def sum(lst):
    sum = 0
    index = 0
    while index < len(lst) and lst[index] % 2 != 0:
        sum = sum + lst[index]
        index = index + 1
    return sum

lst = []
for i in range(100):
    lst.append(random.randint(0,1000))

print(sum(lst))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_90f95bbe4a49428caa9ed0c5e02747b1_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_90f95bbe4a49428caa9ed0c5e02747b1_' + eBookConfig.course;
                    this.page.url = 'http://disqus_90f95bbe4a49428caa9ed0c5e02747b1_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_90f95bbe4a49428caa9ed0c5e02747b1_' + eBookConfig.course;
            disqus_url = 'http://disqus_90f95bbe4a49428caa9ed0c5e02747b1_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_90f95bbe4a49428caa9ed0c5e02747b1" onclick="disqus_90f95bbe4a49428caa9ed0c5e02747b1_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_90f95bbe4a49428caa9ed0c5e02747b1']").attr('data-disqus-identifier', 'disqus_90f95bbe4a49428caa9ed0c5e02747b1_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_12" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_11" class="runestone explainer ac_section alert alert-warning">
Count how many words occur in a list up to and including the first occurrence of the word “sam”.
<textarea data-component="activecode" id=ex_9_11 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def count(lst):
    # your code here
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_13" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    <div id='q13' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_9_12" class="runestone explainer ac_section alert alert-warning">
<p>Although Python provides us with many list methods, it is good practice and very instructive to think about how they are implemented.  Implement a Python function that works like the following:</p>
<ol class="loweralpha simple">
<li>count</li>
<li>in</li>
<li>reverse</li>
<li>index</li>
<li>insert</li>
</ol>

<textarea data-component="activecode" id=ex_9_12 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q13_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q13_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def count(obj, lst):
    count = 0
    for e in lst:
        if e == obj:
            count = count + 1
    return count

def is_in(obj, lst):  # cannot be called in() because in is a reserved keyword
    for e in lst:
        if e == obj:
            return True
    return False

def reverse(lst):
    reversed = []
    for i in range(len(lst)-1, -1, -1): # step through the original list backwards
        reversed.append(lst[i])
    return reversed

def index(obj, lst):
    for i in range(len(lst)):
        if lst[i] == obj:
            return i
    return -1

def insert(obj, index, lst):
    newlst = []
    for i in range(len(lst)):
        if i == index:
            newlst.append(obj)
        newlst.append(lst[i])
    return newlst

lst = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, lst))
print(is_in(4, lst))
print(reverse(lst))
print(index(2, lst))
print(insert('cat', 4, lst))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
                    this.page.url = 'http://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course;
            disqus_url = 'http://disqus_39ee0274e51d4c888cc20b6fefa4069c_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_39ee0274e51d4c888cc20b6fefa4069c" onclick="disqus_39ee0274e51d4c888cc20b6fefa4069c_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_39ee0274e51d4c888cc20b6fefa4069c']").attr('data-disqus-identifier', 'disqus_39ee0274e51d4c888cc20b6fefa4069c_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_14" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_13" class="runestone explainer ac_section alert alert-warning">
<p>Write a function <code class="docutils literal"><span class="pre">replace(s,</span> <span class="pre">old,</span> <span class="pre">new)</span></code> that replaces all occurences of
<code class="docutils literal"><span class="pre">old</span></code> with <code class="docutils literal"><span class="pre">new</span></code> in a string <code class="docutils literal"><span class="pre">s</span></code>:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="s1">&#39;Mississippi&#39;</span><span class="p">,</span> <span class="s1">&#39;i&#39;</span><span class="p">,</span> <span class="s1">&#39;I&#39;</span><span class="p">),</span> <span class="s1">&#39;MIssIssIppI&#39;</span><span class="p">)</span>

<span class="n">s</span> <span class="o">=</span> <span class="s1">&#39;I love spom!  Spom is my favorite food.  Spom, spom, spom, yum!&#39;</span>
<span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="n">s</span><span class="p">,</span> <span class="s1">&#39;om&#39;</span><span class="p">,</span> <span class="s1">&#39;am&#39;</span><span class="p">),</span>
       <span class="s1">&#39;I love spam!  Spam is my favorite food.  Spam, spam, spam, yum!&#39;</span><span class="p">)</span>

<span class="n">test</span><span class="p">(</span><span class="n">replace</span><span class="p">(</span><span class="n">s</span><span class="p">,</span> <span class="s1">&#39;o&#39;</span><span class="p">,</span> <span class="s1">&#39;a&#39;</span><span class="p">),</span>
       <span class="s1">&#39;I lave spam!  Spam is my favarite faad.  Spam, spam, spam, yum!&#39;</span><span class="p">)</span>
</pre></div>
</div>
<p><em>Hint</em>: use the <code class="docutils literal"><span class="pre">split</span></code> and <code class="docutils literal"><span class="pre">join</span></code> methods.</p>

<textarea data-component="activecode" id=ex_9_13 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def replace(s, old, new):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(replace('Mississippi','i','I'),'MIssIssIppI',"Tested replace on input 'Mississippi','i','I'")
        self.assertEqual(replace('Bookkeeper','e','A'),'BookkAApAr',"Tested failed on input 'Bookkeeper','e','A'")
        self.assertEqual(replace('Deeded','e','q'),'Dqqdqd',"Tested failed on input 'Deeded','e','q'")

myTests().main()
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_15" >
    <ol start=15 class=arabic><li class="alert alert-warning">

    <div id='q15' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_9_14" class="runestone explainer ac_section alert alert-warning">
<p>Here are the rules for an L-system that creates something that resembles
a common garden herb.  Implement the following rules and try it.  Use an
angle of 25.7 degrees.</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">H</span>
<span class="n">H</span> <span class="o">--&gt;</span> <span class="n">HFX</span><span class="p">[</span><span class="o">+</span><span class="n">H</span><span class="p">][</span><span class="o">-</span><span class="n">H</span><span class="p">]</span>
<span class="n">X</span> <span class="o">--&gt;</span> <span class="n">X</span><span class="p">[</span><span class="o">-</span><span class="n">FFF</span><span class="p">][</span><span class="o">+</span><span class="n">FFF</span><span class="p">]</span><span class="n">FX</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_9_14 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q15_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q15_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'H':
        newstr = 'HFX[+H][-H]'   # Rule 1
    elif ch == 'X':
        newstr = 'X[-FFF][+FFF]FX'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    savedInfoList = []
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)
        elif cmd == '[':
            savedInfoList.append([aTurtle.heading(), aTurtle.xcor(), aTurtle.ycor()])
            #print(savedInfoList)
        elif cmd == ']':
            newInfo = savedInfoList.pop()
            aTurtle.setheading(newInfo[0])
            aTurtle.setposition(newInfo[1], newInfo[2])


def main():
    inst = createLSystem(4, "H")   # create the string
    print(inst)
    t = turtle.Turtle()            # create the turtle
    wn = turtle.Screen()
    t.up()
    t.back(200)
    t.down()
    t.speed(9)
    drawLsystem(t, inst, 27.5, 5)  # draw the picture

    wn.exitonclick()

main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_10a1fdd1fadb401b8e733afd9d105052_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_10a1fdd1fadb401b8e733afd9d105052_' + eBookConfig.course;
                    this.page.url = 'http://disqus_10a1fdd1fadb401b8e733afd9d105052_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_10a1fdd1fadb401b8e733afd9d105052_' + eBookConfig.course;
            disqus_url = 'http://disqus_10a1fdd1fadb401b8e733afd9d105052_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_10a1fdd1fadb401b8e733afd9d105052" onclick="disqus_10a1fdd1fadb401b8e733afd9d105052_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_10a1fdd1fadb401b8e733afd9d105052']").attr('data-disqus-identifier', 'disqus_10a1fdd1fadb401b8e733afd9d105052_' + eBookConfig.course);
</script>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="strings_ex_16" >
    <ol start=16 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_9_16" class="runestone explainer ac_section alert alert-warning">
<p>Here is another L-System.  Use an Angle of 25.</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">F</span>
<span class="n">F</span> <span class="o">--&gt;</span> <span class="n">F</span><span class="p">[</span><span class="o">-</span><span class="n">F</span><span class="p">]</span><span class="n">F</span><span class="p">[</span><span class="o">+</span><span class="n">F</span><span class="p">]</span><span class="n">F</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_9_16 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    </div>


  
      

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>