<!DOCTYPE html>

<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.19: https://docutils.sourceforge.io/" />

    <title>1. Using Lists &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/637.0fa6cababf056764.css?v=F9484CEF" />
    <link rel="stylesheet" type="text/css" href="../../_static/runestone.3eb60bc006ae9a54.css?v=3A5A250C" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/&#34;&#34;/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script src="../../_static/runestonebase.js"></script>
    <script data-url_root="../../" id="documentation_options" src="../../_static/documentation_options.js"></script>
    <script src="../../_static/jquery.js"></script>
    <script src="../../_static/underscore.js"></script>
    <script src="../../_static/_sphinx_javascript_frameworks_compat.js"></script>
    <script src="../../_static/doctools.js"></script>
    <script src="../../_static/sphinx_highlight.js"></script>
    <script defer="" src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../../_static/runtime.66b619f51b746382.bundle.js?v=C3F69EE5"></script>
    <script defer="" src="../../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../../_static/runestone.ecc6dff5cad8b815.bundle.js?v=AE959292"></script>
    <script async="async" src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-mml-chtml.js"></script>
    <script src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../../_static/jquery-fix.js"></script>
    <script src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../../_static/bootstrap-sphinx.js"></script>
    <script src="../../_static/waypoints.min.js"></script>
    <script src="../../_static/rangy-core.js"></script>
    <script src="../../_static/rangy-textrange.js"></script>
    <script src="../../_static/rangy-cssclassapplier.js"></script>
    <script src="../../_static/user-highlights.js"></script>
    <script src="../../_static/jquery.idle-timer.js"></script>
    <script src="../../_static/processing-1.4.1.min.js"></script>
    <script src="../../_static/jquery.hotkey.js"></script>
    <script src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="2. List Methods" href="ListMethods.html" />
    <link rel="prev" title="4. Turtle Drawing and Strings (L-Systems)" href="../Strings/LSystems.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '6.6.2';
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">ActiveCode Work Area</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">1. Using Lists</a><ul>
<li><a class="reference internal" href="#list-overview">1.1. List Overview</a></li>
<li><a class="reference internal" href="#list-values">1.2. List Values</a></li>
<li><a class="reference internal" href="#list-length">1.3. List Length</a><ul>
<li><a class="reference internal" href="#check-your-understanding">1.3.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#accessing-elements">1.4. Accessing Elements</a><ul>
<li><a class="reference internal" href="#id1">1.4.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#list-membership">1.5. List Membership</a><ul>
<li><a class="reference internal" href="#id2">1.5.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#list-slices">1.6. List Slices</a><ul>
<li><a class="reference internal" href="#id3">1.6.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#strings-are-immutable">1.7. Strings are Immutable</a></li>
<li><a class="reference internal" href="#lists-are-mutable">1.8. Lists are Mutable</a><ul>
<li><a class="reference internal" href="#id4">1.8.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#lists-and-for-loops">1.9. Lists and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops</a></li>
<li><a class="reference internal" href="#practice-problems">1.10. Practice Problems</a><ul>
<li><a class="reference internal" href="#warmup">1.10.1. Warmup</a></li>
<li><a class="reference internal" href="#make-it-fun">1.10.2. Make it Fun</a></li>
<li><a class="reference internal" href="#sum-of-squares">1.10.3. Sum of Squares</a></li>
<li><a class="reference internal" href="#count-odd-numbers">1.10.4. Count Odd Numbers</a></li>
<li><a class="reference internal" href="#count-words-with-length-5">1.10.5. Count Words with Length 5</a></li>
<li><a class="reference internal" href="#look-for-saskatchewan">1.10.6. Look For “saskatchewan”</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4. </span>Turtle Drawing and Strings (L-Systems)" data-toggle="tooltip">
    <a href="../Strings/LSystems.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">2. </span>List Methods' data-toggle="tooltip" >
    <a href="ListMethods.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <section id="using-lists">
<h1><span class="section-number">1. </span>Using Lists<a class="headerlink" href="#using-lists" title="Permalink to this heading">¶</a></h1>
<aside class="topic">
<p class="topic-title">Quick Overview of Day</p>
<p>Creating lists. Checking the length of a list. Access elements of a list using the index operator. Using <code class="docutils literal notranslate"><span class="pre">in</span></code> and <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code>. Taking a slice of a list. Immutability of strings, mutability of lists. For loops with lists.</p>
</aside>

    <div data-component="reveal" class="runestone " id="curriculum_addressed_lists_overview"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide"  style="visibility: hidden;">
    <ul class="simple">
<li><p><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</p></li>
<li><p><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</p></li>
<li><p><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</p></li>
<li><p><strong>CS20-FP3</strong> Construct and utilize functions to create reusable pieces of code.</p></li>
<li><p><strong>CS20-FP4</strong>  Investigate one-dimensional arrays and their applications.</p></li>
</ul>

    </div>
    <section id="list-overview">
<h2><span class="section-number">1.1. </span>List Overview<a class="headerlink" href="#list-overview" title="Permalink to this heading">¶</a></h2>
<p>We’ve been using lists throughout the course, so you probably already feel quite comfortable with them. Some of the following information will be a review for you, but we will spend a bit of time going into more depth than we have done up to now.</p>
<p>A <strong>list</strong> is a sequential collection of Python data values, where each value is identified by an
index. The values that make up a list are called its <strong>elements</strong>. Lists are
similar to strings, which are ordered collections of characters, except that the
elements of a list can have any type and for any one list, the items can be of different types.</p>
</section>
<section id="list-values">
<h2><span class="section-number">1.2. </span>List Values<a class="headerlink" href="#list-values" title="Permalink to this heading">¶</a></h2>
<p>There are several ways to create a new list.  The simplest is to enclose the
elements in square brackets ( <code class="docutils literal notranslate"><span class="pre">[</span></code> and <code class="docutils literal notranslate"><span class="pre">]</span></code>).</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="p">[</span><span class="mi">10</span><span class="p">,</span> <span class="mi">20</span><span class="p">,</span> <span class="mi">30</span><span class="p">,</span> <span class="mi">40</span><span class="p">]</span>
<span class="p">[</span><span class="s2">&quot;spam&quot;</span><span class="p">,</span> <span class="s2">&quot;bungee&quot;</span><span class="p">,</span> <span class="s2">&quot;swallow&quot;</span><span class="p">]</span>
</pre></div>
</div>
<p>The first example is a list of four integers. The second is a list of three
strings. As we said above, the elements of a list don’t have to be the same type.  The following
list contains a string, a float, an integer, and
another list.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="p">[</span><span class="s2">&quot;hello&quot;</span><span class="p">,</span> <span class="mf">2.0</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="p">[</span><span class="mi">10</span><span class="p">,</span> <span class="mi">20</span><span class="p">]]</span>
</pre></div>
</div>
<p>A list within another list is said to be <strong>nested</strong> and the inner list is often called a <strong>sublist</strong>.
Finally, there is a special list that contains no elements. It is called the
empty list and is denoted <code class="docutils literal notranslate"><span class="pre">[]</span></code>.</p>
<p>As you would expect, we can also assign list values to variables and pass lists as parameters to functions.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_1 data-question_label="1.2.1">
<div id=list_overview_1_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
vocabulary = ["iteration", "selection", "control"]
numbers = [17, 123]
empty = []
mixed_list = ["hello", 2.0, 5*2, [10, 20]]

print(numbers)
print(mixed_list)
new_list = [ numbers, vocabulary ]
print(new_list)
</textarea>
</div>
</div>
</section>
<section id="list-length">
<span id="index-0"></span><h2><span class="section-number">1.3. </span>List Length<a class="headerlink" href="#list-length" title="Permalink to this heading">¶</a></h2>
<p>As with strings, the function <code class="docutils literal notranslate"><span class="pre">len</span></code> returns the length of a list (the number
of items in the list).  However, since lists can have items which are themselves lists, it important to note
that <code class="docutils literal notranslate"><span class="pre">len</span></code> only returns the top-most length.  In other words, sublists are considered to be a single
item when counting the length of the list.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_2 data-question_label="1.3.1">
<div id=list_overview_2_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
a_list =  ["hello", 2.0, 5, [10, 20]]
print(len(a_list))
print(len(['Hepburn', 1, ['Saskatoon', 'Regina', 'Prince Albert'], [1, 2, 3]]))
</textarea>
</div>
</div>
<section id="check-your-understanding">
<h3><span class="section-number">1.3.1. </span>Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this heading">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.3.1.1" data-multipleanswers="false"  id="test_question9_2_1"  style="visibility: hidden;">
    <p>lists-overview3: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer"  id="test_question9_2_1_opt_a">4</li><li data-component="feedback">len returns the actual number of items in the list, not the maximum index value.</li>
    
    <li data-component="answer" data-correct='yes' id="test_question9_2_1_opt_b">5</li><li data-component="feedback">Yes, there are 5 items in this list.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.3.1.2" data-multipleanswers="false"  id="test_question9_2_2"  style="visibility: hidden;">
    <p>lists-overview4: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="test_question9_2_2_opt_a">7</li><li data-component="feedback">Yes, there are 7 items in this list even though two of them happen to also be lists.</li>
    
    <li data-component="answer"  id="test_question9_2_2_opt_b">8</li><li data-component="feedback">len returns the number of top level items in the list.  It does not count items in sublists.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="accessing-elements">
<span id="index-1"></span><h2><span class="section-number">1.4. </span>Accessing Elements<a class="headerlink" href="#accessing-elements" title="Permalink to this heading">¶</a></h2>
<p>The syntax for accessing the elements of a list is the same as the syntax for
accessing the characters of a string.  We use the index operator ( <code class="docutils literal notranslate"><span class="pre">[]</span></code> – not to
be confused with an empty list). The expression inside the brackets specifies
the index. Remember that the indices start at 0.  Any integer expression can be used
as an index and as with strings, negative index values will locate items from the right instead
of from the left.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_3 data-question_label="1.4.1">
<div id=list_overview_3_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
numbers = [17, 123, 87, 34, 66, 8398, 44]
print(numbers[2])
print(numbers[9 - 8])
print(numbers[-2])
print(numbers[len(numbers) - 1])
</textarea>
</div>
</div>
<p>Also note that you can use the index operator <code class="docutils literal notranslate"><span class="pre">[]</span></code> more than once in a row to access things inside the list. For example, if you had a list of places in Saskatchewan, you could access both a place, as well as an individual character, as follows:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_4 data-question_label="1.4.2">
<div id=list_overview_4_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
places = ["Saskatoon", "Regina", "Prince Albert", "Moose Jaw", "Yorkton", "Swift Current", "North Battleford"]
print(places[2])
print(places[2][1])
</textarea>
</div>
</div>
<p>The same syntax works if you are trying to access an element of a list inside a list as well. Consider the following:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_5 data-question_label="1.4.3">
<div id=list_overview_5_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_5_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
places = [["Saskatoon", "Regina"], ["Prince Albert", "Moose Jaw"], ["Yorkton", "Swift Current", "North Battleford"]]
print(places[2])
print(places[2][1])
print(places[2][1][3])
</textarea>
</div>
</div>
<section id="id1">
<h3><span class="section-number">1.4.1. </span>Check Your Understanding<a class="headerlink" href="#id1" title="Permalink to this heading">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.4.1.1" data-multipleanswers="false"  id="test_question9_3_1"  style="visibility: hidden;">
    <p>lists-overview8: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="test_question9_3_1_opt_a">[ ]</li><li data-component="feedback">The empty list is at index 4.</li>
    
    <li data-component="answer" data-correct='yes' id="test_question9_3_1_opt_b">3.14</li><li data-component="feedback">Yes, 3.14 is at index 5 since we start counting at 0 and sublists count as one item.</li>
    
    <li data-component="answer"  id="test_question9_3_1_opt_c">False</li><li data-component="feedback">False is at index 6.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.4.1.2" data-multipleanswers="false"  id="test_question9_3_3"  style="visibility: hidden;">
    <p>lists-overview9: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">2</span><span class="p">][</span><span class="mi">0</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="test_question9_3_3_opt_a">56</li><li data-component="feedback">Indexes start with 0, not 1.</li>
    
    <li data-component="answer" data-correct='yes' id="test_question9_3_3_opt_b">c</li><li data-component="feedback">Yes, the first character of the string at index 2 is c</li>
    
    <li data-component="answer"  id="test_question9_3_3_opt_c">cat</li><li data-component="feedback">cat is the item at index 2 but then we index into it further.</li>
    
    <li data-component="answer"  id="test_question9_3_3_opt_d">Error, you cannot have two index values unless you are using slicing.</li><li data-component="feedback">Using more than one index is fine.  You read it from left to right.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="list-membership">
<span id="index-2"></span><h2><span class="section-number">1.5. </span>List Membership<a class="headerlink" href="#list-membership" title="Permalink to this heading">¶</a></h2>
<p><code class="docutils literal notranslate"><span class="pre">in</span></code> and <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> are boolean operators that test membership in a sequence. We
used them previously with strings and they also work here.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_6 data-question_label="1.5.1">
<div id=list_overview_6_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_6_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = ["apple", "orange", "banana", "cherry"]

print("apple" in fruit)
print("pear" in fruit)
</textarea>
</div>
</div>
<section id="id2">
<h3><span class="section-number">1.5.1. </span>Check Your Understanding<a class="headerlink" href="#id2" title="Permalink to this heading">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.5.1.1" data-multipleanswers="false"  id="test_question9_4_1"  style="visibility: hidden;">
    <p>lists-overview11: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="mf">3.14</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="test_question9_4_1_opt_a">True</li><li data-component="feedback">Yes, 3.14 is an item in the list a_list.</li>
    
    <li data-component="answer"  id="test_question9_4_1_opt_b">False</li><li data-component="feedback">There are 7 items in the list, 3.14 is one of them.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.5.1.2" data-multipleanswers="false"  id="test_question9_4_2"  style="visibility: hidden;">
    <p>lists-overview12: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="mi">57</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="test_question9_4_2_opt_a">True</li><li data-component="feedback">in returns True for top level items only.  57 is in a sublist.</li>
    
    <li data-component="answer" data-correct='yes' id="test_question9_4_2_opt_b">False</li><li data-component="feedback">Yes, 57 is not a top level item in a_list.  It is in a sublist.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="list-slices">
<span id="index-3"></span><h2><span class="section-number">1.6. </span>List Slices<a class="headerlink" href="#list-slices" title="Permalink to this heading">¶</a></h2>
<p>The slice operation we saw with strings also work on lists.  Remember that the first index is the starting point for the slice and the second number is one index past the end of the slice (up to but not including that element).  Recall also
that if you omit the first index (before the colon), the slice starts at the
beginning of the sequence. If you omit the second index, the slice goes to the
end of the sequence.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_7 data-question_label="1.6.1">
<div id=list_overview_7_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
a_list = ['a', 'b', 'c', 'd', 'e', 'f']
print(a_list[1:3])
print(a_list[:4])
print(a_list[3:])
print(a_list[:])
</textarea>
</div>
</div>
<section id="id3">
<h3><span class="section-number">1.6.1. </span>Check Your Understanding<a class="headerlink" href="#id3" title="Permalink to this heading">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.6.1.1" data-multipleanswers="false"  id="test_question9_6_1"  style="visibility: hidden;">
    <p>lists-overview14: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="kc">False</span><span class="p">]</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">4</span><span class="p">:])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="test_question9_6_1_opt_a">[ [ ], 3.14, False]</li><li data-component="feedback">Yes, the slice starts at index 4 and goes up to and including the last item.</li>
    
    <li data-component="answer"  id="test_question9_6_1_opt_b">[ [ ], 3.14]</li><li data-component="feedback">By leaving out the upper bound on the slice, we go up to and including the last item.</li>
    
    <li data-component="answer"  id="test_question9_6_1_opt_c">[ [56, 57, "dog"], [ ], 3.14, False]</li><li data-component="feedback">Index values start at 0.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="strings-are-immutable">
<span id="index-4"></span><h2><span class="section-number">1.7. </span>Strings are Immutable<a class="headerlink" href="#strings-are-immutable" title="Permalink to this heading">¶</a></h2>
<p>One thing that makes strings different from Python lists is that
you are not allowed to modify the individual characters in a string.  It is tempting to use the <code class="docutils literal notranslate"><span class="pre">[]</span></code> operator on the left side of an assignment,
with the intention of changing a character in a string.  For example, in the following code, we would like to change the first letter of <code class="docutils literal notranslate"><span class="pre">greeting</span></code>.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=cg08_imm1 data-question_label="1.7.1">
<div id=cg08_imm1_question class="ac_question">

</div>
<textarea data-lang="python" id="cg08_imm1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
greeting = "Hello, world!"
greeting[0] = 'J'            # ERROR!
print(greeting)
</textarea>
</div>
</div>
<p>Instead of producing the output <code class="docutils literal notranslate"><span class="pre">Jello,</span> <span class="pre">world!</span></code>, this code produces the
runtime error <code class="docutils literal notranslate"><span class="pre">TypeError:</span> <span class="pre">'str'</span> <span class="pre">object</span> <span class="pre">does</span> <span class="pre">not</span> <span class="pre">support</span> <span class="pre">item</span> <span class="pre">assignment</span></code>.</p>
<p>Strings are <strong>immutable</strong>, which means you cannot change an existing string. The
best you can do is create a new string that is a variation on the original.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ch08_imm2 data-question_label="1.7.2">
<div id=ch08_imm2_question class="ac_question">

</div>
<textarea data-lang="python" id="ch08_imm2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
greeting = "Hello, world!"
new_greeting = 'J' + greeting[1:]
print(new_greeting)
print(greeting)            # same as it was
</textarea>
</div>
</div>
<p>The solution here is to concatenate a new first letter onto a slice of
<code class="docutils literal notranslate"><span class="pre">greeting</span></code>. This operation has no effect on the original string.</p>
</section>
<section id="lists-are-mutable">
<span id="index-5"></span><h2><span class="section-number">1.8. </span>Lists are Mutable<a class="headerlink" href="#lists-are-mutable" title="Permalink to this heading">¶</a></h2>
<p>Unlike strings, lists are <strong>mutable</strong>.  This means we can change an item in a list by accessing
it directly as part of the assignment statement. Using the indexing operator (square brackets) on the left side of an assignment, we can
update one of the list items.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=ch09_7 data-question_label="1.8.1">
<div id=ch09_7_question class="ac_question">

</div>
<textarea data-lang="python" id="ch09_7_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = ["banana", "apple", "cherry"]
print(fruit)

fruit[0] = "pear"
fruit[-1] = "orange"
print(fruit)
</textarea>
</div>
</div>
<p>An assignment to an element of a list is called <strong>item assignment</strong>. Item
assignment does not work for strings.  Recall that strings are immutable.</p>
<section id="id4">
<h3><span class="section-number">1.8.1. </span>Check Your Understanding<a class="headerlink" href="#id4" title="Permalink to this heading">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.8.1.1" data-multipleanswers="false"  id="test_question9_7_1"  style="visibility: hidden;">
    <p>lists-overview18: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">4</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="n">a_list</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">=</span> <span class="kc">True</span>
<span class="nb">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="test_question9_7_1_opt_a">[4, 2, True, 8, 6, 5]</li><li data-component="feedback">Item assignment does not insert the new item into the list.</li>
    
    <li data-component="answer" data-correct='yes' id="test_question9_7_1_opt_b">[4, 2, True, 6, 5]</li><li data-component="feedback">Yes, the value True is placed in the list at index 2.  It replaces 8.</li>
    
    <li data-component="answer"  id="test_question9_7_1_opt_c">Error, it is illegal to assign</li><li data-component="feedback">Item assignment is allowed with lists.  Lists are mutable.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="lists-and-for-loops">
<h2><span class="section-number">1.9. </span>Lists and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops<a class="headerlink" href="#lists-and-for-loops" title="Permalink to this heading">¶</a></h2>
<p>It is also possible to perform <strong>list traversal</strong> using iteration by item as well as iteration by index.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=chp09_03a data-question_label="1.9.1">
<div id=chp09_03a_question class="ac_question">

</div>
<textarea data-lang="python" id="chp09_03a_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruits = ["apple", "orange", "banana", "cherry"]

# looping by item
for some_fruit in fruits:
    print(some_fruit)
</textarea>
</div>
</div>
<p>It almost reads like natural language: For (every) fruit in (the list of) fruits,
print (the name of the) fruit.</p>
<p>We can also use the indices to access the items in an iterative fashion.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=chp09_03b data-question_label="1.9.2">
<div id=chp09_03b_question class="ac_question">

</div>
<textarea data-lang="python" id="chp09_03b_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruits = ["apple", "orange", "banana", "cherry"]

# looping by index
for position in range(len(fruits)):
    print(fruits[position])
</textarea>
</div>
</div>
<p>In this example, each time through the loop, the variable <code class="docutils literal notranslate"><span class="pre">position</span></code> is used as an index into the
list, printing the <code class="docutils literal notranslate"><span class="pre">position</span></code>-eth element. Note that we used <code class="docutils literal notranslate"><span class="pre">len</span></code> as the upper bound on the range
so that we can iterate correctly no matter how many items are in the list.</p>
<p>Since lists are mutable, it is often desirable to traverse a list, modifying
each of its elements as you go. The following code squares all the numbers from <code class="docutils literal notranslate"><span class="pre">1</span></code> to
<code class="docutils literal notranslate"><span class="pre">5</span></code> using iteration by position.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=chp09_for4 data-question_label="1.9.3">
<div id=chp09_for4_question class="ac_question">

</div>
<textarea data-lang="python" id="chp09_for4_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
numbers = [1, 2, 3, 4, 5]
print(numbers)

for i in range(len(numbers)):
    numbers[i] = numbers[i] ** 2

print(numbers)
</textarea>
</div>
</div>
<p>Take a moment to think about <code class="docutils literal notranslate"><span class="pre">range(len(numbers))</span></code> until you understand how
it works. We are interested here in both the <em>value</em> and its <em>index</em> within the
list, so that we can assign a new value to it.</p>
</section>
<section id="practice-problems">
<h2><span class="section-number">1.10. </span>Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this heading">¶</a></h2>
<section id="warmup">
<h3><span class="section-number">1.10.1. </span>Warmup<a class="headerlink" href="#warmup" title="Permalink to this heading">¶</a></h3>
<p>Create a list called <code class="docutils literal notranslate"><span class="pre">my_list</span></code> with the following six items: <code class="docutils literal notranslate"><span class="pre">76,</span> <span class="pre">92.3,</span> <span class="pre">&quot;hello&quot;,</span> <span class="pre">True,</span> <span class="pre">4,</span> <span class="pre">76</span></code>. Print out the list to make sure it is doing what you intend.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_1 data-question_label="1.10.1.1">
<div id=list_overview_practice_problem_1_question class="ac_question">

</div>
<textarea data-lang="python" id="list_overview_practice_problem_1_editor" 
      data-timelimit=25000  data-codelens="true" data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# your code goes here!
</textarea>
</div>
</div>
</section>
<section id="make-it-fun">
<h3><span class="section-number">1.10.2. </span>Make it Fun<a class="headerlink" href="#make-it-fun" title="Permalink to this heading">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_6 data-question_label="1.10.2.1">
<div id=list_overview_practice_problem_6_question class="ac_question">
<p>Write a function to replace the first element of a list with the word <code class="docutils literal notranslate"><span class="pre">&quot;fun&quot;</span></code>, and then return the altered list. You can assume that the list will always have at least one element.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">make_fun([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;science&quot;]</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">make_fun([&quot;sask&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;]</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">make_fun([&quot;we&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;]</span></code></p>

</div>
<textarea data-lang="python" id="list_overview_practice_problem_6_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def make_fun(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(make_fun(["computer", "science"]),["fun", "science"],'make_fun(["computer", "science"])')
      self.assertEqual(make_fun(["sask", "is", "nice"]),["fun", "is", "nice"],'make_fun(["sask", "is", "nice"])')
      self.assertEqual(make_fun(["we", "live", "in", "saskatchewan"]),["fun", "live", "in", "saskatchewan"],'make_fun(["we", "live", "in", "saskatchewan"])')

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="sum-of-squares">
<h3><span class="section-number">1.10.3. </span>Sum of Squares<a class="headerlink" href="#sum-of-squares" title="Permalink to this heading">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_2 data-question_label="1.10.3.1">
<div id=list_overview_practice_problem_2_question class="ac_question">
<p>Write a function <code class="docutils literal notranslate"><span class="pre">sum_of_squares(a_list)</span></code> that computes the sum of the squares of the numbers in the list <code class="docutils literal notranslate"><span class="pre">a_list</span></code>.  For example, <code class="docutils literal notranslate"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 29, which is the sum of 4+9+16.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">sum_of_squares([2,3,4])</span> <span class="pre">→</span> <span class="pre">29</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">sum_of_squares([0,1,-1])</span> <span class="pre">→</span> <span class="pre">2</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">sum_of_squares([5,12])</span> <span class="pre">→</span> <span class="pre">169</span></code></p>

</div>
<textarea data-lang="python" id="list_overview_practice_problem_2_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def sum_of_squares(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(sum_of_squares([2,3,4]),29,"sum_of_squares([2,3,4])")
      self.assertEqual(sum_of_squares([0,1,-1]),2,"sum_of_squares([0,1,-1])")
      self.assertEqual(sum_of_squares([5,12]),169,"sum_of_squares([5,12])")
      self.assertEqual(sum_of_squares([5,12,14]),365,"sum_of_squares([5,12,14])")
      self.assertEqual(sum_of_squares([5]),25,"sum_of_squares([5])")

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="count-odd-numbers">
<h3><span class="section-number">1.10.4. </span>Count Odd Numbers<a class="headerlink" href="#count-odd-numbers" title="Permalink to this heading">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_3 data-question_label="1.10.4.1">
<div id=list_overview_practice_problem_3_question class="ac_question">
<p>Write a function to count how many odd numbers are in a list.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([1,3,5,7,9])</span> <span class="pre">→</span> <span class="pre">5</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([1,2,3,4,5])</span> <span class="pre">→</span> <span class="pre">3</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_odds([2,4,6,8,10])</span> <span class="pre">→</span> <span class="pre">0</span></code></p>

</div>
<textarea data-lang="python" id="list_overview_practice_problem_3_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def count_odds(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(count_odds([1,3,5,7,9]),5,"count_odds([1,3,5,7,9])")
      self.assertEqual(count_odds([1,2,3,4,5]),3,"count_odds([-1,-2,-3,-4,-5])")
      self.assertEqual(count_odds([2,4,6,8,10]),0,"count_odds([2,4,6,8,10])")
      self.assertEqual(count_odds([0,-1,12,-33]),2,"count_odds([0,-1,12,-33])")

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="count-words-with-length-5">
<h3><span class="section-number">1.10.5. </span>Count Words with Length 5<a class="headerlink" href="#count-words-with-length-5" title="Permalink to this heading">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_4 data-question_label="1.10.5.1">
<div id=list_overview_practice_problem_4_question class="ac_question">
<p>Write a function to count how many words in a list have a length of 5.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">count_words([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">0</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_words([&quot;nobody&quot;,</span> <span class="pre">&quot;knows&quot;])</span> <span class="pre">→</span> <span class="pre">1</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">count_words([&quot;think&quot;,</span> <span class="pre">&quot;pizza&quot;,</span> <span class="pre">&quot;quick&quot;])</span> <span class="pre">→</span> <span class="pre">3</span></code></p>

</div>
<textarea data-lang="python" id="list_overview_practice_problem_4_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def count_words(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(count_words(["computer", "science"]),0,'count_words(["computer", "science"])')
      self.assertEqual(count_words(["nobody", "knows"]),1,'count_words(["nobody", "knows"])')
      self.assertEqual(count_words(["think", "pizza", "quick"]),3,'count_words(["think", "pizza", "quick"])')

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="look-for-saskatchewan">
<h3><span class="section-number">1.10.6. </span>Look For “saskatchewan”<a class="headerlink" href="#look-for-saskatchewan" title="Permalink to this heading">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=list_overview_practice_problem_5 data-question_label="1.10.6.1">
<div id=list_overview_practice_problem_5_question class="ac_question">
<p>Write a function to check if the word “saskatchewan” shows up somewhere in a list. If it does, return <code class="docutils literal notranslate"><span class="pre">True</span></code>. If it doesn’t, return <code class="docutils literal notranslate"><span class="pre">False</span></code>.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">has_sask([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">has_sask([&quot;sask&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;])</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">has_sask([&quot;we&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;])</span> <span class="pre">→</span> <span class="pre">True</span></code></p>

</div>
<textarea data-lang="python" id="list_overview_practice_problem_5_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def has_sask(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(has_sask(["computer", "science"]),False,'has_sask(["computer", "science"])')
      self.assertEqual(has_sask(["sask", "is", "nice"]),False,'has_sask(["sask", "is", "nice"])')
      self.assertEqual(has_sask(["we", "live", "in", "saskatchewan"]),True,'has_sask(["we", "live", "in", "saskatchewan"])')
      self.assertEqual(has_sask(["saskatchewan", "roughriders"]),True,'has_sask(["saskatchewan", "roughriders"])')
      self.assertEqual(has_sask(["saskatoon", "has", "bridges"]),False,'has_sask(["saskatoon", "has", "bridges"])')

myTests().main()
</textarea>
</div>
</div>
</section>
</section>
</section>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - <span class="section-number">4. </span>Turtle Drawing and Strings (L-Systems)' data-toggle="tooltip">
    <a href="../Strings/LSystems.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../Strings/LSystems.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - <span class="section-number">2. </span>List Methods' data-toggle="tooltip" >
    <a href="ListMethods.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="ListMethods.html" >Next Section - <span class="section-number">2. </span>List Methods</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.
    </p>
  </div>
</footer>


  </body>
</html>