<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>1. Using Lists &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="2. List Methods" href="ListMethods.html" />
    <link rel="prev" title="4. Turtle Drawing and Strings (L-Systems)" href="../Strings/LSystems.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">1. Using Lists</a><ul>
<li><a class="reference internal" href="#list-overview">1.1. List Overview</a></li>
<li><a class="reference internal" href="#list-values">1.2. List Values</a></li>
<li><a class="reference internal" href="#list-length">1.3. List Length</a><ul>
<li><a class="reference internal" href="#check-your-understanding">1.3.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#accessing-elements">1.4. Accessing Elements</a><ul>
<li><a class="reference internal" href="#id1">1.4.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#list-membership">1.5. List Membership</a><ul>
<li><a class="reference internal" href="#id2">1.5.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#list-slices">1.6. List Slices</a><ul>
<li><a class="reference internal" href="#id3">1.6.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#strings-are-immutable">1.7. Strings are Immutable</a></li>
<li><a class="reference internal" href="#lists-are-mutable">1.8. Lists are Mutable</a><ul>
<li><a class="reference internal" href="#id4">1.8.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#lists-and-for-loops">1.9. Lists and <code class="docutils literal"><span class="pre">for</span></code> loops</a></li>
<li><a class="reference internal" href="#practice-problems">1.10. Practice Problems</a><ul>
<li><a class="reference internal" href="#warmup">1.10.1. Warmup</a></li>
<li><a class="reference internal" href="#make-it-fun">1.10.2. Make it Fun</a></li>
<li><a class="reference internal" href="#sum-of-squares">1.10.3. Sum of Squares</a></li>
<li><a class="reference internal" href="#count-odd-numbers">1.10.4. Count Odd Numbers</a></li>
<li><a class="reference internal" href="#count-words-with-length-5">1.10.5. Count Words with Length 5</a></li>
<li><a class="reference internal" href="#look-for-saskatchewan">1.10.6. Look For “saskatchewan”</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 4. Turtle Drawing and Strings (L-Systems)" data-toggle="tooltip">
    <a href="../Strings/LSystems.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 2. List Methods' data-toggle="tooltip" >
    <a href="ListMethods.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="using-lists">
<h1>1. Using Lists<a class="headerlink" href="#using-lists" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>Lists overview.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide">
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to create reusable pieces of code.</li>
<li><strong>CS20-FP4</strong>  Investigate one-dimensional arrays and their applications.</li>
</ul>

    </div>
    <div class="section" id="list-overview">
<h2>1.1. List Overview<a class="headerlink" href="#list-overview" title="Permalink to this headline">¶</a></h2>
<p>We’ve been using lists throughout the course, so you probably already feel quite comfortable with them. Some of the following information will be a review for you, but we will spend a bit of time going into more depth than we have done up to now.</p>
<p>A <strong>list</strong> is a sequential collection of Python data values, where each value is identified by an
index. The values that make up a list are called its <strong>elements</strong>. Lists are
similar to strings, which are ordered collections of characters, except that the
elements of a list can have any type and for any one list, the items can be of different types.</p>
</div>
<div class="section" id="list-values">
<h2>1.2. List Values<a class="headerlink" href="#list-values" title="Permalink to this headline">¶</a></h2>
<p>There are several ways to create a new list.  The simplest is to enclose the
elements in square brackets ( <code class="docutils literal"><span class="pre">[</span></code> and <code class="docutils literal"><span class="pre">]</span></code>).</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="p">[</span><span class="mi">10</span><span class="p">,</span> <span class="mi">20</span><span class="p">,</span> <span class="mi">30</span><span class="p">,</span> <span class="mi">40</span><span class="p">]</span>
<span class="p">[</span><span class="s2">&quot;spam&quot;</span><span class="p">,</span> <span class="s2">&quot;bungee&quot;</span><span class="p">,</span> <span class="s2">&quot;swallow&quot;</span><span class="p">]</span>
</pre></div>
</div>
<p>The first example is a list of four integers. The second is a list of three
strings. As we said above, the elements of a list don’t have to be the same type.  The following
list contains a string, a float, an integer, and
another list.</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="p">[</span><span class="s2">&quot;hello&quot;</span><span class="p">,</span> <span class="mf">2.0</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="p">[</span><span class="mi">10</span><span class="p">,</span> <span class="mi">20</span><span class="p">]]</span>
</pre></div>
</div>
<p>A list within another list is said to be <strong>nested</strong> and the inner list is often called a <strong>sublist</strong>.
Finally, there is a special list that contains no elements. It is called the
empty list and is denoted <code class="docutils literal"><span class="pre">[]</span></code>.</p>
<p>As you would expect, we can also assign list values to variables and pass lists as parameters to functions.</p>

<div data-childcomponent="list_overview_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
vocabulary = ["iteration", "selection", "control"]
numbers = [17, 123]
empty = []
mixed_list = ["hello", 2.0, 5*2, [10, 20]]

print(numbers)
print(mixed_list)
new_list = [ numbers, vocabulary ]
print(new_list)
</textarea>
</div>
</div>
<div class="section" id="list-length">
<h2>1.3. List Length<a class="headerlink" href="#list-length" title="Permalink to this headline">¶</a></h2>
<p>As with strings, the function <code class="docutils literal"><span class="pre">len</span></code> returns the length of a list (the number
of items in the list).  However, since lists can have items which are themselves lists, it important to note
that <code class="docutils literal"><span class="pre">len</span></code> only returns the top-most length.  In other words, sublists are considered to be a single
item when counting the length of the list.</p>

<div data-childcomponent="list_overview_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
a_list =  ["hello", 2.0, 5, [10, 20]]
print(len(a_list))
print(len(['Hepburn', 1, ['Saskatoon', 'Regina', 'Prince Albert'], [1, 2, 3]]))
</textarea>
</div>
<div class="section" id="check-your-understanding">
<h3>1.3.1. Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_2_1">
            <p>lists-overview1: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question9_2_1_opt_a">4</li><li data-component="feedback" id="test_question9_2_1_opt_a">len returns the actual number of items in the list, not the maximum index value.</li>
            
            <li data-component="answer" data-correct id="test_question9_2_1_opt_b">5</li><li data-component="feedback" id="test_question9_2_1_opt_b">Yes, there are 5 items in this list.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_2_2">
            <p>lists-overview2: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">a_list</span><span class="p">))</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="test_question9_2_2_opt_a">7</li><li data-component="feedback" id="test_question9_2_2_opt_a">Yes, there are 7 items in this list even though two of them happen to also be lists.</li>
            
            <li data-component="answer"  id="test_question9_2_2_opt_b">8</li><li data-component="feedback" id="test_question9_2_2_opt_b">len returns the number of top level items in the list.  It does not count items in sublists.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="accessing-elements">
<h2>1.4. Accessing Elements<a class="headerlink" href="#accessing-elements" title="Permalink to this headline">¶</a></h2>
<p>The syntax for accessing the elements of a list is the same as the syntax for
accessing the characters of a string.  We use the index operator ( <code class="docutils literal"><span class="pre">[]</span></code> – not to
be confused with an empty list). The expression inside the brackets specifies
the index. Remember that the indices start at 0.  Any integer expression can be used
as an index and as with strings, negative index values will locate items from the right instead
of from the left.</p>

<div data-childcomponent="list_overview_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_3 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
numbers = [17, 123, 87, 34, 66, 8398, 44]
print(numbers[2])
print(numbers[9 - 8])
print(numbers[-2])
print(numbers[len(numbers) - 1])
</textarea>
</div>
<p>Also note that you can use the index operator <code class="docutils literal"><span class="pre">[]</span></code> more than once in a row to access things inside the list. For example, if you had a list of places in Saskatchewan, you could access both a place, as well as an individual character, as follows:</p>

<div data-childcomponent="list_overview_4" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_4 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
places = ["Saskatoon", "Regina", "Prince Albert", "Moose Jaw", "Yorkton", "Swift Current", "North Battleford"]
print(places[2])
print(places[2][1])
</textarea>
</div>
<p>The same syntax works if you are trying to access an element of a list inside a list as well. Consider the following:</p>

<div data-childcomponent="list_overview_5" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_5 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
places = [["Saskatoon", "Regina"], ["Prince Albert", "Moose Jaw"], ["Yorkton", "Swift Current", "North Battleford"]]
print(places[2])
print(places[2][1])
print(places[2][1][3])
</textarea>
</div>
<div class="section" id="id1">
<h3>1.4.1. Check Your Understanding<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_3_1">
            <p>lists-overview3: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question9_3_1_opt_a">[ ]</li><li data-component="feedback" id="test_question9_3_1_opt_a">The empty list is at index 4.</li>
            
            <li data-component="answer" data-correct id="test_question9_3_1_opt_b">3.14</li><li data-component="feedback" id="test_question9_3_1_opt_b">Yes, 3.14 is at index 5 since we start counting at 0 and sublists count as one item.</li>
            
            <li data-component="answer"  id="test_question9_3_1_opt_c">False</li><li data-component="feedback" id="test_question9_3_1_opt_c">False is at index 6.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_3_3">
            <p>lists-overview4: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">2</span><span class="p">][</span><span class="mi">0</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question9_3_3_opt_a">56</li><li data-component="feedback" id="test_question9_3_3_opt_a">Indexes start with 0, not 1.</li>
            
            <li data-component="answer" data-correct id="test_question9_3_3_opt_b">c</li><li data-component="feedback" id="test_question9_3_3_opt_b">Yes, the first character of the string at index 2 is c</li>
            
            <li data-component="answer"  id="test_question9_3_3_opt_c">cat</li><li data-component="feedback" id="test_question9_3_3_opt_c">cat is the item at index 2 but then we index into it further.</li>
            
            <li data-component="answer"  id="test_question9_3_3_opt_d">Error, you cannot have two index values unless you are using slicing.</li><li data-component="feedback" id="test_question9_3_3_opt_d">Using more than one index is fine.  You read it from left to right.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="list-membership">
<h2>1.5. List Membership<a class="headerlink" href="#list-membership" title="Permalink to this headline">¶</a></h2>
<p><code class="docutils literal"><span class="pre">in</span></code> and <code class="docutils literal"><span class="pre">not</span> <span class="pre">in</span></code> are boolean operators that test membership in a sequence. We
used them previously with strings and they also work here.</p>

<div data-childcomponent="list_overview_6" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_6 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = ["apple", "orange", "banana", "cherry"]

print("apple" in fruit)
print("pear" in fruit)
</textarea>
</div>
<div class="section" id="id2">
<h3>1.5.1. Check Your Understanding<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_4_1">
            <p>lists-overview5: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="mf">3.14</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="test_question9_4_1_opt_a">True</li><li data-component="feedback" id="test_question9_4_1_opt_a">Yes, 3.14 is an item in the list a_list.</li>
            
            <li data-component="answer"  id="test_question9_4_1_opt_b">False</li><li data-component="feedback" id="test_question9_4_1_opt_b">There are 7 items in the list, 3.14 is one of them.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_4_2">
            <p>lists-overview6: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="mi">57</span> <span class="ow">in</span> <span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question9_4_2_opt_a">True</li><li data-component="feedback" id="test_question9_4_2_opt_a">in returns True for top level items only.  57 is in a sublist.</li>
            
            <li data-component="answer" data-correct id="test_question9_4_2_opt_b">False</li><li data-component="feedback" id="test_question9_4_2_opt_b">Yes, 57 is not a top level item in a_list.  It is in a sublist.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="list-slices">
<h2>1.6. List Slices<a class="headerlink" href="#list-slices" title="Permalink to this headline">¶</a></h2>
<p>The slice operation we saw with strings also work on lists.  Remember that the first index is the starting point for the slice and the second number is one index past the end of the slice (up to but not including that element).  Recall also
that if you omit the first index (before the colon), the slice starts at the
beginning of the sequence. If you omit the second index, the slice goes to the
end of the sequence.</p>

<div data-childcomponent="list_overview_7" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_7 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
a_list = ['a', 'b', 'c', 'd', 'e', 'f']
print(a_list[1:3])
print(a_list[:4])
print(a_list[3:])
print(a_list[:])
</textarea>
</div>
<div class="section" id="id3">
<h3>1.6.1. Check Your Understanding<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_6_1">
            <p>lists-overview7: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">3</span><span class="p">,</span> <span class="mi">67</span><span class="p">,</span> <span class="s2">&quot;cat&quot;</span><span class="p">,</span> <span class="p">[</span><span class="mi">56</span><span class="p">,</span> <span class="mi">57</span><span class="p">,</span> <span class="s2">&quot;dog&quot;</span><span class="p">],</span> <span class="p">[</span> <span class="p">],</span> <span class="mf">3.14</span><span class="p">,</span> <span class="bp">False</span><span class="p">]</span>
<span class="k">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">[</span><span class="mi">4</span><span class="p">:])</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="test_question9_6_1_opt_a">[ [ ], 3.14, False]</li><li data-component="feedback" id="test_question9_6_1_opt_a">Yes, the slice starts at index 4 and goes up to and including the last item.</li>
            
            <li data-component="answer"  id="test_question9_6_1_opt_b">[ [ ], 3.14]</li><li data-component="feedback" id="test_question9_6_1_opt_b">By leaving out the upper bound on the slice, we go up to and including the last item.</li>
            
            <li data-component="answer"  id="test_question9_6_1_opt_c">[ [56, 57, "dog"], [ ], 3.14, False]</li><li data-component="feedback" id="test_question9_6_1_opt_c">Index values start at 0.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="strings-are-immutable">
<h2>1.7. Strings are Immutable<a class="headerlink" href="#strings-are-immutable" title="Permalink to this headline">¶</a></h2>
<p>One thing that makes strings different from Python lists is that
you are not allowed to modify the individual characters in a string.  It is tempting to use the <code class="docutils literal"><span class="pre">[]</span></code> operator on the left side of an assignment,
with the intention of changing a character in a string.  For example, in the following code, we would like to change the first letter of <code class="docutils literal"><span class="pre">greeting</span></code>.</p>

<div data-childcomponent="cg08_imm1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=cg08_imm1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
greeting = "Hello, world!"
greeting[0] = 'J'            # ERROR!
print(greeting)
</textarea>
</div>
<p>Instead of producing the output <code class="docutils literal"><span class="pre">Jello,</span> <span class="pre">world!</span></code>, this code produces the
runtime error <code class="docutils literal"><span class="pre">TypeError:</span> <span class="pre">'str'</span> <span class="pre">object</span> <span class="pre">does</span> <span class="pre">not</span> <span class="pre">support</span> <span class="pre">item</span> <span class="pre">assignment</span></code>.</p>
<p>Strings are <strong>immutable</strong>, which means you cannot change an existing string. The
best you can do is create a new string that is a variation on the original.</p>

<div data-childcomponent="ch08_imm2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ch08_imm2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
greeting = "Hello, world!"
new_greeting = 'J' + greeting[1:]
print(new_greeting)
print(greeting)            # same as it was
</textarea>
</div>
<p>The solution here is to concatenate a new first letter onto a slice of
<code class="docutils literal"><span class="pre">greeting</span></code>. This operation has no effect on the original string.</p>
</div>
<div class="section" id="lists-are-mutable">
<h2>1.8. Lists are Mutable<a class="headerlink" href="#lists-are-mutable" title="Permalink to this headline">¶</a></h2>
<p>Unlike strings, lists are <strong>mutable</strong>.  This means we can change an item in a list by accessing
it directly as part of the assignment statement. Using the indexing operator (square brackets) on the left side of an assignment, we can
update one of the list items.</p>

<div data-childcomponent="ch09_7" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ch09_7 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = ["banana", "apple", "cherry"]
print(fruit)

fruit[0] = "pear"
fruit[-1] = "orange"
print(fruit)
</textarea>
</div>
<p>An assignment to an element of a list is called <strong>item assignment</strong>. Item
assignment does not work for strings.  Recall that strings are immutable.</p>
<div class="section" id="id4">
<h3>1.8.1. Check Your Understanding<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="test_question9_7_1">
            <p>lists-overview8: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">a_list</span> <span class="o">=</span> <span class="p">[</span><span class="mi">4</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="n">a_list</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">=</span> <span class="bp">True</span>
<span class="k">print</span><span class="p">(</span><span class="n">a_list</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="test_question9_7_1_opt_a">[4, 2, True, 8, 6, 5]</li><li data-component="feedback" id="test_question9_7_1_opt_a">Item assignment does not insert the new item into the list.</li>
            
            <li data-component="answer" data-correct id="test_question9_7_1_opt_b">[4, 2, True, 6, 5]</li><li data-component="feedback" id="test_question9_7_1_opt_b">Yes, the value True is placed in the list at index 2.  It replaces 8.</li>
            
            <li data-component="answer"  id="test_question9_7_1_opt_c">Error, it is illegal to assign</li><li data-component="feedback" id="test_question9_7_1_opt_c">Item assignment is allowed with lists.  Lists are mutable.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="lists-and-for-loops">
<h2>1.9. Lists and <code class="docutils literal"><span class="pre">for</span></code> loops<a class="headerlink" href="#lists-and-for-loops" title="Permalink to this headline">¶</a></h2>
<p>It is also possible to perform <strong>list traversal</strong> using iteration by item as well as iteration by index.</p>

<div data-childcomponent="chp09_03a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=chp09_03a data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruits = ["apple", "orange", "banana", "cherry"]

# looping by item
for some_fruit in fruits:
    print(some_fruit)
</textarea>
</div>
<p>It almost reads like natural language: For (every) fruit in (the list of) fruits,
print (the name of the) fruit.</p>
<p>We can also use the indices to access the items in an iterative fashion.</p>

<div data-childcomponent="chp09_03b" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=chp09_03b data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruits = ["apple", "orange", "banana", "cherry"]

# looping by index
for position in range(len(fruits)):
    print(fruits[position])
</textarea>
</div>
<p>In this example, each time through the loop, the variable <code class="docutils literal"><span class="pre">position</span></code> is used as an index into the
list, printing the <code class="docutils literal"><span class="pre">position</span></code>-eth element. Note that we used <code class="docutils literal"><span class="pre">len</span></code> as the upper bound on the range
so that we can iterate correctly no matter how many items are in the list.</p>
<p>Since lists are mutable, it is often desirable to traverse a list, modifying
each of its elements as you go. The following code squares all the numbers from <code class="docutils literal"><span class="pre">1</span></code> to
<code class="docutils literal"><span class="pre">5</span></code> using iteration by position.</p>

<div data-childcomponent="chp09_for4" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=chp09_for4 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
numbers = [1, 2, 3, 4, 5]
print(numbers)

for i in range(len(numbers)):
    numbers[i] = numbers[i] ** 2

print(numbers)
</textarea>
</div>
<p>Take a moment to think about <code class="docutils literal"><span class="pre">range(len(numbers))</span></code> until you understand how
it works. We are interested here in both the <em>value</em> and its <em>index</em> within the
list, so that we can assign a new value to it.</p>
</div>
<div class="section" id="practice-problems">
<h2>1.10. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<div class="section" id="warmup">
<h3>1.10.1. Warmup<a class="headerlink" href="#warmup" title="Permalink to this headline">¶</a></h3>
<p>Create a list called <code class="docutils literal"><span class="pre">my_list</span></code> with the following six items: <code class="docutils literal"><span class="pre">76,</span> <span class="pre">92.3,</span> <span class="pre">&quot;hello&quot;,</span> <span class="pre">True,</span> <span class="pre">4,</span> <span class="pre">76</span></code>. Print out the list to make sure it is doing what you intend.</p>

<div data-childcomponent="list_overview_practice_problem_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=list_overview_practice_problem_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
# your code goes here!
</textarea>
</div>
</div>
<div class="section" id="make-it-fun">
<h3>1.10.2. Make it Fun<a class="headerlink" href="#make-it-fun" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="list_overview_practice_problem_6" class="runestone explainer ac_section alert alert-warning">
<p>Write a function to replace the first element of a list with the word <code class="docutils literal"><span class="pre">&quot;fun&quot;</span></code>, and then return the altered list. You can assume that the list will always have at least one element.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">make_fun([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;science&quot;]</span></code></p>
<p><code class="docutils literal"><span class="pre">make_fun([&quot;sask&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;]</span></code></p>
<p><code class="docutils literal"><span class="pre">make_fun([&quot;we&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;])</span> <span class="pre">→</span> <span class="pre">[&quot;fun&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;]</span></code></p>

<textarea data-component="activecode" id=list_overview_practice_problem_6 data-lang="python"    data-timelimit=25000   data-audio=''     >
def make_fun(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(make_fun(["computer", "science"]),["fun", "science"],'make_fun(["computer", "science"])')
      self.assertEqual(make_fun(["sask", "is", "nice"]),["fun", "is", "nice"],'make_fun(["sask", "is", "nice"])')
      self.assertEqual(make_fun(["we", "live", "in", "saskatchewan"]),["fun", "live", "in", "saskatchewan"],'make_fun(["we", "live", "in", "saskatchewan"])')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="sum-of-squares">
<h3>1.10.3. Sum of Squares<a class="headerlink" href="#sum-of-squares" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="list_overview_practice_problem_2" class="runestone explainer ac_section alert alert-warning">
<p>Write a function <code class="docutils literal"><span class="pre">sum_of_squares(a_list)</span></code> that computes the sum of the squares of the numbers in the list <code class="docutils literal"><span class="pre">a_list</span></code>.  For example, <code class="docutils literal"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 29, which is the sum of 4+9+16.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">sum_of_squares([2,3,4])</span> <span class="pre">→</span> <span class="pre">29</span></code></p>
<p><code class="docutils literal"><span class="pre">sum_of_squares([0,1,-1])</span> <span class="pre">→</span> <span class="pre">2</span></code></p>
<p><code class="docutils literal"><span class="pre">sum_of_squares([5,12])</span> <span class="pre">→</span> <span class="pre">169</span></code></p>

<textarea data-component="activecode" id=list_overview_practice_problem_2 data-lang="python"    data-timelimit=25000   data-audio=''     >
def sum_of_squares(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(sum_of_squares([2,3,4]),29,"sum_of_squares([2,3,4])")
      self.assertEqual(sum_of_squares([0,1,-1]),2,"sum_of_squares([0,1,-1])")
      self.assertEqual(sum_of_squares([5,12]),169,"sum_of_squares([5,12])")
      self.assertEqual(sum_of_squares([5,12,14]),365,"sum_of_squares([5,12,14])")
      self.assertEqual(sum_of_squares([5]),25,"sum_of_squares([5])")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="count-odd-numbers">
<h3>1.10.4. Count Odd Numbers<a class="headerlink" href="#count-odd-numbers" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="list_overview_practice_problem_3" class="runestone explainer ac_section alert alert-warning">
<p>Write a function to count how many odd numbers are in a list.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">count_odds([1,3,5,7,9])</span> <span class="pre">→</span> <span class="pre">5</span></code></p>
<p><code class="docutils literal"><span class="pre">count_odds([1,2,3,4,5])</span> <span class="pre">→</span> <span class="pre">3</span></code></p>
<p><code class="docutils literal"><span class="pre">count_odds([2,4,6,8,10])</span> <span class="pre">→</span> <span class="pre">0</span></code></p>

<textarea data-component="activecode" id=list_overview_practice_problem_3 data-lang="python"    data-timelimit=25000   data-audio=''     >
def count_odds(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(count_odds([1,3,5,7,9]),5,"count_odds([1,3,5,7,9])")
      self.assertEqual(count_odds([1,2,3,4,5]),3,"count_odds([-1,-2,-3,-4,-5])")
      self.assertEqual(count_odds([2,4,6,8,10]),0,"count_odds([2,4,6,8,10])")
      self.assertEqual(count_odds([0,-1,12,-33]),2,"count_odds([0,-1,12,-33])")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="count-words-with-length-5">
<h3>1.10.5. Count Words with Length 5<a class="headerlink" href="#count-words-with-length-5" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="list_overview_practice_problem_4" class="runestone explainer ac_section alert alert-warning">
<p>Write a function to count how many words in a list have a length of 5.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">count_words([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">0</span></code></p>
<p><code class="docutils literal"><span class="pre">count_words([&quot;nobody&quot;,</span> <span class="pre">&quot;knows&quot;])</span> <span class="pre">→</span> <span class="pre">1</span></code></p>
<p><code class="docutils literal"><span class="pre">count_words([&quot;think&quot;,</span> <span class="pre">&quot;pizza&quot;,</span> <span class="pre">&quot;quick&quot;])</span> <span class="pre">→</span> <span class="pre">3</span></code></p>

<textarea data-component="activecode" id=list_overview_practice_problem_4 data-lang="python"    data-timelimit=25000   data-audio=''     >
def count_words(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(count_words(["computer", "science"]),0,'count_words(["computer", "science"])')
      self.assertEqual(count_words(["nobody", "knows"]),1,'count_words(["nobody", "knows"])')
      self.assertEqual(count_words(["think", "pizza", "quick"]),3,'count_words(["think", "pizza", "quick"])')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="look-for-saskatchewan">
<h3>1.10.6. Look For “saskatchewan”<a class="headerlink" href="#look-for-saskatchewan" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="list_overview_practice_problem_5" class="runestone explainer ac_section alert alert-warning">
<p>Write a function to check if the word “saskatchewan” shows up somewhere in a list. If it does, return <code class="docutils literal"><span class="pre">True</span></code>. If it doesn’t, return <code class="docutils literal"><span class="pre">False</span></code>.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">has_sask([&quot;computer&quot;,</span> <span class="pre">&quot;science&quot;])</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal"><span class="pre">has_sask([&quot;sask&quot;,</span> <span class="pre">&quot;is&quot;,</span> <span class="pre">&quot;nice&quot;])</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal"><span class="pre">has_sask([&quot;we&quot;,</span> <span class="pre">&quot;live&quot;,</span> <span class="pre">&quot;in&quot;,</span> <span class="pre">&quot;saskatchewan&quot;])</span> <span class="pre">→</span> <span class="pre">True</span></code></p>

<textarea data-component="activecode" id=list_overview_practice_problem_5 data-lang="python"    data-timelimit=25000   data-audio=''     >
def has_sask(a_list):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(has_sask(["computer", "science"]),False,'has_sask(["computer", "science"])')
      self.assertEqual(has_sask(["sask", "is", "nice"]),False,'has_sask(["sask", "is", "nice"])')
      self.assertEqual(has_sask(["we", "live", "in", "saskatchewan"]),True,'has_sask(["we", "live", "in", "saskatchewan"])')
      self.assertEqual(has_sask(["saskatchewan", "roughriders"]),True,'has_sask(["saskatchewan", "roughriders"])')
      self.assertEqual(has_sask(["saskatoon", "has", "bridges"]),False,'has_sask(["saskatoon", "has", "bridges"])')

myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 4. Turtle Drawing and Strings (L-Systems)' data-toggle="tooltip">
    <a href="../Strings/LSystems.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../Strings/LSystems.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 2. List Methods' data-toggle="tooltip" >
    <a href="ListMethods.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="ListMethods.html" >Next Section - 2. List Methods</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>