<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>9.22. Exercises &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = false;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">9.22. Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>9.22. Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic">
<li><div id='q1' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>What is the result of each of the following:</p>
<ol class="loweralpha simple">
<li>‘Python’[1]</li>
<li>“Strings are sequences of characters.”[5]</li>
<li>len(“wonderful”)</li>
<li>‘Mystery’[:4]</li>
<li>‘p’ in ‘Pineapple’</li>
<li>‘apple’ in ‘Pineapple’</li>
<li>‘pear’ not in ‘Pineapple’</li>
<li>‘apple’ &gt; ‘pineapple’</li>
<li>‘pineapple’ &lt; ‘Peach’</li>
</ol>
</div><div data-component="tab" data-tabname="Answer" >
<ol class="loweralpha simple">
<li>‘Python’[1] = ‘y’</li>
<li>‘Strings are sequences of characters.’[5] = ‘g’</li>
<li>len(‘wonderful’) = 9</li>
<li>‘Mystery’[:4] = ‘Myst’</li>
<li>‘p’ in ‘Pineapple’ = True</li>
<li>‘apple’ in ‘Pineapple’ = True</li>
<li>‘pear’ not in ‘Pineapple’ = True</li>
<li>‘apple’ &gt; ‘pineapple’ = False</li>
<li>‘pineapple’ &lt; ‘Peach’ = False</li>
</ol>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_dc2457710a924d9283b12f42a31d2b27_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course;
                    this.page.url = 'http://disqus_dc2457710a924d9283b12f42a31d2b27_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course;
            disqus_url = 'http://disqus_dc2457710a924d9283b12f42a31d2b27_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_dc2457710a924d9283b12f42a31d2b27" onclick="disqus_dc2457710a924d9283b12f42a31d2b27_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_dc2457710a924d9283b12f42a31d2b27']").attr('data-disqus-identifier', 'disqus_dc2457710a924d9283b12f42a31d2b27_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_2" class="runestone explainer ac_section alert alert-warning">
<p>In Robert McCloskey’s
book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack,
Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">prefixes</span> <span class="o">=</span> <span class="s2">&quot;JKLMNOPQ&quot;</span>
<span class="n">suffix</span> <span class="o">=</span> <span class="s2">&quot;ack&quot;</span>

<span class="k">for</span> <span class="n">p</span> <span class="ow">in</span> <span class="n">prefixes</span><span class="p">:</span>
    <span class="k">print</span><span class="p">(</span><span class="n">p</span> <span class="o">+</span> <span class="n">suffix</span><span class="p">)</span>
</pre></div>
</div>
<p>Of course, that’s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>

<textarea data-component="activecode" id=ex_8_2 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>
</li>
<li><div id='q3' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_3" class="runestone explainer ac_section alert alert-warning">
<p>Assign to a variable in your program a triple-quoted string that contains
your favorite paragraph of text - perhaps a poem, a speech, instructions
to bake a cake, some inspirational verses, etc.</p>
<p>Write a function that counts the number of alphabetic characters (a through z, or A through Z) in your text and then keeps track of how many are the letter ‘e’.  Your function should print an analysis of the text like this:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">Your</span> <span class="n">text</span> <span class="n">contains</span> <span class="mi">243</span> <span class="n">alphabetic</span> <span class="n">characters</span><span class="p">,</span> <span class="n">of</span> <span class="n">which</span> <span class="mi">109</span> <span class="p">(</span><span class="mf">44.8</span><span class="o">%</span><span class="p">)</span> <span class="n">are</span> <span class="s1">&#39;e&#39;</span><span class="o">.</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_8_3 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def count(p):
    # your code here
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q3_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q3_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def count(p):
    lows = "abcdefghijklmnopqrstuvwxyz"
    ups =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"

    numberOfe = 0
    totalChars = 0
    for achar in p:
        if achar in lows or achar in ups:
            totalChars = totalChars + 1
            if achar == 'e':
                numberOfe = numberOfe + 1

    percent_with_e = (numberOfe / totalChars) * 100
    print("Your text contains", totalChars, "alphabetic characters of which", numberOfe, "(", percent_with_e, "%)", "are 'e'.")


p = '''
"If the automobile had followed the same development cycle as the computer, a
Rolls-Royce would today cost $100, get a million miles per gallon, and explode
once a year, killing everyone inside."
-Robert Cringely
'''

count(p)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_bf88b1c4616d43f289c798b56a43b01c_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course;
                    this.page.url = 'http://disqus_bf88b1c4616d43f289c798b56a43b01c_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course;
            disqus_url = 'http://disqus_bf88b1c4616d43f289c798b56a43b01c_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_bf88b1c4616d43f289c798b56a43b01c" onclick="disqus_bf88b1c4616d43f289c798b56a43b01c_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_bf88b1c4616d43f289c798b56a43b01c']").attr('data-disqus-identifier', 'disqus_bf88b1c4616d43f289c798b56a43b01c_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_4" class="runestone explainer ac_section alert alert-warning">
<p>Print out a neatly formatted multiplication table, up to 12 x 12.</p>

<textarea data-component="activecode" id=ex_8_4 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>
</li>
<li><div id='q5' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_7_10" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that will return the number of digits in an integer.</p>

<textarea data-component="activecode" id=ex_7_10 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def numDigits(n):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
      self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
      self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
      self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")



myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q5_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q5_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def numDigits(n):
    n_str = str(n)
    return len(n_str)


print(numDigits(50))
print(numDigits(20000))
print(numDigits(1))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
                    this.page.url = 'http://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course;
            disqus_url = 'http://disqus_bfd6f74a183c4682b29c72c4411200fb_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_bfd6f74a183c4682b29c72c4411200fb" onclick="disqus_bfd6f74a183c4682b29c72c4411200fb_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_bfd6f74a183c4682b29c72c4411200fb']").attr('data-disqus-identifier', 'disqus_bfd6f74a183c4682b29c72c4411200fb_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_5" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that reverses its string argument.</p>

<textarea data-component="activecode" id=ex_8_5 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def reverse(astring):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
      self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
      self.assertEqual(reverse(""),"","Tested reverse on input of ''")




myTests().main()
</textarea>
</div>
</li>
<li><div id='q7' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function that mirrors its string argument,
generating a string containing the original string and the string backwards.</p>

<div data-childcomponent="ex_8_6" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ex_8_6 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def mirror(mystr):
    # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(mirror("good"),"gooddoog","Tested mirror on input of 'good'")
        self.assertEqual(mirror("Python"),"PythonnohtyP","Tested mirror on input of 'Python'")
        self.assertEqual(mirror(""),"","Tested mirror on input of ''")
        self.assertEqual(mirror("a"),"aa","Tested mirror on input of 'a'")


myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q7_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q7_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
from test import testEqual

def reverse(mystr):
    reversed = ''
    for char in mystr:
        reversed = char + reversed
    return reversed

def mirror(mystr):
    return mystr + reverse(mystr)

testEqual(mirror('good'), 'gooddoog')
testEqual(mirror('Python'), 'PythonnohtyP')
testEqual(mirror(''), '')
testEqual(mirror('a'), 'aa')
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_70b7ac515456497c952a2de5caa27ab9_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
                    this.page.url = 'http://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course;
            disqus_url = 'http://disqus_70b7ac515456497c952a2de5caa27ab9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_70b7ac515456497c952a2de5caa27ab9" onclick="disqus_70b7ac515456497c952a2de5caa27ab9_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_70b7ac515456497c952a2de5caa27ab9']").attr('data-disqus-identifier', 'disqus_70b7ac515456497c952a2de5caa27ab9_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_7" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that removes all occurrences of a given letter from a string.</p>

<textarea data-component="activecode" id=ex_8_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def remove_letter(theLetter, theString):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
      self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
      self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")



myTests().main()
</textarea>
</div>
</li>
<li><div id='q9' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_8" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that recognizes palindromes. (Hint: use your <code class="docutils literal"><span class="pre">reverse</span></code> function to make this easy!).</p>

<textarea data-component="activecode" id=ex_8_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def is_palindrome(myStr):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_palindrome("abba"),True,"Tested is_palindrome on input of 'abba'")
        self.assertEqual(is_palindrome("abab"),False,"Tested is_palindrome on input of 'abab'")
        self.assertEqual(is_palindrome("straw warts"),True,"Tested is_palindrome on input of 'straw warts'")
        self.assertEqual(is_palindrome("a"),True,"Tested is_palindrome on input of 'a'")
        self.assertEqual(is_palindrome(""),True,"Tested is_palindrome on input of ''")




myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q9_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q9_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
from test import testEqual

def reverse(mystr):
    reversed = ''
    for char in mystr:
        reversed = char + reversed
    return reversed

def is_palindrome(myStr):
    if myStr in reverse(myStr):
        return True
    else:
        return False

testEqual(is_palindrome('abba'), True)
testEqual(is_palindrome('abab'), False)
testEqual(is_palindrome('straw warts'), True)
testEqual(is_palindrome('a'), True)
testEqual(is_palindrome(''), True)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_605923545bb849f7b8d41bbf823518e9_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course;
                    this.page.url = 'http://disqus_605923545bb849f7b8d41bbf823518e9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course;
            disqus_url = 'http://disqus_605923545bb849f7b8d41bbf823518e9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_605923545bb849f7b8d41bbf823518e9" onclick="disqus_605923545bb849f7b8d41bbf823518e9_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_605923545bb849f7b8d41bbf823518e9']").attr('data-disqus-identifier', 'disqus_605923545bb849f7b8d41bbf823518e9_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Write a function that counts how many non-overlapping occurences of a substring appear in a string.</p>

<div data-childcomponent="ex_8_9" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ex_8_9 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def count(substr,theStr):
    # your code here


====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(count("is","Mississippi"),2,"Tested count on inputs of 'is' and 'Mississippi'")
          self.assertEqual(count("an","banana"),2,"Tested count on inputs of 'an' and 'banana'")
          self.assertEqual(count("ana","banana"),2,"Tested count on inputs of 'ana' and 'banana'")
          self.assertEqual(count("nana","banana"),1,"Tested count on inputs of 'nana' and 'banana'")
          self.assertEqual(count("nanan","banana"),0,"Tested count on inputs of 'nanan' and 'banana'")
          self.assertEqual(count("aaa","aaaaaa"),4,"Tested count on input of 'aaa' and 'aaaaaa'")




myTests().main()
</textarea>
</div>
</li>
<li><div id='q11' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_10" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that removes the first occurrence of a string from another string.</p>

<textarea data-component="activecode" id=ex_8_10 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
def remove(substr,theStr):
    # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(remove("an","banana"),"bana","Tested remove on inputs of 'an' and 'banana'")
            self.assertEqual(remove("cyc","bicycle"),"bile","Tested remove on inputs of 'cyc' and 'bicycle'")
            self.assertEqual(remove("iss","Mississippi"),"Missippi","Tested remove on inputs of 'iss' and 'Mississippi'")
            self.assertEqual(remove("egg","bicycle"),"bicycle","Tested remove on inputs of 'egg' and 'bicycle'")



myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q11_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q11_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
from test import testEqual

def remove(substr,theStr):
    index = theStr.find(substr)
    if index < 0: # substr doesn't exist in theStr
        return theStr
    return_str = theStr[:index] + theStr[index+len(substr):]
    return return_str

testEqual(remove('an', 'banana'), 'bana')
testEqual(remove('cyc', 'bicycle'), 'bile')
testEqual(remove('iss', 'Mississippi'), 'Missippi')
testEqual(remove('egg', 'bicycle'), 'bicycle')
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_2f2772134b604a6498748138542d312d_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course;
                    this.page.url = 'http://disqus_2f2772134b604a6498748138542d312d_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course;
            disqus_url = 'http://disqus_2f2772134b604a6498748138542d312d_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_2f2772134b604a6498748138542d312d" onclick="disqus_2f2772134b604a6498748138542d312d_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_2f2772134b604a6498748138542d312d']").attr('data-disqus-identifier', 'disqus_2f2772134b604a6498748138542d312d_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_11" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that removes all occurrences of a string from another string.</p>

<textarea data-component="activecode" id=ex_8_11 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def remove_all(substr,theStr):
    # your code here



====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_all("an","banana"),"ba","Tested remove_all on inputs of 'an' and 'banana'")
      self.assertEqual(remove_all("cyc","bicycle"),"bile","Tested remove_all on inputs of 'cyc' and 'bicycle'")
      self.assertEqual(remove_all("iss","Mississippi"),"Mippi","Tested remove_all on inputs of 'iss' and 'Mississippi'")
      self.assertEqual(remove_all("eggs","bicycle"),"bicycle","Tested remove_all on inputs of 'eggs' and 'bicycle'")



myTests().main()
</textarea>
</div>
</li>
<li><div id='q13' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_12" class="runestone explainer ac_section alert alert-warning">
<p>Here is another interesting L-System called a Hilbert curve.  Use 90 degrees:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">L</span>
<span class="n">L</span> <span class="o">-&gt;</span> <span class="o">+</span><span class="n">RF</span><span class="o">-</span><span class="n">LFL</span><span class="o">-</span><span class="n">FR</span><span class="o">+</span>
<span class="n">R</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">LF</span><span class="o">+</span><span class="n">RFR</span><span class="o">+</span><span class="n">FL</span><span class="o">-</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_8_12 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q13_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q13_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'L':
        newstr = '+RF-LFL-FR+'   # Rule 1
    elif ch == 'R':
        newstr = '-LF+RFR+FL-'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(4, "L")  # create the string
    print(inst)
    t = turtle.Turtle()           # create the turtle
    wn = turtle.Screen()

    t.up()
    t.back(200)
    t.down()
    t.speed(9)
    drawLsystem(t, inst, 90, 5)   # draw the picture
                                  # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_ab823200fac64461a9e88f53b75f5795_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course;
                    this.page.url = 'http://disqus_ab823200fac64461a9e88f53b75f5795_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course;
            disqus_url = 'http://disqus_ab823200fac64461a9e88f53b75f5795_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_ab823200fac64461a9e88f53b75f5795" onclick="disqus_ab823200fac64461a9e88f53b75f5795_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_ab823200fac64461a9e88f53b75f5795']").attr('data-disqus-identifier', 'disqus_ab823200fac64461a9e88f53b75f5795_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_13" class="runestone explainer ac_section alert alert-warning">
<p>Here is a dragon curve.  Use 90 degrees.:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">FX</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">X</span><span class="o">+</span><span class="n">YF</span><span class="o">+</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">FX</span><span class="o">-</span><span class="n">Y</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_8_13 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</li>
<li><div id='q15' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_14" class="runestone explainer ac_section alert alert-warning">
<p>Here is something called an arrowhead curve.  Use 60 degrees.:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">YF</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">YF</span><span class="o">+</span><span class="n">XF</span><span class="o">+</span><span class="n">Y</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="n">XF</span><span class="o">-</span><span class="n">YF</span><span class="o">-</span><span class="n">X</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_8_14 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q15_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q15_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'X':
        newstr = 'YF+XF+Y'   # Rule 1
    elif ch == 'Y':
        newstr = 'XF-YF-X'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(5, "YF")  # create the string
    print(inst)
    t = turtle.Turtle()            # create the turtle
    wn = turtle.Screen()

    t.speed(9)
    drawLsystem(t, inst, 60, 5)    # draw the picture
                                   # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_9b2dfba083a64d5c894f873af2e93a1b_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course;
                    this.page.url = 'http://disqus_9b2dfba083a64d5c894f873af2e93a1b_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course;
            disqus_url = 'http://disqus_9b2dfba083a64d5c894f873af2e93a1b_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_9b2dfba083a64d5c894f873af2e93a1b" onclick="disqus_9b2dfba083a64d5c894f873af2e93a1b_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_9b2dfba083a64d5c894f873af2e93a1b']").attr('data-disqus-identifier', 'disqus_9b2dfba083a64d5c894f873af2e93a1b_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_15" class="runestone explainer ac_section alert alert-warning">
<p>Try the Peano-Gosper curve.  Use 60 degrees.:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">FX</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="n">X</span><span class="o">+</span><span class="n">YF</span><span class="o">++</span><span class="n">YF</span><span class="o">-</span><span class="n">FX</span><span class="o">--</span><span class="n">FXFX</span><span class="o">-</span><span class="n">YF</span><span class="o">+</span>
<span class="n">Y</span> <span class="o">-&gt;</span> <span class="o">-</span><span class="n">FX</span><span class="o">+</span><span class="n">YFYF</span><span class="o">++</span><span class="n">YF</span><span class="o">+</span><span class="n">FX</span><span class="o">--</span><span class="n">FX</span><span class="o">-</span><span class="n">Y</span>
</pre></div>
</div>

<textarea data-component="activecode" id=ex_8_15 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</li>
<li><div id='q17' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_16" class="runestone explainer ac_section alert alert-warning">
<blockquote>
<div><p>The Sierpinski Triangle.  Use 60 degrees.:</p>
<div class="highlight-default"><div class="highlight"><pre><span></span><span class="n">FXF</span><span class="o">--</span><span class="n">FF</span><span class="o">--</span><span class="n">FF</span>
<span class="n">F</span> <span class="o">-&gt;</span> <span class="n">FF</span>
<span class="n">X</span> <span class="o">-&gt;</span> <span class="o">--</span><span class="n">FXF</span><span class="o">++</span><span class="n">FXF</span><span class="o">++</span><span class="n">FXF</span><span class="o">--</span>
</pre></div>
</div>
</div></blockquote>

<textarea data-component="activecode" id=ex_8_16 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q17_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q17_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'F':
        newstr = 'FF'   # Rule 1
    elif ch == 'X':
        newstr = '--FXF++FXF++FXF--'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)

def main():
    inst = createLSystem(5, "FXF--FF--FF")   # create the string
    print(inst)
    t = turtle.Turtle()           # create the turtle
    wn = turtle.Screen()
    t.up()
    t.back(200)
    t.right(90)
    t.forward(100)
    t.left(90)
    t.down()
    t.speed(9)

    drawLsystem(t, inst, 60, 5)   # draw the picture
                                  # angle 90, segment length 5
    wn.exitonclick()

main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_912a5f19d3964dc2af7a067dcd832c7e_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course;
                    this.page.url = 'http://disqus_912a5f19d3964dc2af7a067dcd832c7e_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course;
            disqus_url = 'http://disqus_912a5f19d3964dc2af7a067dcd832c7e_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_912a5f19d3964dc2af7a067dcd832c7e" onclick="disqus_912a5f19d3964dc2af7a067dcd832c7e_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_912a5f19d3964dc2af7a067dcd832c7e']").attr('data-disqus-identifier', 'disqus_912a5f19d3964dc2af7a067dcd832c7e_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_17" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that implements a substitution cipher.  In a substitution
cipher one letter is substituted for another to garble the message.  For
example A -&gt; Q, B -&gt; T, C -&gt; G etc.  your function should take two
parameters, the message you want to encrypt, and a string that represents
the mapping of the 26 letters in the alphabet.  Your function should
return a string that is the encrypted version of the message.</p>

<textarea data-component="activecode" id=ex_8_17 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>
</li>
<li><div id='q19' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_18" class="runestone explainer ac_section alert alert-warning">
<p>Write a function that decrypts the message from the previous exercise.  It
should also take two parameters.  The encrypted message,
and the mixed up alphabet.  The function should return a string that is
the same as the original unencrypted message.</p>

<textarea data-component="activecode" id=ex_8_18 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q19_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q19_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def encrypt(message, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    encrypted = ''
    for char in message:
        if char == ' ':
            encrypted = encrypted + ' '
        else:
            pos = alphabet.index(char)
            encrypted = encrypted + cipher[pos]
    return encrypted

def decrypt(encrypted, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    decrypted = ''
    for char in encrypted:
        if char == ' ':
            decrypted = decrypted + ' '
        else:
            pos = cipher.index(char)
            decrypted = decrypted + alphabet[pos]
    return decrypted


cipher = "badcfehgjilknmporqtsvuxwzy"

encrypted = encrypt('hello world', cipher)
print encrypted

decrypted = decrypt(encrypted, cipher)
print(decrypted)
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_d7d1ca00bfff4e7bacf886386fb3302e_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course;
                    this.page.url = 'http://disqus_d7d1ca00bfff4e7bacf886386fb3302e_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course;
            disqus_url = 'http://disqus_d7d1ca00bfff4e7bacf886386fb3302e_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_d7d1ca00bfff4e7bacf886386fb3302e" onclick="disqus_d7d1ca00bfff4e7bacf886386fb3302e_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_d7d1ca00bfff4e7bacf886386fb3302e']").attr('data-disqus-identifier', 'disqus_d7d1ca00bfff4e7bacf886386fb3302e_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_19" class="runestone explainer ac_section alert alert-warning">
<p>Write a function called  <code class="docutils literal"><span class="pre">remove_dups</span></code> that takes a string and creates a new string by only adding those characters that are not already present.  In other words,
there will never be a duplicate letter added to the new string.</p>

<textarea data-component="activecode" id=ex_8_19 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def remove_dups(astring):
    # your code here


print(remove_dups("mississippi"))   #should print misp

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_dups("pooh"),"poh","Tested remove_dups on string 'pooh'")
      self.assertEqual(remove_dups("mississippi"),"misp","Tested remove_dups on string 'mississippi'")
      self.assertEqual(remove_dups("potato"),"pota","Tested remove_dups on string 'potato'")
      self.assertEqual(remove_dups("bookkeeper"),"bokepr","Tested remove_dups on string 'bookkeeper'")
      self.assertEqual(remove_dups("oo"),"o","Tested remove_dups on string 'oo'")

myTests().main()
</textarea>
</div>
</li>
<li><div id='q21' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ex_8_20" class="runestone explainer ac_section alert alert-warning">
<p>Write a function called <code class="docutils literal"><span class="pre">rot13</span></code> that uses the Caesar cipher to encrypt a message.
The Caesar cipher works like a substitution cipher but each character is replaced
by the character 13 characters to ‘its right’ in the alphabet.  So for example
the letter a becomes the letter n.  If a letter is past the middle of the alphabet
then the counting wraps around to the letter a again, so n becomes a, o becomes b
and so on.  <em>Hint:</em> Whenever you talk about things wrapping around its a good idea
to think of modulo arithmetic.</p>

<textarea data-component="activecode" id=ex_8_20 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def rot13(mess):
    # Your code here

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('Since rot13 is symmetric you should see this message')))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="q21_answer" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=q21_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def rot13(mess):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted = ''
    for char in mess:
        if char == ' ':
            encrypted = encrypted + ' '
        else:
            rotated_index = alphabet.index(char) + 13
            if rotated_index < 26:
                encrypted = encrypted + alphabet[rotated_index]
            else:
                encrypted = encrypted + alphabet[rotated_index % 26]
    return encrypted

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('since rot thirteen is symmetric you should see this message')))
</textarea>
</div>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_49e1151bb7864a3287a6b6ae1c84db16_disqus(source) {
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course;
                    this.page.url = 'http://disqus_49e1151bb7864a3287a6b6ae1c84db16_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';
            disqus_identifier = 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course;
            disqus_url = 'http://disqus_49e1151bb7864a3287a6b6ae1c84db16_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_49e1151bb7864a3287a6b6ae1c84db16" onclick="disqus_49e1151bb7864a3287a6b6ae1c84db16_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_49e1151bb7864a3287a6b6ae1c84db16']").attr('data-disqus-identifier', 'disqus_49e1151bb7864a3287a6b6ae1c84db16_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li>
<div data-childcomponent="ex_8_22" class="runestone explainer ac_section alert alert-warning">
<p>Modify this code so it prints each subtotal, the total cost, and average price
to exactly two decimal places.</p>

<textarea data-component="activecode" id=ex_8_22 data-lang="python"  data-hidecode="true"  data-timelimit=60 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def checkout():
    total = 0
    count = 0
    moreItems = True
    while moreItems:
        price = float(input('Enter price of item (0 when done): '))
        if price != 0:
            count = count + 1
            total = total + price
            print('Subtotal: $', total)
        else:
            moreItems = False
    average = total / count
    print('Total items:', count)
    print('Total $', total)
    print('Average price per item: $', average)

checkout()
</textarea>
</div>
</li>
</ol>
</div>


  
      

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>