<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>30. More Information About Strings &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="Loops Practice Quiz" href="../PracticeQuizzes/Loops.html" />
    <link rel="prev" title="29. More Information About Strings" href="StringIntro.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">30. More Information About Strings</a><ul>
<li><a class="reference internal" href="#traversal-and-the-for-loop-by-index">30.1. Traversal and the <code class="docutils literal"><span class="pre">for</span></code> Loop: By Index</a><ul>
<li><a class="reference internal" href="#check-your-understanding">30.1.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#practice-problems">30.2. Practice Problems</a><ul>
<li><a class="reference internal" href="#con-cat">30.2.1. Con Cat</a></li>
<li><a class="reference internal" href="#duckling-names">30.2.2. Duckling Names</a></li>
<li><a class="reference internal" href="#first-letter-of-a-word">30.2.3. First Letter of a Word</a></li>
<li><a class="reference internal" href="#password-length">30.2.4. Password Length</a></li>
<li><a class="reference internal" href="#removing-the-start-and-end">30.2.5. Removing the Start and End</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 29. More Information About Strings" data-toggle="tooltip">
    <a href="StringIntro.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Loops Practice Quiz' data-toggle="tooltip" >
    <a href="../PracticeQuizzes/Loops.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="more-information-about-strings">
<h1>30. More Information About Strings<a class="headerlink" href="#more-information-about-strings" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>Explore string traversal (for loop by index), <code class="docutils literal"><span class="pre">in</span></code> and <code class="docutils literal"><span class="pre">not</span> <span class="pre">in</span></code> operators, accumulator pattern with strings, and using strings to control turtle movement.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed"   data-showtitle="Curriculum Objectives Addressed In This Section" data-hidetitle="Hide">
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to encapsulate reusable pieces of code.</li>
</ul>

    </div>
    <div class="section" id="traversal-and-the-for-loop-by-index">
<h2>30.1. Traversal and the <code class="docutils literal"><span class="pre">for</span></code> Loop: By Index<a class="headerlink" href="#traversal-and-the-for-loop-by-index" title="Permalink to this headline">¶</a></h2>
<p>It is also possible to use the <code class="docutils literal"><span class="pre">range</span></code> function to systematically generate the indices of the characters.  The <code class="docutils literal"><span class="pre">for</span></code> loop can then be used to iterate over these positions.
These positions can be used together with the indexing operator to access the individual
characters in the string.</p>
<p>Consider the following codelens example.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="string_for_loop_by_index_1"></div>
<p class="cl_caption"><span class="cl_caption_text"> (string_for_loop_by_index_1)</span> </p>
</div>
<script type="text/javascript">
var string_for_loop_by_index_1_trace = {"code": "fruit = \"apple\"\nfor current_index in range(5):\n    current_char = fruit[current_index]\n    print(current_char)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple"}, "ordered_globals": ["fruit"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\ne\n"}, {"line": 2, "event": "return", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\ne\n"}]};
var string_for_loop_by_index_1_vis;

$(document).ready(function() {
    try {
        string_for_loop_by_index_1_vis = new ExecutionVisualizer('string_for_loop_by_index_1',string_for_loop_by_index_1_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(string_for_loop_by_index_1_vis,'string_for_loop_by_index_1');
        styleButtons('string_for_loop_by_index_1');
        allVisualizers.push(string_for_loop_by_index_1_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component string_for_loop_by_index_1_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#string_for_loop_by_index_1_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'string_for_loop_by_index_1'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (string_for_loop_by_index_1_vis) {
        string_for_loop_by_index_1_vis.redrawConnectors();
    }
});
</script>
</div>
<p>The index positions in “apple” are 0,1,2,3 and 4.  This is exactly the same sequence of integers returned by <code class="docutils literal"><span class="pre">range(5)</span></code>.  The first time through the for loop, <code class="docutils literal"><span class="pre">current_index</span></code> will be 0 and the “a” will be printed.  Then, <code class="docutils literal"><span class="pre">current_index</span></code> will be reassigned to 1 and “p” will be displayed.  This will repeat for all the range values up to but not including 5.  Since “e” has index 4, this will be exactly right to show all
of the characters.</p>
<p>In order to make the iteration more general, we can use the <code class="docutils literal"><span class="pre">len</span></code> function to provide the bound for <code class="docutils literal"><span class="pre">range</span></code>.  This is a very common pattern for traversing any sequence by position.  Make sure you understand why the range function behaves
correctly when using <code class="docutils literal"><span class="pre">len</span></code> of the string as its parameter value.</p>

<div data-childcomponent="string_for_loop_by_index_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_for_loop_by_index_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "apple"
for current_index in range(len(fruit)):
    print(fruit[current_index])
</textarea>
</div>
<p>You may also note that iteration by position allows the programmer to control the direction of the
traversal by changing the sequence of index values.  Recall that we can create ranges that count down as
well as up so the following code will print the characters from right to left.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="string_for_loop_by_index_3"></div>
<p class="cl_caption"><span class="cl_caption_text"> (string_for_loop_by_index_3)</span> </p>
</div>
<script type="text/javascript">
var string_for_loop_by_index_3_trace = {"code": "fruit = \"apple\"\nfor current_index in range(len(fruit)-1, -1, -1):\n    print(fruit[current_index])", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple"}, "ordered_globals": ["fruit"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\na\n"}, {"line": 2, "event": "return", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\na\n"}]};
var string_for_loop_by_index_3_vis;

$(document).ready(function() {
    try {
        string_for_loop_by_index_3_vis = new ExecutionVisualizer('string_for_loop_by_index_3',string_for_loop_by_index_3_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(string_for_loop_by_index_3_vis,'string_for_loop_by_index_3');
        styleButtons('string_for_loop_by_index_3');
        allVisualizers.push(string_for_loop_by_index_3_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component string_for_loop_by_index_3_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#string_for_loop_by_index_3_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'string_for_loop_by_index_3'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (string_for_loop_by_index_3_vis) {
        string_for_loop_by_index_3_vis.redrawConnectors();
    }
});
</script>
</div>
<p>Trace the values of <code class="docutils literal"><span class="pre">current_index</span></code> and satisfy yourself that they are correct.  In particular, note the start and end of the range.</p>
<div class="section" id="check-your-understanding">
<h3>30.1.1. Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_for_loop_by_index_test_question_1">
            <p>more-info-about-strings1: How many times is the letter o printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">current_index</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">)):</span>
    <span class="k">if</span> <span class="n">current_index</span> <span class="o">%</span> <span class="mi">2</span> <span class="o">==</span> <span class="mi">0</span><span class="p">:</span>
        <span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="n">current_index</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_a">0</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_a">The for loop visits each index but the selection only prints some of them.</li>
            
            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_b">1</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_b">o is at positions 4 and 8</li>
            
            <li data-component="answer" data-correct id="string_for_loop_by_index_test_question_1_opt_c">2</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_c">Yes, it will print all the characters in even index positions and the o character appears both times in an even location.</li>
            
            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_d">Error, the for statement cannot have an if inside.</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_d">The for statement can have any statements inside, including if as well as for.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="practice-problems">
<h2>30.2. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems. You can either work directly in the textbook, or use Thonny. Either way, copy/paste your finished code into Thonny and save your solution into your Computer Science 20 folder when you finish!</p>
<div class="section" id="con-cat">
<h3>30.2.1. Con Cat<a class="headerlink" href="#con-cat" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in the name of a cat, then prints out a hello message. For example, if the user types in <code class="docutils literal"><span class="pre">Garfield</span></code>, your program could print something like <code class="docutils literal"><span class="pre">Good</span> <span class="pre">to</span> <span class="pre">see</span> <span class="pre">you,</span> <span class="pre">Garfield!</span></code>. <strong>Be sure to use concatenation in your solution!</strong></p>

<div data-childcomponent="practice_problem_string_into_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_1 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Hello message to a cat
</textarea>
</div>
</div>
<div class="section" id="duckling-names">
<h3>30.2.2. Duckling Names<a class="headerlink" href="#duckling-names" title="Permalink to this headline">¶</a></h3>
<p>In Robert McCloskey’s book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>

<div data-childcomponent="practice_problem_string_into_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_2 data-lang="python"    data-timelimit=25000   data-audio=''     >
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    # add a conditional statement here...
    print(letter + suffix)
</textarea>
</div>
<p>Of course, that’s not quite right because Ouack and Quack are misspelled. See if you can fix it without changing the value of either the <code class="docutils literal"><span class="pre">prefixes</span></code> or <code class="docutils literal"><span class="pre">suffix</span></code> variables! <em>You should do this by adding a conditional statement within the for loop.</em></p>
</div>
<div class="section" id="first-letter-of-a-word">
<h3>30.2.3. First Letter of a Word<a class="headerlink" href="#first-letter-of-a-word" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_into_3" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal"><span class="pre">word</span></code> that returns <code class="docutils literal"><span class="pre">True</span></code> if the <code class="docutils literal"><span class="pre">word</span></code> begins with the letter “t” or “c”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;thing&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;concatenation&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;warman&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_into_3 data-lang="python"    data-timelimit=25000   data-audio=''     >
def starts_with_tc(word):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(starts_with_tc("thing"),True,'starts_with_tc("thing")')
        self.assertEqual(starts_with_tc("concatenation"),True,'starts_with_tc("concatenation")')
        self.assertEqual(starts_with_tc("warman"),False,'starts_with_tc("warman")')
        self.assertEqual(starts_with_tc("Thing"),False,'starts_with_tc("Thing")')
        self.assertEqual(starts_with_tc("candy"),True,'starts_with_tc("candy")')
        self.assertEqual(starts_with_tc("saskatoon"),False,'starts_with_tc("saskatoon")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="password-length">
<h3>30.2.4. Password Length<a class="headerlink" href="#password-length" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_into_4" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal"><span class="pre">password</span></code> that returns <code class="docutils literal"><span class="pre">True</span></code> if the <code class="docutils literal"><span class="pre">password</span></code> is between 8 and 32 characters (inclusive), and <code class="docutils literal"><span class="pre">False</span></code> otherwise. <em>Please note there is much more to a strong password than just the length of the string!</em></p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;123456789&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;qwerty&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;cray-topnotch-tampa-anthem-trial&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_into_4 data-lang="python"    data-timelimit=25000   data-audio=''     >
def password_length(password):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(password_length("123456789"),True,'password_length("123456789")')
        self.assertEqual(password_length("qwerty"),False,'password_length("qwerty")')
        self.assertEqual(password_length("cray-topnotch-tampa-anthem-trial"),True,'password_length("cray-topnotch-tampa-anthem-trial")')
        self.assertEqual(password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt"),False,'password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt")')
        self.assertEqual(password_length("fixate"),False,'password_length("fixate")')
        self.assertEqual(password_length("8X?3a739"),True,'password_length("8X?3a739")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="removing-the-start-and-end">
<h3>30.2.5. Removing the Start and End<a class="headerlink" href="#removing-the-start-and-end" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in three inputs from the user:</p>
<ul class="simple">
<li>a word/sentence</li>
<li>a starting number</li>
<li>an ending number</li>
</ul>
<p>After taking in the input, the program should print out a portion of the word/sentence that was entered. For example, if the word the user entered was <code class="docutils literal"><span class="pre">Saskatchewan</span></code>, the starting number was 3 and the ending number was 2, the program should print <code class="docutils literal"><span class="pre">katchew</span></code>. Notice that in the output, it is the same as the word, except that the first 3 characters and last 2 characters have been removed.</p>
<p>To be sure you understand the challenge, consider the following example:</p>
<ul class="simple">
<li>word = “unimaginatively”</li>
<li>starting_number = 3</li>
<li>ending_number = 2</li>
<li>output should be <code class="docutils literal"><span class="pre">maginative</span></code></li>
</ul>

<div data-childcomponent="practice_problem_string_into_5" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_5 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Remove the Start and End
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 29. More Information About Strings' data-toggle="tooltip">
    <a href="StringIntro.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="StringIntro.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Loops Practice Quiz' data-toggle="tooltip" >
    <a href="../PracticeQuizzes/Loops.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="../PracticeQuizzes/Loops.html" >Next Section - Loops Practice Quiz</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>