<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>2. String Traversal (by index), and the Accumulator Pattern with Strings &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=808B982A" />
    <link rel="stylesheet" type="text/css" href="../../_static/activecode.css?v=F471F95B" />
    <link rel="stylesheet" type="text/css" href="../../_static/codemirror.css?v=1CE53237" />
    <link rel="stylesheet" type="text/css" href="../../_static/clickable.css?v=58C4E18A" />
    <link rel="stylesheet" type="text/css" href="../../_static/pytutor.css?v=239FC54E" />
    <link rel="stylesheet" type="text/css" href="../../_static/modal-basic.css?v=8A1A42E4" />
    <link rel="stylesheet" type="text/css" href="../../_static/datafile.css?v=C184EB8A" />
    <link rel="stylesheet" type="text/css" href="../../_static/dragndrop.css?v=B9BAA0D1" />
    <link rel="stylesheet" type="text/css" href="../../_static/fitb.css?v=3BA0E63E" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../../_static/parsons.css?v=2F3E13EE" />
    <link rel="stylesheet" type="text/css" href="../../_static/js_lib/prettify.css?v=2A364B6B" />
    <link rel="stylesheet" type="text/css" href="../../_static/poll.css?v=E220AC33" />
    <link rel="stylesheet" type="text/css" href="../../_static/reveal.css?v=4A81D381" />
    <link rel="stylesheet" type="text/css" href="../../_static/showEval.css?v=C6382F3" />
    <link rel="stylesheet" type="text/css" href="../../_static/tabbedstuff.css?v=424F1CC9" />
    <link rel="stylesheet" type="text/css" href="../../_static/video.css?v=E5385B3B" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../../" src="../../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/language_data.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js?v=F1C496C8"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js?v=BAE6E360"></script>
    <script type="text/javascript" src="../../_static/codemirror.js?v=94F552C9"></script>
    <script type="text/javascript" src="../../_static/xml.js?v=236EDFF7"></script>
    <script type="text/javascript" src="../../_static/css.js?v=623153CE"></script>
    <script type="text/javascript" src="../../_static/python.js?v=14A28F1F"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js?v=33B91944"></script>
    <script type="text/javascript" src="../../_static/javascript.js?v=6F38CA04"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/CLDRPluralRuleParser.js?v=599B42B0"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.js?v=88A1430F"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.messagestore.js?v=AE15A4A1"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.fallbacks.js?v=8ABDFE7F"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.language.js?v=747279FF"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.parser.js?v=7B020E13"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.emitter.js?v=660E85C3"></script>
    <script type="text/javascript" src="../../_static/jquery_i18n/jquery.i18n.emitter.bidi.js?v=54B21948"></script>
    <script type="text/javascript" src="../../_static/activecode-i18n.en.js?v=3FBAE230"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js?v=DC7AC815"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js?v=48A5D93E"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js?v=9FA573E1"></script>
    <script type="text/javascript" src="../../_static/clike.js?v=2DCE0756"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js?v=5150FA78"></script>
    <script type="text/javascript" src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script type="text/javascript" src="../../_static/mchoice.js?v=8329CA55"></script>
    <script type="text/javascript" src="../../_static/timedmc.js?v=B010A457"></script>
    <script type="text/javascript" src="../../_static/timed.js?v=6789FE30"></script>
    <script type="text/javascript" src="../../_static/clickable.js?v=9560E65E"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js?v=93989C23"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js?v=836E01C"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js?v=B64CF0A8"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js?v=42F9EBEA"></script>
    <script type="text/javascript" src="../../_static/pytutor.js?v=82B022B9"></script>
    <script type="text/javascript" src="../../_static/codelens.js?v=CABECD49"></script>
    <script type="text/javascript" src="../../_static/datafile.js?v=3ED1A80C"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js?v=2B6403E8"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js?v=378F237E"></script>
    <script type="text/javascript" src="../../_static/fitb.js?v=8E40E713"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js?v=2AA17071"></script>
    <script type="text/javascript" src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script type="text/javascript" src="../../_static/js_lib/prettify.js?v=CA7E8928"></script>
    <script type="text/javascript" src="../../_static/js_lib/hammer.min.js?v=73056367"></script>
    <script type="text/javascript" src="../../_static/parsons.js?v=93EFFF39"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js?v=2EEA9E5A"></script>
    <script type="text/javascript" src="../../_static/poll.js?v=6C59AE9F"></script>
    <script type="text/javascript" src="../../_static/reveal.js?v=112CBFA2"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js?v=5EB6170D"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js?v=B3A117D7"></script>
    <script type="text/javascript" src="../../_static/showEval.js?v=C5CE0C81"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js?v=5BE6106"></script>
    <script type="text/javascript" src="../../_static/runestonevideo.js?v=53E888DE"></script>
    <script type="text/javascript" src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script type="text/javascript" src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script type="text/javascript" src="../../_static/Blob.js?v=AEC2447"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="3. String Methods" href="StringMethods.html" />
    <link rel="prev" title="1. String Index, Length, Slicing and Traversal" href="StringIntro.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.4.4';
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">ActiveCode Work Area</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">2. String Traversal (by index), and the Accumulator Pattern with Strings</a><ul>
<li><a class="reference internal" href="#what-does-this-program-do">2.1. What Does This Program Do?</a></li>
<li><a class="reference internal" href="#traversal-and-the-for-loop-by-index">2.2. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index</a><ul>
<li><a class="reference internal" href="#check-your-understanding">2.2.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#the-in-and-not-in-operators">2.3. The <code class="docutils literal notranslate"><span class="pre">in</span></code> and <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> operators</a></li>
<li><a class="reference internal" href="#the-accumulator-pattern-with-strings">2.4. The Accumulator Pattern with Strings</a><ul>
<li><a class="reference internal" href="#id1">2.4.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#practice-problems">2.5. Practice Problems</a><ul>
<li><a class="reference internal" href="#even-letters-of-a-word">2.5.1. Even Letters of a Word</a></li>
<li><a class="reference internal" href="#reverse-me">2.5.2. Reverse Me</a></li>
<li><a class="reference internal" href="#letter-destroyer">2.5.3. Letter Destroyer</a></li>
<li><a class="reference internal" href="#qwerty-finder">2.5.4. QWERTY Finder</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 1. String Index, Length, Slicing and Traversal" data-toggle="tooltip">
    <a href="StringIntro.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 3. String Methods' data-toggle="tooltip" >
    <a href="StringMethods.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="string-traversal-by-index-and-the-accumulator-pattern-with-strings">
<h1>2. String Traversal (by index), and the Accumulator Pattern with Strings<a class="headerlink" href="#string-traversal-by-index-and-the-accumulator-pattern-with-strings" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>WDTPD string questions. Explore string traversal (for loop by index), <code class="docutils literal notranslate"><span class="pre">in</span></code> and <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> operators, and the accumulator pattern with strings. Work on practice problems covering these ideas.</p>
</div>

    <div data-component="reveal" id="curriculum_addressed_string_accumulator_pattern"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide" >
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to create reusable pieces of code.</li>
</ul>

    </div>
    <div class="section" id="what-does-this-program-do">
<h2>2.1. What Does This Program Do?<a class="headerlink" href="#what-does-this-program-do" title="Permalink to this headline">¶</a></h2>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">Your teacher may choose to use the following examples as a class activity, by displaying the  examples, and having you take a guess as to what you think each will do before running the code.</p>
</div>
<p><strong>What will the following programs output? Why?</strong></p>

<div data-childcomponent="wdtpd_string_manipulation_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_string_manipulation_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
         data-caption='What will this program print?' >
title = "Pirates of the Caribbean"
print(title[1])
</textarea>
</div>

<div data-childcomponent="wdtpd_string_manipulation_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_string_manipulation_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
         data-caption='What will this program print?' >
title = "Pirates of the Caribbean"
print(title[-3])
</textarea>
</div>

<div data-childcomponent="wdtpd_string_manipulation_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_string_manipulation_3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
         data-caption='What will this program print?' >
title = "Pirates of the Caribbean"
print(title[2:6])
</textarea>
</div>

<div data-childcomponent="wdtpd_string_manipulation_4" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=wdtpd_string_manipulation_4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
         data-caption='What will this program print?' >
fruit = "pineapple"
for poorly_named_variable in fruit:
    print(poorly_named_variable)
</textarea>
</div>
</div>
<div class="section" id="traversal-and-the-for-loop-by-index">
<h2>2.2. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index<a class="headerlink" href="#traversal-and-the-for-loop-by-index" title="Permalink to this headline">¶</a></h2>
<p>It is also possible to use the <code class="docutils literal notranslate"><span class="pre">range</span></code> function to systematically generate the indices of the characters.  The <code class="docutils literal notranslate"><span class="pre">for</span></code> loop can then be used to iterate over these positions.
These positions can be used together with the indexing operator to access the individual
characters in the string.</p>
<p>Consider the following codelens example.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="string_for_loop_by_index_1"></div>
<p class="cl_caption"><span class="cl_caption_text"> (string_for_loop_by_index_1)</span> </p>
</div>
<script type="text/javascript">
var string_for_loop_by_index_1_trace = {"code": "fruit = \"apple\"\nfor current_index in range(5):\n    current_char = fruit[current_index]\n    print(current_char)", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple"}, "ordered_globals": ["fruit"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "a"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "p"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "l"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 4, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\ne\n"}, {"line": 2, "event": "return", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4, "current_char": "e"}, "ordered_globals": ["fruit", "current_index", "current_char"], "stack_to_render": [], "heap": {}, "stdout": "a\np\np\nl\ne\n"}]};
var string_for_loop_by_index_1_vis;

$(document).ready(function() {
    try {
        string_for_loop_by_index_1_vis = new ExecutionVisualizer('string_for_loop_by_index_1',string_for_loop_by_index_1_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(string_for_loop_by_index_1_vis,'string_for_loop_by_index_1');
        styleButtons('string_for_loop_by_index_1');
        allVisualizers.push(string_for_loop_by_index_1_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component string_for_loop_by_index_1_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#string_for_loop_by_index_1_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'string_for_loop_by_index_1'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (string_for_loop_by_index_1_vis) {
        string_for_loop_by_index_1_vis.redrawConnectors();
    }
});
</script>
</div>
<p>The index positions in “apple” are 0,1,2,3 and 4.  This is exactly the same sequence of integers returned by <code class="docutils literal notranslate"><span class="pre">range(5)</span></code>.  The first time through the for loop, <code class="docutils literal notranslate"><span class="pre">current_index</span></code> will be 0 and the “a” will be printed.  Then, <code class="docutils literal notranslate"><span class="pre">current_index</span></code> will be reassigned to 1 and “p” will be displayed.  This will repeat for all the range values up to but not including 5.  Since “e” has index 4, this will be exactly right to show all
of the characters.</p>
<p>In order to make the iteration more general, we can use the <code class="docutils literal notranslate"><span class="pre">len</span></code> function to provide the bound for <code class="docutils literal notranslate"><span class="pre">range</span></code>.  This is a very common pattern for traversing any sequence by position.  Make sure you understand why the range function behaves
correctly when using <code class="docutils literal notranslate"><span class="pre">len</span></code> of the string as its parameter value.</p>

<div data-childcomponent="string_for_loop_by_index_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_for_loop_by_index_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
          >
fruit = "apple"
for current_index in range(len(fruit)):
    print(fruit[current_index])
</textarea>
</div>
<p>You may also note that iteration by position allows the programmer to control the direction of the
traversal by changing the sequence of index values.  Recall that we can create ranges that count down as well as up so the following code will print the characters from right to left.</p>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">Remember that when using the <code class="docutils literal notranslate"><span class="pre">range()</span></code> function with three arguments, the arguments are interpreted as <code class="docutils literal notranslate"><span class="pre">starting_number</span></code>, <code class="docutils literal notranslate"><span class="pre">ending_number</span></code>, and <code class="docutils literal notranslate"><span class="pre">amount_to_step_by</span></code>. The range function begins at the <code class="docutils literal notranslate"><span class="pre">starting_number</span></code>, and goes up to, but not including, the <code class="docutils literal notranslate"><span class="pre">ending_number</span></code>.</p>
</div>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="string_for_loop_by_index_3"></div>
<p class="cl_caption"><span class="cl_caption_text"> (string_for_loop_by_index_3)</span> </p>
</div>
<script type="text/javascript">
var string_for_loop_by_index_3_trace = {"code": "fruit = \"apple\"\nfor current_index in range(len(fruit)-1, -1, -1):\n    print(fruit[current_index])", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple"}, "ordered_globals": ["fruit"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 4}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 3}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 2}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 1}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\n"}, {"line": 3, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\n"}, {"line": 2, "event": "step_line", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\na\n"}, {"line": 2, "event": "return", "func_name": "<module>", "globals": {"fruit": "apple", "current_index": 0}, "ordered_globals": ["fruit", "current_index"], "stack_to_render": [], "heap": {}, "stdout": "e\nl\np\np\na\n"}]};
var string_for_loop_by_index_3_vis;

$(document).ready(function() {
    try {
        string_for_loop_by_index_3_vis = new ExecutionVisualizer('string_for_loop_by_index_3',string_for_loop_by_index_3_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(string_for_loop_by_index_3_vis,'string_for_loop_by_index_3');
        styleButtons('string_for_loop_by_index_3');
        allVisualizers.push(string_for_loop_by_index_3_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component string_for_loop_by_index_3_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#string_for_loop_by_index_3_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'string_for_loop_by_index_3'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (string_for_loop_by_index_3_vis) {
        string_for_loop_by_index_3_vis.redrawConnectors();
    }
});
</script>
</div>
<p>Trace the values of <code class="docutils literal notranslate"><span class="pre">current_index</span></code> and satisfy yourself that they are correct.  In particular, note the start and end of the range.</p>
<div class="section" id="check-your-understanding">
<h3>2.2.1. Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_for_loop_by_index_test_question_1">
            <p>string-acc1: How many times is the letter o printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">current_index</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">)):</span>
  <span class="k">if</span> <span class="n">current_index</span> <span class="o">%</span> <span class="mi">2</span> <span class="o">==</span> <span class="mi">0</span><span class="p">:</span>
      <span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="n">current_index</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_a">0</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_a">The for loop visits each index but the selection only prints some of them.</li>
            
            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_b">1</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_b">o is at positions 4 and 8</li>
            
            <li data-component="answer" data-correct id="string_for_loop_by_index_test_question_1_opt_c">2</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_c">Yes, it will print all the characters in even index positions and the o character appears both times in an even location.</li>
            
            <li data-component="answer"  id="string_for_loop_by_index_test_question_1_opt_d">Error, the for statement cannot have an if inside.</li><li data-component="feedback" id="string_for_loop_by_index_test_question_1_opt_d">The for statement can have any statements inside, including if as well as for.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="the-in-and-not-in-operators">
<span id="index-0"></span><h2>2.3. The <code class="docutils literal notranslate"><span class="pre">in</span></code> and <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> operators<a class="headerlink" href="#the-in-and-not-in-operators" title="Permalink to this headline">¶</a></h2>
<p>The <code class="docutils literal notranslate"><span class="pre">in</span></code> operator tests if one string is a substring of another:</p>

<div data-childcomponent="string_in_not_in_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_in_not_in_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
          >
print('p' in 'apple')
print('i' in 'apple')
print('ap' in 'apple')
print('pa' in 'apple')
</textarea>
</div>
<p>Note that a string is a substring of itself, and the empty string is a
substring of any other string. <em>(Also note that computer scientists
like to think about these edge cases quite carefully!)</em></p>

<div data-childcomponent="string_in_not_in_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_in_not_in_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
          >
print('a' in 'a')
print('apple' in 'apple')
print('' in 'a')
print('' in 'apple')
</textarea>
</div>
<p>The <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> operator returns the logical opposite result of <code class="docutils literal notranslate"><span class="pre">in</span></code>.</p>

<div data-childcomponent="string_in_not_in_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_in_not_in_3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
          >
print('x' not in 'apple')
</textarea>
</div>
</div>
<div class="section" id="the-accumulator-pattern-with-strings">
<span id="index-1"></span><h2>2.4. The Accumulator Pattern with Strings<a class="headerlink" href="#the-accumulator-pattern-with-strings" title="Permalink to this headline">¶</a></h2>
<p>Combining the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator with string concatenation using <code class="docutils literal notranslate"><span class="pre">+</span></code> and the accumulator pattern, we can
write a function that removes all the vowels from a string.  The idea is to start with a string and iterate over each character, checking to see if the character is a vowel.  As we process the characters, we will build up a new string consisting of only the nonvowel characters.  To do this, we use the accumulator pattern.</p>
<p>Remember that the accumulator pattern allows us to keep a “running total”.  With strings, we are not accumulating a numeric total.  Instead we are accumulating characters onto a string.</p>

<div data-childcomponent="string_accumulator_pattern_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_accumulator_pattern_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''     
          >
def remove_vowels(original_string):
    vowels = "aeiouAEIOU"
    string_without_vowels = ""
    for this_char in original_string:
        if this_char not in vowels:
            string_without_vowels = string_without_vowels + this_char
    return string_without_vowels

print(remove_vowels("compsci"))
print(remove_vowels("aAbEefIijOopUus"))
</textarea>
</div>
<p>Line 5 uses the <code class="docutils literal notranslate"><span class="pre">not</span> <span class="pre">in</span></code> operator to check whether the current character is not in the string <code class="docutils literal notranslate"><span class="pre">vowels</span></code>. The alternative to using this operator would be to write a very large <code class="docutils literal notranslate"><span class="pre">if</span></code> statement that checks each of the individual vowel characters.  Note we would need to use logical <code class="docutils literal notranslate"><span class="pre">and</span></code> to be sure that the character is not any of the vowels.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">if</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;a&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;e&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;i&#39;</span>  <span class="ow">and</span>
   <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;o&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;u&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;A&#39;</span>  <span class="ow">and</span>
   <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;E&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;I&#39;</span>  <span class="ow">and</span> <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;O&#39;</span>  <span class="ow">and</span>
   <span class="n">each_char</span> <span class="o">!=</span> <span class="s1">&#39;U&#39;</span><span class="p">:</span>

     <span class="n">string_without_vowels</span> <span class="o">=</span> <span class="n">string_without_vowels</span> <span class="o">+</span> <span class="n">each_char</span>
</pre></div>
</div>
<p>Look carefully at line 6 in the above program (<code class="docutils literal notranslate"><span class="pre">string_without_vowels</span> <span class="pre">=</span> <span class="pre">string_without_vowels</span> <span class="pre">+</span> <span class="pre">each_char</span></code>).  We will do this for every character that is not a vowel.  This should look
very familiar.  As we were describing earlier, it is an example of the accumulator pattern, this time using a string to “accumulate” the final result.
In words it says that the new value of <code class="docutils literal notranslate"><span class="pre">string_without_vowels</span></code> will be the old value of <code class="docutils literal notranslate"><span class="pre">string_without_vowels</span></code> concatenated with
the value of <code class="docutils literal notranslate"><span class="pre">each_char</span></code>.  We are building the result string character by character.</p>
<p>Take a close look also at the initialization of <code class="docutils literal notranslate"><span class="pre">string_without_vowels</span></code>.  We start with an empty string and then begin adding
new characters to the end.</p>
<p>Step through the function using codelens to see the accumulator variable grow.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="string_accumulator_pattern_2"></div>
<p class="cl_caption"><span class="cl_caption_text"> (string_accumulator_pattern_2)</span> </p>
</div>
<script type="text/javascript">
var string_accumulator_pattern_2_trace = {"code": "def remove_vowels(original_string):\n    vowels = \"aeiouAEIOU\"\n    string_without_vowels = \"\"\n    for each_char in original_string:\n        if each_char not in vowels:\n            string_without_vowels = string_without_vowels + each_char\n    return string_without_vowels\n\nprint(remove_vowels(\"compsci\"))", "trace": [{"line": 1, "event": "step_line", "func_name": "<module>", "globals": {}, "ordered_globals": [], "stack_to_render": [], "heap": {}, "stdout": ""}, {"line": 9, "event": "step_line", "func_name": "<module>", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 1, "event": "call", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci"}, "ordered_varnames": ["original_string"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 2, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci"}, "ordered_varnames": ["original_string"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 3, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU"}, "ordered_varnames": ["original_string", "vowels"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": ""}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "c", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "c", "each_char": "o"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "c", "each_char": "o"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "c", "each_char": "m"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "c", "each_char": "m"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cm", "each_char": "m"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cm", "each_char": "p"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cm", "each_char": "p"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmp", "each_char": "p"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmp", "each_char": "s"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmp", "each_char": "s"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmps", "each_char": "s"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmps", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 6, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmps", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmpsc", "each_char": "c"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 5, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmpsc", "each_char": "i"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 4, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmpsc", "each_char": "i"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 7, "event": "step_line", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmpsc", "each_char": "i"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 7, "event": "return", "func_name": "remove_vowels", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [{"func_name": "remove_vowels", "is_parent": false, "frame_id": 1, "parent_frame_id_list": [], "encoded_locals": {"original_string": "compsci", "vowels": "aeiouAEIOU", "string_without_vowels": "cmpsc", "each_char": "i", "__return__": "cmpsc"}, "ordered_varnames": ["original_string", "vowels", "string_without_vowels", "each_char", "__return__"], "is_zombie": false, "is_highlighted": true, "unique_hash": "remove_vowels_f1"}], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": ""}, {"line": 9, "event": "return", "func_name": "<module>", "globals": {"remove_vowels": ["REF", 1]}, "ordered_globals": ["remove_vowels"], "stack_to_render": [], "heap": {"1": ["FUNCTION", "remove_vowels(original_string)", null]}, "stdout": "cmpsc\n"}]};
var string_accumulator_pattern_2_vis;

$(document).ready(function() {
    try {
        string_accumulator_pattern_2_vis = new ExecutionVisualizer('string_accumulator_pattern_2',string_accumulator_pattern_2_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(string_accumulator_pattern_2_vis,'string_accumulator_pattern_2');
        styleButtons('string_accumulator_pattern_2');
        allVisualizers.push(string_accumulator_pattern_2_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component string_accumulator_pattern_2_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#string_accumulator_pattern_2_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'string_accumulator_pattern_2'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (string_accumulator_pattern_2_vis) {
        string_accumulator_pattern_2_vis.redrawConnectors();
    }
});
</script>
</div>
<div class="section" id="id1">
<h3>2.4.1. Check Your Understanding<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h3>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_accumulator_pattern_test_question_1">
            <p>string-acc2: What is printed by the following statements:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">some_string</span> <span class="o">=</span> <span class="s2">&quot;ball&quot;</span>
<span class="n">another_string</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
<span class="k">for</span> <span class="n">item</span> <span class="ow">in</span> <span class="n">some_string</span><span class="p">:</span>
    <span class="n">another_string</span> <span class="o">=</span> <span class="n">another_string</span> <span class="o">+</span> <span class="n">item</span>
<span class="k">print</span><span class="p">(</span><span class="n">another_string</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="string_accumulator_pattern_test_question_1_opt_a">ball</li><li data-component="feedback" id="string_accumulator_pattern_test_question_1_opt_a">Yes, the repeated concatenation will cause another_string to become the same as some_string.</li>
            
            <li data-component="answer"  id="string_accumulator_pattern_test_question_1_opt_b">llab</li><li data-component="feedback" id="string_accumulator_pattern_test_question_1_opt_b">Look again at the *order* of the concatenation!</li>
            

            </ul>
            </div>
            
            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_accumulator_pattern_test_question_2">
            <p>string-acc3: What is printed by the following statements:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">some_string</span> <span class="o">=</span> <span class="s2">&quot;ball&quot;</span>
<span class="n">another_string</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
<span class="k">for</span> <span class="n">item</span> <span class="ow">in</span> <span class="n">some_string</span><span class="p">:</span>
    <span class="n">another_string</span> <span class="o">=</span> <span class="n">item</span> <span class="o">+</span> <span class="n">another_string</span>
<span class="k">print</span><span class="p">(</span><span class="n">another_string</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_accumulator_pattern_test_question_2_opt_a">ball</li><li data-component="feedback" id="string_accumulator_pattern_test_question_2_opt_a">Look again at the *order* of the concatenation!</li>
            
            <li data-component="answer" data-correct id="string_accumulator_pattern_test_question_2_opt_b">llab</li><li data-component="feedback" id="string_accumulator_pattern_test_question_2_opt_b">Yes, the order is reversed due to the order of the concatenation.</li>
            

            </ul>
            </div>
            
    <div data-component="reveal" id="reveal_string_manipulation_reversed_order_of_concatenation"   data-showtitle="Need Help?" data-hidetitle="Hide Help" >
    <p>Consider the following table, which shows the values of the variables as the code iterates:</p>
<table border="1" class="docutils">
<colgroup>
<col width="17%" />
<col width="22%" />
<col width="24%" />
<col width="37%" />
</colgroup>
<thead valign="bottom">
<tr class="row-odd"><th class="head">iteration</th>
<th class="head">item</th>
<th class="head">another_string</th>
<th class="head">item + another_string</th>
</tr>
</thead>
<tbody valign="top">
<tr class="row-even"><td>1</td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;b&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;b&quot;</span></code></td>
</tr>
<tr class="row-odd"><td>2</td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;a&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;b&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;ab&quot;</span></code></td>
</tr>
<tr class="row-even"><td>3</td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;l&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;ab&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;lab&quot;</span></code></td>
</tr>
<tr class="row-odd"><td>4</td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;l&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;lab&quot;</span></code></td>
<td><code class="docutils literal notranslate"><span class="pre">&quot;llab&quot;</span></code></td>
</tr>
</tbody>
</table>

    </div>
    
        <div class="runestone" style="max-width: none;">
        <pre data-component="parsons" id="string_accumulator_patten_parsons_1"      >
        <span data-question>string-acc4: Construct a block of code that correctly implements the accumulator pattern with strings. After the code has finished executing, ``new_word`` is printed, and will have the same value as ``original_word``.</span>original_word = "clockwork"
new_word = ""
for letter in original_word:
    new_word = new_word + letter
print(new_word)
        </pre>
        </div>
    </div>
</div>
<div class="section" id="practice-problems">
<h2>2.5. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems. You can either work directly in the textbook, or use Thonny. Either way, be sure to save your solution into your Computer Science 20 folder when you finish!</p>
<p><strong>Hint: For each of the following, you will want to use the accumulator pattern with strings. In other words, you first need to create an empty string, then concatenate letters onto it as needed.</strong></p>
<div class="section" id="even-letters-of-a-word">
<h3>2.5.1. Even Letters of a Word<a class="headerlink" href="#even-letters-of-a-word" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_more_info_1" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal notranslate"><span class="pre">word</span></code> that returns the even letters of the <code class="docutils literal notranslate"><span class="pre">word</span></code> (the first letter is even, since we start counting our index values at 0). For example, given the <code class="docutils literal notranslate"><span class="pre">word</span></code> “Saskatoon”, the function should return “Ssaon”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">even_letters(&quot;Saskatoon&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;Ssaon&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">even_letters(&quot;Saskatchewan&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;Ssacea&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">even_letters(&quot;Roughriders&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;Ruhies&quot;</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_more_info_1 data-lang="python" 
      data-timelimit=25000   data-enabledownload="true" 
    data-audio=''     
          >
def even_letters(word):
    return ""

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(even_letters("Saskatoon"),"Ssaon",'even_letters("Saskatoon")')
        self.assertEqual(even_letters("Saskatchewan"),"Ssacea",'even_letters("Saskatchewan")')
        self.assertEqual(even_letters("Roughriders"),"Ruhies",'even_letters("Roughriders")')
        self.assertEqual(even_letters("Tisdale"),"Tsae",'even_letters("Tisdale")')
        self.assertEqual(even_letters("Hepburn"),"Hpun",'even_letters("Hepburn")')
        self.assertEqual(even_letters("Neuanlage"),"Nunae",'even_letters("Neuanlage")')
        self.assertEqual(even_letters("Blumenheim"),"Buehi",'even_letters("Blumenheim")')

myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="reveal_string_manipulation_even_letters_hint"   data-showtitle="Need a Hint?" data-hidetitle="Hide Hint" >
    <p>Since you are trying to get only the even letters, you need to know which index value you are currently on. That means we need to traverse the string using a <code class="docutils literal notranslate"><span class="pre">for</span> <span class="pre">counter</span> <span class="pre">in</span> <span class="pre">range()</span></code> style loop. The accumulator pattern will look something like this:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">even_letters</span><span class="p">(</span><span class="n">word</span><span class="p">):</span>
    <span class="n">new_word</span> <span class="o">=</span> <span class="s2">&quot;&quot;</span>
    <span class="k">for</span> <span class="n">counter</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">word</span><span class="p">)):</span>
        <span class="c1"># add a question here...</span>
        <span class="n">new_word</span> <span class="o">=</span> <span class="n">new_word</span> <span class="o">+</span> <span class="n">word</span><span class="p">[</span><span class="n">counter</span><span class="p">]</span>
    <span class="k">return</span> <span class="n">new_word</span>
</pre></div>
</div>

    </div>
    </div>
<div class="section" id="reverse-me">
<h3>2.5.2. Reverse Me<a class="headerlink" href="#reverse-me" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_more_info_2" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal notranslate"><span class="pre">word</span></code> that returns the <code class="docutils literal notranslate"><span class="pre">word</span></code> spelled backwards. For example, if the <code class="docutils literal notranslate"><span class="pre">word</span></code> was “Saskatoon”, the function should return “nootaksaS”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">reverse_me(&quot;Saskatoon&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;nootaksaS&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">reverse_me(&quot;Saskatchewan&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;nawehctaksaS&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">reverse_me(&quot;Roughriders&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;sredirhguoR&quot;</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_more_info_2 data-lang="python" 
      data-timelimit=25000   data-enabledownload="true" 
    data-audio=''     
          >
def reverse_me(word):
    return ""

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(reverse_me("Saskatoon"),"nootaksaS",'reverse_me("Saskatoon")')
        self.assertEqual(reverse_me("Saskatchewan"),"nawehctaksaS",'reverse_me("Saskatchewan")')
        self.assertEqual(reverse_me("Roughriders"),"sredirhguoR",'reverse_me("Roughriders")')
        self.assertEqual(reverse_me("Oxbow"),"wobxO",'reverse_me("Oxbow")')
        self.assertEqual(reverse_me("Lumsden"),"nedsmuL",'reverse_me("Lumsden")')
        self.assertEqual(reverse_me("La Loche"),"ehcoL aL",'reverse_me("La Loche")')
        self.assertEqual(reverse_me("Waskesiu"),"uiseksaW",'reverse_me("Waskesiu")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="letter-destroyer">
<h3>2.5.3. Letter Destroyer<a class="headerlink" href="#letter-destroyer" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_more_info_3" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with two parameters, <code class="docutils literal notranslate"><span class="pre">word</span></code> and <code class="docutils literal notranslate"><span class="pre">letter_to_destroy</span></code>. The function should return the <code class="docutils literal notranslate"><span class="pre">word</span></code>, but without any of the occurrences of the <code class="docutils literal notranslate"><span class="pre">letter_to_destroy</span></code> in the string. For example, if the <code class="docutils literal notranslate"><span class="pre">word</span></code> was “Saskatoon”, and the <code class="docutils literal notranslate"><span class="pre">letter_to_destroy</span></code> was ‘o’, the function should return “Saskatn”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_destroyer(&quot;Saskatoon&quot;,</span> <span class="pre">&quot;o&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;Saskatn&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_destroyer(&quot;Saskatchewan&quot;,</span> <span class="pre">&quot;a&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;Ssktchewn&quot;</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">letter_destroyer(&quot;roughriders&quot;,</span> <span class="pre">&quot;r&quot;)</span> <span class="pre">→</span> <span class="pre">&quot;oughides&quot;</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_more_info_3 data-lang="python" 
      data-timelimit=25000   data-enabledownload="true" 
    data-audio=''     
          >
def letter_destroyer(word, letter_to_destroy):
    return ""

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(letter_destroyer("Saskatoon", "o"),"Saskatn",'letter_destroyer("Saskatoon", "o")')
        self.assertEqual(letter_destroyer("Saskatchewan", "a"),"Ssktchewn",'letter_destroyer("Saskatchewan", "a")')
        self.assertEqual(letter_destroyer("roughriders", "r"),"oughides",'letter_destroyer("roughriders", "r")')
        self.assertEqual(letter_destroyer("hotshots", "h"),"otsots",'letter_destroyer("hotshots", "h")')
        self.assertEqual(letter_destroyer("hotshots", "o"),"htshts",'letter_destroyer("hotshots", "o")')
        self.assertEqual(letter_destroyer("hotshots", "t"),"hoshos",'letter_destroyer("hotshots", "t")')
        self.assertEqual(letter_destroyer("hotshots", "s"),"hothot",'letter_destroyer("hotshots", "s")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="qwerty-finder">
<h3>2.5.4. QWERTY Finder<a class="headerlink" href="#qwerty-finder" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_more_info_4" class="runestone explainer ac_section alert alert-warning">
<p>Create a function that takes in a single parameter <code class="docutils literal notranslate"><span class="pre">word</span></code> and returns the location of the first occurrence of one of the following letters: “qwerty”. If none of these letters exist in the word, have the function return -1. For example, if the <code class="docutils literal notranslate"><span class="pre">word</span></code> was “Saskatoon”, the function should return 5 (the index value for the <code class="docutils literal notranslate"><span class="pre">t</span></code> in “Saskatoon”).</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">qwerty_finder(&quot;Saskatoon&quot;)</span> <span class="pre">→</span> <span class="pre">5</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">qwerty_finder(&quot;Naomi&quot;)</span> <span class="pre">→</span> <span class="pre">-1</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">qwerty_finder(&quot;bunnyhug&quot;)</span> <span class="pre">→</span> <span class="pre">4</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_more_info_4 data-lang="python" 
      data-timelimit=25000   data-enabledownload="true" 
    data-audio=''     
          >
def qwerty_finder(word):
    return ""

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(qwerty_finder("Saskatoon"), 5,'qwerty_finder("Saskatoon")')
        self.assertEqual(qwerty_finder("Naomi"), -1,'qwerty_finder("Naomi")')
        self.assertEqual(qwerty_finder("bunnyhug"), 4,'qwerty_finder("bunnyhug")')
        self.assertEqual(qwerty_finder("Schellenberg"), 3,'qwerty_finder("Schellenberg")')
        self.assertEqual(qwerty_finder("roughriders"), 0,'qwerty_finder("roughriders")')
        self.assertEqual(qwerty_finder("junk"),-1 ,'qwerty_finder("junk")')
        self.assertEqual(qwerty_finder("university"), 4,'qwerty_finder("university")')

myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 1. String Index, Length, Slicing and Traversal' data-toggle="tooltip">
    <a href="StringIntro.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="StringIntro.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 3. String Methods' data-toggle="tooltip" >
    <a href="StringMethods.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="StringMethods.html" >Next Section - 3. String Methods</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.
    </p>
  </div>
</footer>


  </body>
</html>