<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>29. String Index, Length, Slicing and Traversal &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../../_static/jquery.js"></script>
    <script type="text/javascript" src="../../_static/underscore.js"></script>
    <script type="text/javascript" src="../../_static/doctools.js"></script>
    <script type="text/javascript" src="../../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../../_static/codemirror.js"></script>
    <script type="text/javascript" src="../../_static/xml.js"></script>
    <script type="text/javascript" src="../../_static/css.js"></script>
    <script type="text/javascript" src="../../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../../_static/python.js"></script>
    <script type="text/javascript" src="../../_static/javascript.js"></script>
    <script type="text/javascript" src="../../_static/activecode.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/clike.js"></script>
    <script type="text/javascript" src="../../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../../_static/animationbase.js"></script>
    <script type="text/javascript" src="../../_static/mchoice.js"></script>
    <script type="text/javascript" src="../../_static/timedmc.js"></script>
    <script type="text/javascript" src="../../_static/timed.js"></script>
    <script type="text/javascript" src="../../_static/clickable.js"></script>
    <script type="text/javascript" src="../../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../../_static/pytutor.js"></script>
    <script type="text/javascript" src="../../_static/codelens.js"></script>
    <script type="text/javascript" src="../../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../../_static/datafile.js"></script>
    <script type="text/javascript" src="../../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../../_static/fitb.js"></script>
    <script type="text/javascript" src="../../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../../_static/parsons.js"></script>
    <script type="text/javascript" src="../../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../../_static/poll.js"></script>
    <script type="text/javascript" src="../../_static/reveal.js"></script>
    <script type="text/javascript" src="../../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../../_static/showEval.js"></script>
    <script type="text/javascript" src="../../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="30. More Information About Strings" href="MoreAboutStrings.html" />
    <link rel="prev" title="28. Image Processing With Conditionals" href="../MoreAboutIteration/ImageProcessingWithSelection.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '3.0.5';
</script>



  </head>
  <body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/cs20sk/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">29. String Index, Length, Slicing and Traversal</a><ul>
<li><a class="reference internal" href="#a-collection-data-type">29.1. A Collection Data Type</a></li>
<li><a class="reference internal" href="#concatenation-reminder">29.2. Concatenation Reminder</a></li>
<li><a class="reference internal" href="#index-operator-working-with-the-characters-of-a-string">29.3. Index Operator: Working with the Characters of a String</a><ul>
<li><a class="reference internal" href="#check-your-understanding">29.3.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#length">29.4. Length</a><ul>
<li><a class="reference internal" href="#id2">29.4.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#the-slice-operator">29.5. The Slice Operator</a><ul>
<li><a class="reference internal" href="#id3">29.5.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#traversal-and-the-for-loop-by-item">29.6. Traversal and the <code class="docutils literal"><span class="pre">for</span></code> Loop: By Item</a><ul>
<li><a class="reference internal" href="#id4">29.6.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#practice-problems">29.7. Practice Problems</a><ul>
<li><a class="reference internal" href="#con-cat">29.7.1. Con Cat</a></li>
<li><a class="reference internal" href="#duckling-names">29.7.2. Duckling Names</a></li>
<li><a class="reference internal" href="#first-letter-of-a-word">29.7.3. First Letter of a Word</a></li>
<li><a class="reference internal" href="#password-length">29.7.4. Password Length</a></li>
<li><a class="reference internal" href="#removing-the-start-and-end">29.7.5. Removing the Start and End</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 28. Image Processing With Conditionals" data-toggle="tooltip">
    <a href="../MoreAboutIteration/ImageProcessingWithSelection.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 30. More Information About Strings' data-toggle="tooltip" >
    <a href="MoreAboutStrings.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="string-index-length-slicing-and-traversal">
<h1>29. String Index, Length, Slicing and Traversal<a class="headerlink" href="#string-index-length-slicing-and-traversal" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title first">Quick Overview of Day</p>
<p>Explore string operators, index values, length of strings, string slicing, and traversing a string using a for loop (by item).</p>
</div>

    <div data-component="reveal" id="curriculum_addressed"   data-showtitle="Curriculum Objectives Addressed In This Section" data-hidetitle="Hide">
    <ul class="simple">
<li><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</li>
<li><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</li>
<li><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</li>
<li><strong>CS20-FP3</strong> Construct and utilize functions to encapsulate reusable pieces of code.</li>
</ul>

    </div>
    <p>Throughout the portion of the course, we have used strings to represent words or phrases that we
wanted to print out.  Our definition was simple:  a string is simply some characters inside quotes.  We will now explore strings in much more detail.</p>
<div class="section" id="a-collection-data-type">
<h2>29.1. A Collection Data Type<a class="headerlink" href="#a-collection-data-type" title="Permalink to this headline">¶</a></h2>
<p>So far we have seen built-in types like: <code class="docutils literal"><span class="pre">int</span></code>, <code class="docutils literal"><span class="pre">float</span></code>,
<code class="docutils literal"><span class="pre">bool</span></code>, <code class="docutils literal"><span class="pre">str</span></code> and we’ve seen lists.
<code class="docutils literal"><span class="pre">int</span></code>, <code class="docutils literal"><span class="pre">float</span></code>, and
<code class="docutils literal"><span class="pre">bool</span></code> are considered to be simple or primitive data types because their values are not composed
of any smaller parts.  They cannot be broken down.
On the other hand, strings and lists are different from the others because they
are made up of smaller pieces.  In the case of strings, they are made up of smaller
strings each containing one <strong>character</strong>.</p>
<p>Types that are comprised of smaller pieces are called <strong>collection data types</strong>.
Depending on what we are doing, we may want to treat a collection data type as a
single entity (the whole), or we may want to access its parts. This ambiguity is useful.</p>
<p>Strings can be defined as sequential collections of characters.  This means that the individual characters
that make up the string are assumed to be in a particular order from left to right.</p>
<p>A string that contains no characters, often referred to as the <strong>empty string</strong>, is still considered to be a string.  It is simply a sequence of zero characters and is represented by ‘’ or “” (two single or two double quotes with nothing in between).</p>
</div>
<div class="section" id="concatenation-reminder">
<h2>29.2. Concatenation Reminder<a class="headerlink" href="#concatenation-reminder" title="Permalink to this headline">¶</a></h2>
<p>As we have seen before, you cannot perform mathematical operations on strings, even if the
strings look like numbers. The one exception to this rule is that the <code class="docutils literal"><span class="pre">+</span></code> operator does work with strings, but for strings, the <code class="docutils literal"><span class="pre">+</span></code> operator represents <strong>concatenation</strong>, not addition.  As we have learned previously, concatenation means joining the two operands by linking them end-to-end. For example:</p>

<div data-childcomponent="concatenation_reminder_demo" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=concatenation_reminder_demo data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "banana"
bakedGood = " nut bread"
print(fruit + bakedGood)
</textarea>
</div>
</div>
<div class="section" id="index-operator-working-with-the-characters-of-a-string">
<h2>29.3. Index Operator: Working with the Characters of a String<a class="headerlink" href="#index-operator-working-with-the-characters-of-a-string" title="Permalink to this headline">¶</a></h2>
<p>The <strong>indexing operator</strong> (Python uses square brackets to enclose the index)
selects a single character from a string.  The characters are accessed by their position or
index value.  For example, in the string shown below, the 14 characters are indexed left to right from postion 0 to position 13.</p>
<img alt="index values" src="../../_images/indexvalues1.png" />
<p>It is also the case that the positions are named from right to left using negative numbers where -1 is the rightmost
index and so on.
Note that the character at index 9 (or -5) is the blank/space character.</p>

<div data-childcomponent="string_index_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_index_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
place = "Saskatoon Sask"
some_char = place[4]
print(some_char)

last_char = place[-1]
print(last_char)
</textarea>
</div>
<p>The expression <code class="docutils literal"><span class="pre">place[4]</span></code> selects the character at index 4 from <code class="docutils literal"><span class="pre">place</span></code>, and creates a new
string containing just this one character. The variable <code class="docutils literal"><span class="pre">some_char</span></code> refers to the result.</p>
<p>Remember that computer scientists often start counting
from zero. The letter at index zero of <code class="docutils literal"><span class="pre">&quot;Saskatoon</span> <span class="pre">Sask&quot;</span></code> is <code class="docutils literal"><span class="pre">S</span></code>.  So at
position <code class="docutils literal"><span class="pre">[4]</span></code> we have the letter <code class="docutils literal"><span class="pre">a</span></code>.</p>
<p>If you want the zero-eth letter of a string, you just put 0, or any expression
with the value 0, in the brackets.  Give it a try.</p>
<p>The expression in brackets is called an <strong>index</strong>. An index specifies a member
of an ordered collection.  In this case the collection of characters in the string. The index
<em>indicates</em> which character you want. It can be any integer
expression so long as it evaluates to a valid index value.</p>
<p>Note that indexing returns a <em>string</em> — Python has no special type for a single character.
It is just a string of length 1.</p>
<div class="section" id="check-your-understanding">
<h3>29.3.1. Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_index_test_question_1">
            <p>more-info-about-strings1: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_index_test_question_1_opt_a">t</li><li data-component="feedback" id="string_index_test_question_1_opt_a">Index locations do not start with 1, they start with 0.</li>
            
            <li data-component="answer" data-correct id="string_index_test_question_1_opt_b">h</li><li data-component="feedback" id="string_index_test_question_1_opt_b">Yes, index locations start with 0.</li>
            
            <li data-component="answer"  id="string_index_test_question_1_opt_c">c</li><li data-component="feedback" id="string_index_test_question_1_opt_c">sentence[-3] would return c, counting from right to left.</li>
            
            <li data-component="answer"  id="string_index_test_question_1_opt_d">Error, you cannot use the [ ] operator with a string.</li><li data-component="feedback" id="string_index_test_question_1_opt_d">[ ] is the index operator</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_index_test_question_2">
            <p>more-info-about-strings2: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">+</span> <span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="string_index_test_question_2_opt_a">tr</li><li data-component="feedback" id="string_index_test_question_2_opt_a">Yes, indexing operator has precedence over concatenation.</li>
            
            <li data-component="answer"  id="string_index_test_question_2_opt_b">ps</li><li data-component="feedback" id="string_index_test_question_2_opt_b">p is at location 0, not 2.</li>
            
            <li data-component="answer"  id="string_index_test_question_2_opt_c">nn</li><li data-component="feedback" id="string_index_test_question_2_opt_c">n is at location 5, not 2.</li>
            
            <li data-component="answer"  id="string_index_test_question_2_opt_d">Error, you cannot use the [ ] operator with the + operator.</li><li data-component="feedback" id="string_index_test_question_2_opt_d">[ ] operator returns a string that can be concatenated with another string.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="length">
<h2>29.4. Length<a class="headerlink" href="#length" title="Permalink to this headline">¶</a></h2>
<p>The <code class="docutils literal"><span class="pre">len</span></code> function, when applied to a string, returns the number of characters in a string.</p>

<div data-childcomponent="string_length_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_length_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "Banana"
print(len(fruit))
</textarea>
</div>
<p>To get the last letter of a string, you might be tempted to try something like
this:</p>

<div data-childcomponent="string_length_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_length_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "Banana"
size = len(fruit)
last = fruit[size]       # ERROR!
print(last)
</textarea>
</div>
<p>That won’t work. It causes the runtime error
<code class="docutils literal"><span class="pre">IndexError:</span> <span class="pre">string</span> <span class="pre">index</span> <span class="pre">out</span> <span class="pre">of</span> <span class="pre">range</span></code>. The reason is that there is no
letter at index position 6 in <code class="docutils literal"><span class="pre">&quot;Banana&quot;</span></code>.
Since we started counting at zero, the six indexes are
numbered 0 to 5. To get the last character, we have to subtract 1 from
the length.  Give it a try in the example above.</p>

<div data-childcomponent="string_length_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_length_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "Banana"
size = len(fruit)
last_char = fruit[size-1]
print(last_char)
</textarea>
</div>
<p>Alternatively in Python, we can use <strong>negative indices</strong>, which count backward from the
end of the string. The expression <code class="docutils literal"><span class="pre">fruit[-1]</span></code> yields the last letter,
<code class="docutils literal"><span class="pre">fruit[-2]</span></code> yields the second to last, and so on.  Try it!
Most other languages do <em>not</em> allow the negative indices, but they are a handy feature of Python!</p>
<div class="section" id="id2">
<h3>29.4.1. Check Your Understanding<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_length_test_question_1">
            <p>more-info-about-strings3: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">))</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_length_test_question_1_opt_a">11</li><li data-component="feedback" id="string_length_test_question_1_opt_a">The blank counts as a character.</li>
            
            <li data-component="answer" data-correct id="string_length_test_question_1_opt_b">12</li><li data-component="feedback" id="string_length_test_question_1_opt_b">Yes, there are 12 characters in the string.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_length_test_question_2">
            <p>more-info-about-strings4: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">)</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_length_test_question_2_opt_a">o</li><li data-component="feedback" id="string_length_test_question_2_opt_a">Take a look at the index calculation again, len(sentence)-5.</li>
            
            <li data-component="answer" data-correct id="string_length_test_question_2_opt_b">r</li><li data-component="feedback" id="string_length_test_question_2_opt_b">Yes, len(sentence) is 12 and 12-5 is 7.  Use 7 as index and remember to start counting with 0.</li>
            
            <li data-component="answer"  id="string_length_test_question_2_opt_c">s</li><li data-component="feedback" id="string_length_test_question_2_opt_c">sentence is at index 11</li>
            
            <li data-component="answer"  id="string_length_test_question_2_opt_d">Error, len(sentence) is 12 and there is no index 12.</li><li data-component="feedback" id="string_length_test_question_2_opt_d">You subtract 5 before using the index operator so it will work.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_length_test_question_3">
            <p>more-info-about-strings5: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="string_length_test_question_3_opt_a">c</li><li data-component="feedback" id="string_length_test_question_3_opt_a">Yes, 3 characters from the end.</li>
            
            <li data-component="answer"  id="string_length_test_question_3_opt_b">k</li><li data-component="feedback" id="string_length_test_question_3_opt_b">Count backward 3 characters.</li>
            
            <li data-component="answer"  id="string_length_test_question_3_opt_c">s</li><li data-component="feedback" id="string_length_test_question_3_opt_c">When expressed with a negative index the last character s is at index -1.</li>
            
            <li data-component="answer"  id="string_length_test_question_3_opt_d">Error, negative indices are illegal.</li><li data-component="feedback" id="string_length_test_question_3_opt_d">Python does use negative indices to count backward from the end.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="the-slice-operator">
<h2>29.5. The Slice Operator<a class="headerlink" href="#the-slice-operator" title="Permalink to this headline">¶</a></h2>
<p>A substring of a string is called a <strong>slice</strong>. Selecting a slice is similar to
selecting a character:</p>

<div data-childcomponent="string_slice_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_slice_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
places = "Regina, Estevan, and Melville"
print(places[0:6])
print(places[8:15])
print(places[21:29])
</textarea>
</div>
<p>The <cite>slice</cite> operator <code class="docutils literal"><span class="pre">[n:m]</span></code> returns the part of the string from the n’th character
to the m’th character, including the first but excluding the last. In other words,  start with the character at index n and
go up to but do not include the character at index m.
This
behavior may seem counter-intuitive but if you recall the <code class="docutils literal"><span class="pre">range</span></code> function, it did not include its end
point either.</p>
<p>If you omit the first index (before the colon), the slice starts at the
beginning of the string. If you omit the second index, the slice goes to the
end of the string.</p>
<p>There is no Index Out Of Range exception for a slice. A slice is forgiving and shifts any offending index to something legal.</p>

<div data-childcomponent="string_slice_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_slice_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
fruit = "banana"
print(fruit[:3])
print(fruit[3:])
print(fruit[3:-10])
print(fruit[3:99])
</textarea>
</div>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">What do you think <code class="docutils literal"><span class="pre">fruit[:]</span></code> means?</p>
</div>
<div class="section" id="id3">
<h3>29.5.1. Check Your Understanding<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_length_test_question_1">
            <p>more-info-about-strings6: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">])</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_length_test_question_1_opt_a">python</li><li data-component="feedback" id="string_length_test_question_1_opt_a">That would be sentence[0:6].</li>
            
            <li data-component="answer"  id="string_length_test_question_1_opt_b">rocks</li><li data-component="feedback" id="string_length_test_question_1_opt_b">That would be sentence[7:].</li>
            
            <li data-component="answer" data-correct id="string_length_test_question_1_opt_c">hon r</li><li data-component="feedback" id="string_length_test_question_1_opt_c">Yes, start with the character at index 3 and go up to but not include the character at index 8.</li>
            
            <li data-component="answer"  id="string_length_test_question_1_opt_d">Error, you cannot have two numbers inside the [ ].</li><li data-component="feedback" id="string_length_test_question_1_opt_d">This is called slicing, not indexing.  It requires a start and an end.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_length_test_question_2">
            <p>more-info-about-strings7: What is printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">7</span><span class="p">:</span><span class="mi">11</span><span class="p">]</span> <span class="o">*</span> <span class="mi">3</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer" data-correct id="string_length_test_question_2_opt_a">rockrockrock</li><li data-component="feedback" id="string_length_test_question_2_opt_a">Yes, rock starts at 7 and goes through 10.  Repeat it 3 times.</li>
            
            <li data-component="answer"  id="string_length_test_question_2_opt_b">rock rock rock</li><li data-component="feedback" id="string_length_test_question_2_opt_b">Repetition does not add a space.</li>
            
            <li data-component="answer"  id="string_length_test_question_2_opt_c">rocksrocksrocks</li><li data-component="feedback" id="string_length_test_question_2_opt_c">Slicing will not include the character at index 11.  Just up to it (10 in this case).</li>
            
            <li data-component="answer"  id="string_length_test_question_2_opt_d">Error, you cannot use repetition with slicing.</li><li data-component="feedback" id="string_length_test_question_2_opt_d">The slice will happen first, then the repetition.  So it is ok.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="traversal-and-the-for-loop-by-item">
<h2>29.6. Traversal and the <code class="docutils literal"><span class="pre">for</span></code> Loop: By Item<a class="headerlink" href="#traversal-and-the-for-loop-by-item" title="Permalink to this headline">¶</a></h2>
<p>A lot of computations involve processing a collection one item at a time.  For strings this means
that we would like to process one character at a time.
Often we start at the beginning, select each character in turn, do something
to it, and continue until the end. This pattern of processing is called a
<strong>traversal</strong>.</p>
<p>We have previously seen that the <code class="docutils literal"><span class="pre">for</span></code> statement can iterate over the items of a sequence (a list of names in the case below).</p>

<div data-childcomponent="string_for_loop_by_item_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_for_loop_by_item_1 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
for a_name in ["Joe", "Amy", "Brad", "Angelina", "Zuki", "Thandi", "Paris"]:
    invitation = "Hi " + a_name + ".  Please come to my party on Saturday!"
    print(invitation)
</textarea>
</div>
<p>Recall that the loop variable takes on each value in the sequence of names.  The body is performed once for each name.  The same was true for the sequence of integers created by the <code class="docutils literal"><span class="pre">range</span></code> function.</p>

<div data-childcomponent="string_for_loop_by_item_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_for_loop_by_item_2 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
for a_value in range(10):
    print(a_value)
</textarea>
</div>
<p>Since a string is simply a sequence of characters, the <code class="docutils literal"><span class="pre">for</span></code> loop iterates over each character automatically.</p>

<div data-childcomponent="string_for_loop_by_item_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=string_for_loop_by_item_3 data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
for a_char in "Go Spot Go":
    print(a_char)
</textarea>
</div>
<p>The loop variable <code class="docutils literal"><span class="pre">a_char</span></code> is automatically reassigned each character in the string “Go Spot Go”.
We will refer to this type of sequence iteration as <strong>iteration by item</strong>.
Note that it is only possible to process the characters one at a time from left to right.</p>
<div class="section" id="id4">
<h3>29.6.1. Check Your Understanding<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>

            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_for_loop_by_item_test_question_1">
            <p>more-info-about-strings8: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">:</span>
    <span class="k">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_a">10</li><li data-component="feedback" id="string_for_loop_by_item_test_question_1_opt_a">Iteration by item will process once for each item in the sequence.</li>
            
            <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_b">11</li><li data-component="feedback" id="string_for_loop_by_item_test_question_1_opt_b">The blank is part of the sequence.</li>
            
            <li data-component="answer" data-correct id="string_for_loop_by_item_test_question_1_opt_c">12</li><li data-component="feedback" id="string_for_loop_by_item_test_question_1_opt_c">Yes, there are 12 characters, including the blank.</li>
            
            <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_d">Error, the for statement needs to use the range function.</li><li data-component="feedback" id="string_for_loop_by_item_test_question_1_opt_d">The for statement can iterate over a sequence item by item.</li>
            

            </ul>
            </div>
            
            <div class="runestone">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="string_for_loop_by_item_test_question_2">
            <p>more-info-about-strings9: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">]:</span>
    <span class="k">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

            <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_a">4</li><li data-component="feedback" id="string_for_loop_by_item_test_question_2_opt_a">Slice returns a sequence that can be iterated over.</li>
            
            <li data-component="answer" data-correct id="string_for_loop_by_item_test_question_2_opt_b">5</li><li data-component="feedback" id="string_for_loop_by_item_test_question_2_opt_b">Yes, The blank is part of the sequence returned by slice</li>
            
            <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_c">6</li><li data-component="feedback" id="string_for_loop_by_item_test_question_2_opt_c">Check the result of s[3:8].  It does not include the item at index 8.</li>
            
            <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_d">Error, the for statement cannot use slice.</li><li data-component="feedback" id="string_for_loop_by_item_test_question_2_opt_d">Slice returns a sequence.</li>
            

            </ul>
            </div>
            </div>
</div>
<div class="section" id="practice-problems">
<h2>29.7. Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems. You can either work directly in the textbook, or use Thonny. Either way, copy/paste your finished code into Thonny and save your solution into your Computer Science 20 folder when you finish!</p>
<div class="section" id="con-cat">
<h3>29.7.1. Con Cat<a class="headerlink" href="#con-cat" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in the name of a cat, then prints out a hello message. For example, if the user types in <code class="docutils literal"><span class="pre">Garfield</span></code>, your program could print something like <code class="docutils literal"><span class="pre">Good</span> <span class="pre">to</span> <span class="pre">see</span> <span class="pre">you,</span> <span class="pre">Garfield!</span></code>. <strong>Be sure to use concatenation in your solution!</strong></p>

<div data-childcomponent="practice_problem_string_into_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_1 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Hello message to a cat
</textarea>
</div>
</div>
<div class="section" id="duckling-names">
<h3>29.7.2. Duckling Names<a class="headerlink" href="#duckling-names" title="Permalink to this headline">¶</a></h3>
<p>In Robert McCloskey’s book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>

<div data-childcomponent="practice_problem_string_into_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_2 data-lang="python"    data-timelimit=25000   data-audio=''     >
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    # add a conditional statement here...
    print(letter + suffix)
</textarea>
</div>
<p>Of course, that’s not quite right because Ouack and Quack are misspelled. See if you can fix it without changing the value of either the <code class="docutils literal"><span class="pre">prefixes</span></code> or <code class="docutils literal"><span class="pre">suffix</span></code> variables! <em>You should do this by adding a conditional statement within the for loop.</em></p>
</div>
<div class="section" id="first-letter-of-a-word">
<h3>29.7.3. First Letter of a Word<a class="headerlink" href="#first-letter-of-a-word" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_into_3" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal"><span class="pre">word</span></code> that returns <code class="docutils literal"><span class="pre">True</span></code> if the <code class="docutils literal"><span class="pre">word</span></code> begins with the letter “t” or “c”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;thing&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;concatenation&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">starts_with_tc(&quot;warman&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_into_3 data-lang="python"    data-timelimit=25000   data-audio=''     >
def starts_with_tc(word):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(starts_with_tc("thing"),True,'starts_with_tc("thing")')
        self.assertEqual(starts_with_tc("concatenation"),True,'starts_with_tc("concatenation")')
        self.assertEqual(starts_with_tc("warman"),False,'starts_with_tc("warman")')
        self.assertEqual(starts_with_tc("Thing"),False,'starts_with_tc("Thing")')
        self.assertEqual(starts_with_tc("candy"),True,'starts_with_tc("candy")')
        self.assertEqual(starts_with_tc("saskatoon"),False,'starts_with_tc("saskatoon")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="password-length">
<h3>29.7.4. Password Length<a class="headerlink" href="#password-length" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div data-childcomponent="practice_problem_string_into_4" class="runestone explainer ac_section alert alert-warning">
<p>Create a function with a single parameter <code class="docutils literal"><span class="pre">password</span></code> that returns <code class="docutils literal"><span class="pre">True</span></code> if the <code class="docutils literal"><span class="pre">password</span></code> is between 8 and 32 characters (inclusive), and <code class="docutils literal"><span class="pre">False</span></code> otherwise. <em>Please note there is much more to a strong password than just the length of the string!</em></p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;123456789&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;qwerty&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal"><span class="pre">password_length(&quot;cray-topnotch-tampa-anthem-trial&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>

<textarea data-component="activecode" id=practice_problem_string_into_4 data-lang="python"    data-timelimit=25000   data-audio=''     >
def password_length(password):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(password_length("123456789"),True,'password_length("123456789")')
        self.assertEqual(password_length("qwerty"),False,'password_length("qwerty")')
        self.assertEqual(password_length("cray-topnotch-tampa-anthem-trial"),True,'password_length("cray-topnotch-tampa-anthem-trial")')
        self.assertEqual(password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt"),False,'password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt")')
        self.assertEqual(password_length("fixate"),False,'password_length("fixate")')
        self.assertEqual(password_length("8X?3a739"),True,'password_length("8X?3a739")')

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="removing-the-start-and-end">
<h3>29.7.5. Removing the Start and End<a class="headerlink" href="#removing-the-start-and-end" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in three inputs from the user:</p>
<ul class="simple">
<li>a word/sentence</li>
<li>a starting number</li>
<li>an ending number</li>
</ul>
<p>After taking in the input, the program should print out a portion of the word/sentence that was entered. For example, if the word the user entered was <code class="docutils literal"><span class="pre">Saskatchewan</span></code>, the starting number was 3 and the ending number was 2, the program should print <code class="docutils literal"><span class="pre">katchew</span></code>. Notice that in the output, it is the same as the word, except that the first 3 characters and last 2 characters have been removed.</p>
<p>To be sure you understand the challenge, consider the following example:</p>
<ul class="simple">
<li>word = “unimaginatively”</li>
<li>starting_number = 3</li>
<li>ending_number = 2</li>
<li>output should be <code class="docutils literal"><span class="pre">maginative</span></code></li>
</ul>

<div data-childcomponent="practice_problem_string_into_5" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=practice_problem_string_into_5 data-lang="python"    data-timelimit=25000   data-audio=''     >
# Remove the Start and End
</textarea>
</div>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 28. Image Processing With Conditionals' data-toggle="tooltip">
    <a href="../MoreAboutIteration/ImageProcessingWithSelection.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../MoreAboutIteration/ImageProcessingWithSelection.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 30. More Information About Strings' data-toggle="tooltip" >
    <a href="MoreAboutStrings.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="MoreAboutStrings.html" >Next Section - 30. More Information About Strings</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>. 
    </p>
  </div>
</footer>


  </body>
</html>