<!DOCTYPE html>

<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" /><meta name="generator" content="Docutils 0.17.1: http://docutils.sourceforge.net/" />

    <title>1. String Index, Length, Slicing and Traversal &#8212; Computer Science 20 Saskatchewan</title>
    <link rel="stylesheet" href="../../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../../_static/pygments.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/basic.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/accessibility.css?v=50109A2C" />
    <link rel="stylesheet" type="text/css" href="../../_static/CodeChat.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/matrixeq.css?v=58BC228F" />
    <link rel="stylesheet" type="text/css" href="../../_static/webgldemo.css?v=B36B48D" />
    <link rel="stylesheet" type="text/css" href="../../_static/webglinteractive.css?v=910DF1E1" />
    <link rel="stylesheet" type="text/css" href="../../_static/bootstrap-sphinx.css" />
    <link rel="stylesheet" type="text/css" href="../../_static/637.fafafbd97df8a0d1.css?v=B5361BC8" />
    <link rel="stylesheet" type="text/css" href="../../_static/runestone.e4d5592da655219f.css?v=32648AE5" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/&#34;&#34;/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script src="../../_static/runestonebase.js"></script>
    <script data-url_root="../../" id="documentation_options" src="../../_static/documentation_options.js"></script>
    <script src="../../_static/jquery.js"></script>
    <script src="../../_static/underscore.js"></script>
    <script src="../../_static/doctools.js"></script>
    <script defer="" src="../../_static/animationbase.js?v=A043C3A7"></script>
    <script defer="" src="../../_static/html4css1.css?v=4533ACD3"></script>
    <script defer="" src="../../_static/matrixeq.js?v=D64CFA5A"></script>
    <script defer="" src="../../_static/webglinteractive.js?v=2B768384"></script>
    <script defer="" src="../../_static/FileSaver.min.js?v=8408C997"></script>
    <script defer="" src="../../_static/runtime.c894f7032a5bccea.bundle.js?v=D20A5287"></script>
    <script defer="" src="../../_static/637.d54be67956c5c660.bundle.js?v=915B1670"></script>
    <script defer="" src="../../_static/runestone.5b7c431cb6d3e0db.bundle.js?v=36CB01A8"></script>
    <script src="../../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../../_static/jquery-fix.js"></script>
    <script src="../../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script src="../../_static/bootstrap-sphinx.js"></script>
    <script src="../../_static/waypoints.min.js"></script>
    <script src="../../_static/rangy-core.js"></script>
    <script src="../../_static/rangy-textrange.js"></script>
    <script src="../../_static/rangy-cssclassapplier.js"></script>
    <script src="../../_static/user-highlights.js"></script>
    <script src="../../_static/jquery.idle-timer.js"></script>
    <script src="../../_static/processing-1.4.1.min.js"></script>
    <script src="../../_static/jquery.hotkey.js"></script>
    <script src="../../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../../genindex.html" />
    <link rel="search" title="Search" href="../../search.html" />
    <link rel="next" title="2. String Traversal (by index), and the Accumulator Pattern with Strings" href="MoreAboutStrings.html" />
    <link rel="prev" title="4. Functions and Image Processing" href="../MoreAboutIteration/ImageProcessingFunctions.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'cs20sk';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'cs20sk';
  eBookConfig.runestone_version = '6.2.2';
</script>



  </head><body>


<script>
// Prevent F5 from reloading the page
document.addEventListener('keydown', function(event) {
    if (event.key === 'F5' || (event.ctrlKey && event.key === 'r')) {
        event.preventDefault();
        console.log('Page reload is disabled to prevent loss of work.');
    }
});
</script>

<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../../_static/img/logo_small.png") no-repeat 0px 0px;' href="../../index.html">&nbsp; </a>
        <a class="navbar-brand" href="../../index.html">Computer Science 20 - Saskatchewan</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../../index.html'>Table of Contents</a></li>
            
            <li><a href='/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">ActiveCode Work Area</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">1. String Index, Length, Slicing and Traversal</a><ul>
<li><a class="reference internal" href="#a-collection-data-type">1.1. A Collection Data Type</a></li>
<li><a class="reference internal" href="#concatenation-reminder">1.2. Concatenation Reminder</a></li>
<li><a class="reference internal" href="#f-strings-reminder">1.3. F-Strings Reminder</a></li>
<li><a class="reference internal" href="#index-operator-working-with-the-characters-of-a-string">1.4. Index Operator: Working with the Characters of a String</a><ul>
<li><a class="reference internal" href="#check-your-understanding">1.4.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#length">1.5. Length</a><ul>
<li><a class="reference internal" href="#id1">1.5.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#the-slice-operator">1.6. The Slice Operator</a><ul>
<li><a class="reference internal" href="#id2">1.6.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#traversal-and-the-for-loop-by-item">1.7. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Item</a><ul>
<li><a class="reference internal" href="#id3">1.7.1. Check Your Understanding</a></li>
</ul>
</li>
<li><a class="reference internal" href="#practice-problems">1.8. Practice Problems</a><ul>
<li><a class="reference internal" href="#con-cat">1.8.1. Con Cat</a></li>
<li><a class="reference internal" href="#duckling-names">1.8.2. Duckling Names</a></li>
<li><a class="reference internal" href="#first-letter-of-a-word">1.8.3. First Letter of a Word</a></li>
<li><a class="reference internal" href="#password-length">1.8.4. Password Length</a></li>
<li><a class="reference internal" href="#removing-the-start-and-end">1.8.5. Removing the Start and End</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - <span class="section-number">4. </span>Functions and Image Processing" data-toggle="tooltip">
    <a href="../MoreAboutIteration/ImageProcessingFunctions.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - <span class="section-number">2. </span>String Traversal (by index), and the Accumulator Pattern with Strings' data-toggle="tooltip" >
    <a href="MoreAboutStrings.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <section id="string-index-length-slicing-and-traversal">
<h1><span class="section-number">1. </span>String Index, Length, Slicing and Traversal<a class="headerlink" href="#string-index-length-slicing-and-traversal" title="Permalink to this headline">¶</a></h1>
<div class="topic">
<p class="topic-title">Quick Overview of Day</p>
<p>Explore string operators, index values, length of strings, string slicing, and traversing a string using a for loop (by item).</p>
</div>

    <div data-component="reveal" class="runestone " id="curriculum_addressed_string_exploration"   data-showtitle="Curriculum Outcomes Addressed In This Section" data-hidetitle="Hide"  style="visibility: hidden;">
    <ul class="simple">
<li><p><strong>CS20-CP1</strong> Apply various problem-solving strategies to solve programming problems throughout Computer Science 20.</p></li>
<li><p><strong>CS20-FP1</strong> Utilize different data types, including integer, floating point, Boolean and string, to solve programming problems.</p></li>
<li><p><strong>CS20-FP2</strong> Investigate how control structures affect program flow.</p></li>
<li><p><strong>CS20-FP3</strong> Construct and utilize functions to create reusable pieces of code.</p></li>
</ul>

    </div>
    <p>Throughout the portion of the course, we have used strings to represent words or phrases that we
wanted to print out.  Our definition was simple:  a string is simply some characters inside quotes.  We will now explore strings in much more detail.</p>
<section id="a-collection-data-type">
<h2><span class="section-number">1.1. </span>A Collection Data Type<a class="headerlink" href="#a-collection-data-type" title="Permalink to this headline">¶</a></h2>
<p>So far we have seen built-in types like: <code class="docutils literal notranslate"><span class="pre">int</span></code>, <code class="docutils literal notranslate"><span class="pre">float</span></code>,
<code class="docutils literal notranslate"><span class="pre">bool</span></code>, <code class="docutils literal notranslate"><span class="pre">str</span></code> and we’ve seen lists.
<code class="docutils literal notranslate"><span class="pre">int</span></code>, <code class="docutils literal notranslate"><span class="pre">float</span></code>, and
<code class="docutils literal notranslate"><span class="pre">bool</span></code> are considered to be simple or primitive data types because their values are not composed
of any smaller parts.  They cannot be broken down.
On the other hand, strings and lists are different from the others because they
are made up of smaller pieces.  In the case of strings, they are made up of smaller
strings each containing one <strong>character</strong>.</p>
<p>Types that are comprised of smaller pieces are called <strong>collection data types</strong>.
Depending on what we are doing, we may want to treat a collection data type as a
single entity (the whole), or we may want to access its parts. This ambiguity is useful.</p>
<p>Strings can be defined as sequential collections of characters.  This means that the individual characters
that make up the string are assumed to be in a particular order from left to right.</p>
<p>A string that contains no characters, often referred to as the <strong>empty string</strong>, is still considered to be a string.  It is simply a sequence of zero characters and is represented by ‘’ or “” (two single or two double quotes with nothing in between).</p>
</section>
<section id="concatenation-reminder">
<h2><span class="section-number">1.2. </span>Concatenation Reminder<a class="headerlink" href="#concatenation-reminder" title="Permalink to this headline">¶</a></h2>
<p>As we have seen before, you cannot perform mathematical operations on strings, even if the
strings look like numbers. The one exception to this rule is that the <code class="docutils literal notranslate"><span class="pre">+</span></code> operator does work with strings, but for strings, the <code class="docutils literal notranslate"><span class="pre">+</span></code> operator represents <strong>concatenation</strong>, not addition.  As we have learned previously, concatenation means joining the two operands by linking them end-to-end. For example:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=concatenation_reminder_demo data-question_label="1.2.1">
<div id=concatenation_reminder_demo_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="concatenation_reminder_demo_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "banana"
baked_good = " nut bread"
print(fruit + baked_good)
</textarea>
</div>
</div>
</section>
<section id="f-strings-reminder">
<h2><span class="section-number">1.3. </span>F-Strings Reminder<a class="headerlink" href="#f-strings-reminder" title="Permalink to this headline">¶</a></h2>
<p>Although concatenation works well for simple use cases, it is often preferable to use an f-string to
format strings, which allows you to insert the value of any variable by enclosing the variable in curly
brackets. To indicate it is a formatted string, add a <code class="docutils literal notranslate"><span class="pre">f</span></code> before the string quotes begin.
For example, the concatenation example above could also be written as:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=f_strings_reminder_demo data-question_label="1.3.1">
<div id=f_strings_reminder_demo_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="f_strings_reminder_demo_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "banana"
baked_good = "nut bread"
print(f"{fruit} {baked_good}")
</textarea>
</div>
</div>
</section>
<section id="index-operator-working-with-the-characters-of-a-string">
<span id="index-0"></span><h2><span class="section-number">1.4. </span>Index Operator: Working with the Characters of a String<a class="headerlink" href="#index-operator-working-with-the-characters-of-a-string" title="Permalink to this headline">¶</a></h2>
<p>The <strong>indexing operator</strong> (Python uses square brackets to enclose the index)
selects a single character from a string.  The characters are accessed by their position or
index value.  For example, in the string shown below, the 14 characters are indexed left to right from postion 0 to position 13.</p>
<img alt="index values" src="../../_images/indexvalues.png" />
<p>It is also the case that the positions are named from right to left using negative numbers where -1 is the rightmost
index and so on.
Note that the character at index 9 (or -5) is the blank/space character.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_index_1 data-question_label="1.4.1">
<div id=string_index_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_index_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
place = "Saskatoon Sask"
some_char = place[4]
print(some_char)

last_char = place[-1]
print(last_char)
</textarea>
</div>
</div>
<p>The expression <code class="docutils literal notranslate"><span class="pre">place[4]</span></code> selects the character at index 4 from <code class="docutils literal notranslate"><span class="pre">place</span></code>, and creates a new
string containing just this one character. The variable <code class="docutils literal notranslate"><span class="pre">some_char</span></code> refers to the result.</p>
<p>Remember that computer scientists often start counting
from zero. The letter at index zero of <code class="docutils literal notranslate"><span class="pre">&quot;Saskatoon</span> <span class="pre">Sask&quot;</span></code> is <code class="docutils literal notranslate"><span class="pre">S</span></code>.  So at
position <code class="docutils literal notranslate"><span class="pre">[4]</span></code> we have the letter <code class="docutils literal notranslate"><span class="pre">a</span></code>.</p>
<p>If you want the zero-eth letter of a string, you just put 0, or any expression
with the value 0, in the brackets.  Give it a try.</p>
<p>The expression in brackets is called an <strong>index</strong>. An index specifies a member
of an ordered collection.  In this case the collection of characters in the string. The index
<em>indicates</em> which character you want. It can be any integer
expression so long as it evaluates to a valid index value.</p>
<p>Note that indexing returns a <em>string</em> — Python has no special type for a single character.
It is just a string of length 1.</p>
<section id="check-your-understanding">
<h3><span class="section-number">1.4.1. </span>Check Your Understanding<a class="headerlink" href="#check-your-understanding" title="Permalink to this headline">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.4.1.1" data-multipleanswers="false"  id="string_index_test_question_1"  style="visibility: hidden;">
    <p>string-exploration4: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_index_test_question_1_opt_a">t</li><li data-component="feedback">Index locations do not start with 1, they start with 0.</li>
    
    <li data-component="answer" data-correct='yes' id="string_index_test_question_1_opt_b">h</li><li data-component="feedback">Yes, index locations start with 0.</li>
    
    <li data-component="answer"  id="string_index_test_question_1_opt_c">c</li><li data-component="feedback">sentence[-3] would return c, counting from right to left.</li>
    
    <li data-component="answer"  id="string_index_test_question_1_opt_d">Error, you cannot use the [ ] operator with a string.</li><li data-component="feedback">[ ] is the index operator</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.4.1.2" data-multipleanswers="false"  id="string_index_test_question_2"  style="visibility: hidden;">
    <p>string-exploration5: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">2</span><span class="p">]</span> <span class="o">+</span> <span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="string_index_test_question_2_opt_a">tr</li><li data-component="feedback">Yes, indexing operator has precedence over concatenation.</li>
    
    <li data-component="answer"  id="string_index_test_question_2_opt_b">ps</li><li data-component="feedback">p is at location 0, not 2.</li>
    
    <li data-component="answer"  id="string_index_test_question_2_opt_c">nn</li><li data-component="feedback">n is at location 5, not 2.</li>
    
    <li data-component="answer"  id="string_index_test_question_2_opt_d">Error, you cannot use the [ ] operator with the + operator.</li><li data-component="feedback">[ ] operator returns a string that can be concatenated with another string.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="length">
<span id="index-1"></span><h2><span class="section-number">1.5. </span>Length<a class="headerlink" href="#length" title="Permalink to this headline">¶</a></h2>
<p>The <code class="docutils literal notranslate"><span class="pre">len</span></code> function, when applied to a string, returns the number of characters in a string.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_length_1 data-question_label="1.5.1">
<div id=string_length_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_length_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "Banana"
print(len(fruit))
</textarea>
</div>
</div>
<p>To get the last letter of a string, you might be tempted to try something like
this:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_length_2 data-question_label="1.5.2">
<div id=string_length_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_length_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "Banana"
size = len(fruit)
last = fruit[size]       # ERROR!
print(last)
</textarea>
</div>
</div>
<p>That won’t work. It causes the runtime error
<code class="docutils literal notranslate"><span class="pre">IndexError:</span> <span class="pre">string</span> <span class="pre">index</span> <span class="pre">out</span> <span class="pre">of</span> <span class="pre">range</span></code>. The reason is that there is no
letter at index position 6 in <code class="docutils literal notranslate"><span class="pre">&quot;Banana&quot;</span></code>.
Since we started counting at zero, the six indexes are
numbered 0 to 5. To get the last character, we have to subtract 1 from
the length.  Give it a try in the example above.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_length_3 data-question_label="1.5.3">
<div id=string_length_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_length_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "Banana"
size = len(fruit)
last_char = fruit[size-1]
print(last_char)
</textarea>
</div>
</div>
<p>Alternatively in Python, we can use <strong>negative indices</strong>, which count backward from the
end of the string. The expression <code class="docutils literal notranslate"><span class="pre">fruit[-1]</span></code> yields the last letter,
<code class="docutils literal notranslate"><span class="pre">fruit[-2]</span></code> yields the second to last, and so on.  Try it!
Most other languages do <em>not</em> allow the negative indices, but they are a handy feature of Python!</p>
<section id="id1">
<h3><span class="section-number">1.5.1. </span>Check Your Understanding<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.5.1.1" data-multipleanswers="false"  id="string_length_test_question_1"  style="visibility: hidden;">
    <p>string-exploration9: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">))</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_length_test_question_1_opt_a">11</li><li data-component="feedback">The blank counts as a character.</li>
    
    <li data-component="answer" data-correct='yes' id="string_length_test_question_1_opt_b">12</li><li data-component="feedback">Yes, there are 12 characters in the string.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.5.1.2" data-multipleanswers="false"  id="string_length_test_question_2"  style="visibility: hidden;">
    <p>string-exploration10: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="nb">len</span><span class="p">(</span><span class="n">sentence</span><span class="p">)</span><span class="o">-</span><span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_length_test_question_2_opt_a">o</li><li data-component="feedback">Take a look at the index calculation again, len(sentence)-5.</li>
    
    <li data-component="answer" data-correct='yes' id="string_length_test_question_2_opt_b">r</li><li data-component="feedback">Yes, len(sentence) is 12 and 12-5 is 7.  Use 7 as index and remember to start counting with 0.</li>
    
    <li data-component="answer"  id="string_length_test_question_2_opt_c">s</li><li data-component="feedback">sentence is at index 11</li>
    
    <li data-component="answer"  id="string_length_test_question_2_opt_d">Error, len(sentence) is 12 and there is no index 12.</li><li data-component="feedback">You subtract 5 before using the index operator so it will work.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.5.1.3" data-multipleanswers="false"  id="string_length_test_question_3"  style="visibility: hidden;">
    <p>string-exploration11: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="o">-</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="string_length_test_question_3_opt_a">c</li><li data-component="feedback">Yes, 3 characters from the end.</li>
    
    <li data-component="answer"  id="string_length_test_question_3_opt_b">k</li><li data-component="feedback">Count backward 3 characters.</li>
    
    <li data-component="answer"  id="string_length_test_question_3_opt_c">s</li><li data-component="feedback">When expressed with a negative index the last character s is at index -1.</li>
    
    <li data-component="answer"  id="string_length_test_question_3_opt_d">Error, negative indices are illegal.</li><li data-component="feedback">Python does use negative indices to count backward from the end.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="the-slice-operator">
<span id="index-2"></span><h2><span class="section-number">1.6. </span>The Slice Operator<a class="headerlink" href="#the-slice-operator" title="Permalink to this headline">¶</a></h2>
<p>A substring of a string is called a <strong>slice</strong>. Selecting a slice is similar to
selecting a character:</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_slice_1 data-question_label="1.6.1">
<div id=string_slice_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_slice_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
places = "Regina, Estevan, and Melville"
print(places[0:6])
print(places[8:15])
print(places[21:29])
</textarea>
</div>
</div>
<p>The <cite>slice</cite> operator <code class="docutils literal notranslate"><span class="pre">[n:m]</span></code> returns the part of the string from the n’th character
to the m’th character, including the first but excluding the last. In other words,  start with the character at index n and
go up to but do not include the character at index m.
This
behavior may seem counter-intuitive but if you recall the <code class="docutils literal notranslate"><span class="pre">range</span></code> function, it did not include its end
point either.</p>
<p>If you omit the first index (before the colon), the slice starts at the
beginning of the string. If you omit the second index, the slice goes to the
end of the string.</p>
<p>There is no Index Out Of Range exception for a slice. A slice is forgiving and shifts any offending index to something legal.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_slice_2 data-question_label="1.6.2">
<div id=string_slice_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_slice_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
fruit = "banana"
print(fruit[:3])
print(fruit[3:])
print(fruit[3:-10])
print(fruit[3:99])
</textarea>
</div>
</div>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>What do you think <code class="docutils literal notranslate"><span class="pre">fruit[:]</span></code> means?</p>
</div>
<section id="id2">
<h3><span class="section-number">1.6.1. </span>Check Your Understanding<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.6.1.1" data-multipleanswers="false"  id="string_slice_test_question_1"  style="visibility: hidden;">
    <p>string-exploration14: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">])</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_slice_test_question_1_opt_a">python</li><li data-component="feedback">That would be sentence[0:6].</li>
    
    <li data-component="answer"  id="string_slice_test_question_1_opt_b">rocks</li><li data-component="feedback">That would be sentence[7:].</li>
    
    <li data-component="answer" data-correct='yes' id="string_slice_test_question_1_opt_c">hon r</li><li data-component="feedback">Yes, start with the character at index 3 and go up to but not include the character at index 8.</li>
    
    <li data-component="answer"  id="string_slice_test_question_1_opt_d">Error, you cannot have two numbers inside the [ ].</li><li data-component="feedback">This is called slicing, not indexing.  It requires a start and an end.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.6.1.2" data-multipleanswers="false"  id="string_slice_test_question_2"  style="visibility: hidden;">
    <p>string-exploration15: What is printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">sentence</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="nb">print</span><span class="p">(</span><span class="n">sentence</span><span class="p">[</span><span class="mi">7</span><span class="p">:</span><span class="mi">11</span><span class="p">]</span> <span class="o">*</span> <span class="mi">3</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer" data-correct='yes' id="string_slice_test_question_2_opt_a">rockrockrock</li><li data-component="feedback">Yes, rock starts at 7 and goes through 10.  Repeat it 3 times.</li>
    
    <li data-component="answer"  id="string_slice_test_question_2_opt_b">rock rock rock</li><li data-component="feedback">Repetition does not add a space.</li>
    
    <li data-component="answer"  id="string_slice_test_question_2_opt_c">rocksrocksrocks</li><li data-component="feedback">Slicing will not include the character at index 11.  Just up to it (10 in this case).</li>
    
    <li data-component="answer"  id="string_slice_test_question_2_opt_d">Error, you cannot use repetition with slicing.</li><li data-component="feedback">The slice will happen first, then the repetition.  So it is ok.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="traversal-and-the-for-loop-by-item">
<h2><span class="section-number">1.7. </span>Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Item<a class="headerlink" href="#traversal-and-the-for-loop-by-item" title="Permalink to this headline">¶</a></h2>
<p>A lot of computations involve processing a collection one item at a time.  For strings this means
that we would like to process one character at a time.
Often we start at the beginning, select each character in turn, do something
to it, and continue until the end. This pattern of processing is called a
<strong>traversal</strong>.</p>
<p>We have previously seen that the <code class="docutils literal notranslate"><span class="pre">for</span></code> statement can iterate over the items of a sequence (a list of names in the case below).</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_for_loop_by_item_1 data-question_label="1.7.1">
<div id=string_for_loop_by_item_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_for_loop_by_item_1_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
for a_name in ["Joe", "Amy", "Brad", "Angelina", "Zuki", "Thandi", "Paris"]:
    invitation = f"Hi {a_name}.  Please come to my party on Saturday!"
    print(invitation)
</textarea>
</div>
</div>
<p>Recall that the loop variable takes on each value in the sequence of names.  The body is performed once for each name.  The same was true for the sequence of integers created by the <code class="docutils literal notranslate"><span class="pre">range</span></code> function.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_for_loop_by_item_2 data-question_label="1.7.2">
<div id=string_for_loop_by_item_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_for_loop_by_item_2_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
for a_value in range(10):
    print(a_value)
</textarea>
</div>
</div>
<p>Since a string is simply a sequence of characters, the <code class="docutils literal notranslate"><span class="pre">for</span></code> loop iterates over each character automatically.</p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=string_for_loop_by_item_3 data-question_label="1.7.3">
<div id=string_for_loop_by_item_3_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="string_for_loop_by_item_3_editor" 
      data-timelimit=25000  data-codelens="true"   
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
for a_char in "Go Spot Go":
    print(a_char)
</textarea>
</div>
</div>
<p>The loop variable <code class="docutils literal notranslate"><span class="pre">a_char</span></code> is automatically reassigned each character in the string “Go Spot Go”.
We will refer to this type of sequence iteration as <strong>iteration by item</strong>.
Note that it is only possible to process the characters one at a time from left to right.</p>
<section id="id3">
<h3><span class="section-number">1.7.1. </span>Check Your Understanding<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>

    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.7.1.1" data-multipleanswers="false"  id="string_for_loop_by_item_test_question_1"  style="visibility: hidden;">
    <p>string-exploration19: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_a">10</li><li data-component="feedback">Iteration by item will process once for each item in the sequence.</li>
    
    <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_b">11</li><li data-component="feedback">The blank is part of the sequence.</li>
    
    <li data-component="answer" data-correct='yes' id="string_for_loop_by_item_test_question_1_opt_c">12</li><li data-component="feedback">Yes, there are 12 characters, including the blank.</li>
    
    <li data-component="answer"  id="string_for_loop_by_item_test_question_1_opt_d">Error, the for statement needs to use the range function.</li><li data-component="feedback">The for statement can iterate over a sequence item by item.</li>
    

    </ul>
    </div>
    
    <div class="runestone ">
    <ul data-component="multiplechoice" data-question_label="1.7.1.2" data-multipleanswers="false"  id="string_for_loop_by_item_test_question_2"  style="visibility: hidden;">
    <p>string-exploration20: How many times is the word HELLO printed by the following statements?</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;python rocks&quot;</span>
<span class="k">for</span> <span class="n">ch</span> <span class="ow">in</span> <span class="n">s</span><span class="p">[</span><span class="mi">3</span><span class="p">:</span><span class="mi">8</span><span class="p">]:</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;HELLO&quot;</span><span class="p">)</span>
</pre></div>
</div>

    <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_a">4</li><li data-component="feedback">Slice returns a sequence that can be iterated over.</li>
    
    <li data-component="answer" data-correct='yes' id="string_for_loop_by_item_test_question_2_opt_b">5</li><li data-component="feedback">Yes, The blank is part of the sequence returned by slice</li>
    
    <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_c">6</li><li data-component="feedback">Check the result of s[3:8].  It does not include the item at index 8.</li>
    
    <li data-component="answer"  id="string_for_loop_by_item_test_question_2_opt_d">Error, the for statement cannot use slice.</li><li data-component="feedback">Slice returns a sequence.</li>
    

    </ul>
    </div>
    </section>
</section>
<section id="practice-problems">
<h2><span class="section-number">1.8. </span>Practice Problems<a class="headerlink" href="#practice-problems" title="Permalink to this headline">¶</a></h2>
<p>Try the following practice problems. You can either work directly in the textbook, or use Thonny. Either way, be sure to save your solution into your Computer Science 20 folder when you finish!</p>
<section id="con-cat">
<h3><span class="section-number">1.8.1. </span>Con Cat<a class="headerlink" href="#con-cat" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in the name of a cat, then prints out a hello message. For example, if the user types in <code class="docutils literal notranslate"><span class="pre">Garfield</span></code>, your program could print something like <code class="docutils literal notranslate"><span class="pre">Good</span> <span class="pre">to</span> <span class="pre">see</span> <span class="pre">you,</span> <span class="pre">Garfield!</span></code>. <strong>You can use concatenation or f-strings in your solution!</strong></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=practice_problem_string_into_1 data-question_label="1.8.1.1">
<div id=practice_problem_string_into_1_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="practice_problem_string_into_1_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Hello message to a cat
</textarea>
</div>
</div>
</section>
<section id="duckling-names">
<h3><span class="section-number">1.8.2. </span>Duckling Names<a class="headerlink" href="#duckling-names" title="Permalink to this headline">¶</a></h3>
<p>In Robert McCloskey’s book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.  This code below attempts to output these names in order. Unfortunately, the output is not quite right because Ouack and Quack are misspelled. See if you can fix it without changing the value of either the <code class="docutils literal notranslate"><span class="pre">prefixes</span></code> or <code class="docutils literal notranslate"><span class="pre">suffix</span></code> variables! <em>You should do this by adding a conditional statement within the for loop.</em></p>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=practice_problem_string_into_2 data-question_label="1.8.2.1">
<div id=practice_problem_string_into_2_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="practice_problem_string_into_2_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
prefixes = "JKLMNOPQ"
suffix = "ack"

for letter in prefixes:
    # add a conditional statement here...
    print(letter + suffix)
</textarea>
</div>
</div>
</section>
<section id="first-letter-of-a-word">
<h3><span class="section-number">1.8.3. </span>First Letter of a Word<a class="headerlink" href="#first-letter-of-a-word" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=practice_problem_string_into_3 data-question_label="1.8.3.1">
<div id=practice_problem_string_into_3_question class="ac_question col-md-12">
<p>Create a function with a single parameter <code class="docutils literal notranslate"><span class="pre">word</span></code> that returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">word</span></code> begins with the letter “t” or “c”.</p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">starts_with_tc(&quot;thing&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">starts_with_tc(&quot;concatenation&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">starts_with_tc(&quot;warman&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>

</div>
<textarea data-lang="python" id="practice_problem_string_into_3_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def starts_with_tc(word):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(starts_with_tc("thing"),True,'starts_with_tc("thing")')
        self.assertEqual(starts_with_tc("concatenation"),True,'starts_with_tc("concatenation")')
        self.assertEqual(starts_with_tc("warman"),False,'starts_with_tc("warman")')
        self.assertEqual(starts_with_tc("Thing"),False,'starts_with_tc("Thing")')
        self.assertEqual(starts_with_tc("candy"),True,'starts_with_tc("candy")')
        self.assertEqual(starts_with_tc("saskatoon"),False,'starts_with_tc("saskatoon")')

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="password-length">
<h3><span class="section-number">1.8.4. </span>Password Length<a class="headerlink" href="#password-length" title="Permalink to this headline">¶</a></h3>
<div class="admonition note">
<p class="admonition-title">Note</p>
<p>The only thing you need to do for this question is to complete the function definition! <strong>You do not need to call the function</strong>, as that will be done automatically for you.</p>
</div>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=practice_problem_string_into_4 data-question_label="1.8.4.1">
<div id=practice_problem_string_into_4_question class="ac_question col-md-12">
<p>Create a function with a single parameter <code class="docutils literal notranslate"><span class="pre">password</span></code> that returns <code class="docutils literal notranslate"><span class="pre">True</span></code> if the <code class="docutils literal notranslate"><span class="pre">password</span></code> is between 8 and 32 characters (inclusive), and <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise. <em>Please note there is much more to a strong password than just the length of the string!</em></p>
<p><strong>Examples:</strong></p>
<p><code class="docutils literal notranslate"><span class="pre">password_length(&quot;123456789&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">password_length(&quot;qwerty&quot;)</span> <span class="pre">→</span> <span class="pre">False</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">password_length(&quot;cray-topnotch-tampa-anthem-trial&quot;)</span> <span class="pre">→</span> <span class="pre">True</span></code></p>

</div>
<textarea data-lang="python" id="practice_problem_string_into_4_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
def password_length(password):
    return False

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(password_length("123456789"),True,'password_length("123456789")')
        self.assertEqual(password_length("qwerty"),False,'password_length("qwerty")')
        self.assertEqual(password_length("cray-topnotch-tampa-anthem-trial"),True,'password_length("cray-topnotch-tampa-anthem-trial")')
        self.assertEqual(password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt"),False,'password_length("7mvRVNwB9JAeVtdWywB3MfGsU6+Rnz4PyVDNCnUt")')
        self.assertEqual(password_length("fixate"),False,'password_length("fixate")')
        self.assertEqual(password_length("8X?3a739"),True,'password_length("8X?3a739")')

myTests().main()
</textarea>
</div>
</div>
</section>
<section id="removing-the-start-and-end">
<h3><span class="section-number">1.8.5. </span>Removing the Start and End<a class="headerlink" href="#removing-the-start-and-end" title="Permalink to this headline">¶</a></h3>
<p>Create a program that takes in three inputs from the user:</p>
<ul class="simple">
<li><p>a word/sentence</p></li>
<li><p>a starting number</p></li>
<li><p>an ending number</p></li>
</ul>
<p>After taking in the input, the program should print out a portion of the word/sentence that was entered. For example, if the word the user entered was <code class="docutils literal notranslate"><span class="pre">Saskatchewan</span></code>, the starting number was 3 and the ending number was 2, the program should print <code class="docutils literal notranslate"><span class="pre">katchew</span></code>. Notice that in the output, it is the same as the word, except that the first 3 characters and last 2 characters have been removed.</p>
<p>To be sure you understand the challenge, consider the following example:</p>
<ul class="simple">
<li><p>word = “unimaginatively”</p></li>
<li><p>starting_number = 3</p></li>
<li><p>ending_number = 2</p></li>
<li><p>output should be <code class="docutils literal notranslate"><span class="pre">maginative</span></code></p></li>
</ul>

<div class="runestone explainer ac_section ">
<div data-component="activecode" id=practice_problem_string_into_5 data-question_label="1.8.5.1">
<div id=practice_problem_string_into_5_question class="ac_question col-md-12">

</div>
<textarea data-lang="python" id="practice_problem_string_into_5_editor" 
      data-timelimit=25000   data-enabledownload="true"  
    data-audio=''      
           data-wasm=/_static
     style="visibility: hidden;">
# Remove the Start and End
</textarea>
</div>
</div>
</section>
</section>
</section>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - <span class="section-number">4. </span>Functions and Image Processing' data-toggle="tooltip">
    <a href="../MoreAboutIteration/ImageProcessingFunctions.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../MoreAboutIteration/ImageProcessingFunctions.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - <span class="section-number">2. </span>String Traversal (by index), and the Accumulator Pattern with Strings' data-toggle="tooltip" >
    <a href="MoreAboutStrings.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="MoreAboutStrings.html" >Next Section - <span class="section-number">2. </span>String Traversal (by index), and the Accumulator Pattern with Strings</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
    Created by Dan Schellenberg. <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons Licence" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/80x15.png" /></a> This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.
    </p>
  </div>
</footer>


  </body>
</html>